<quiz>
    <question type="stack">
        <name>
            <text>1.5.2.19 Umstellen und Substitution mit Einheiten [Transposition and Substitution with Units]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        In the design of orifice plate flowmeters, the volumetric flowrate, \(Q\) \((\textrm{m}^{2}\textrm{s}^{-1})\), is given by \[Q = CO\sqrt{\dfrac{2g\Delta h}{1 - O^{2}/P^{2}}}\] where \(C\) is a dimensionless discharge coefficient, \(\Delta h \textrm{ (m)}\)
                is the head difference across the orifice plate and \(O\) \((\textrm{m}^{2})\) is the area of the orifice and \(P\) \((\textrm{m}^{2})\) is the area of the pipe.
    </span>
    <span class="multilang" lang="de">
        Bei der Konstruktion von Messblenden-Durchflussmessern ist der Volumendurchfluss, \(Q\) \((\textrm{m}^{2}\textrm{s}^{- 1})\), gegeben durch \[Q = CO\sqrt{\dfrac{2g\Delta h}{1 - O^{2}/P^{2}}}\] wobei \(C\) ein dimensionsloser Abflussbeiwert, \(\Delta h \textrm{ (m)}\)
                die Druckdifferenz über der Blende und \(O\) \((\textrm{m}^{2})\) die Fläche der Blende und \(P\) \((\textrm{m}^ {2})\) die Fläche des Rohrs ist.
    </span>
</p>
<br/>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Rearrange the equation to solve for the area of the orifice, \(O\), in terms of the other variables
        </span>
        <span class="multilang" lang="de">
            Stellen Sie die Gleichung um, um sie für die Öffnungsfläche \(O\) in Bezug auf die anderen Variablen zu lösen.
        </span>
        <br/>
        \(O = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                A volumetric flowrate of {@aq*10^6@} \(\textrm{cm}^{3}\textrm{s}^{-1}\) passes through a {@diam*100@} \(\textrm{cm}\) inside diameter pipe. Assuming a discharge coefficient of {@float(ac)@}, calculate the required orifice diameter, so that the
                                        head difference across the orifice plate is {@ah*1000@} \(\textrm{mm}\).
            </span>
            <span class="multilang" lang="de">
                Ein Volumenstrom von {@aq*10^6@} \(\textrm{cm}^{3}\textrm{s}^{-1}\) verläuft durch ein Rohr mit Innendurchmesser {@diam*100@} \(\textrm{cm}\). Berechnen Sie unter Annahme eines Durchflusskoeffizienten von {@float(ac)@} den erforderlichen Öffnungsdurchmesser, sodass die Druckdifferenz über der Blende {@ah*1000@} \(\textrm{mm}\) beträgt.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                (Give your answer as a decimal to three decimal places, note that trailing zeroes are required i.e. \(2.3\) must be written as \(2.300\).)
            </span>
            <span class="multilang" lang="de">
                (Geben Sie Ihr Ergebnis als Dezimalzahl mit drei Dezimalstellen an. Beachten Sie, dass nachgestellte Nullen erforderlich sind, d. h. \(2.3\) muss als \(2.300\) geschrieben werden.)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Orifice diameter =
            </span>
            <span class="multilang" lang="de">
                Öffnungsdurchmesser =
            </span>
            [[input:ans2]] [[validation:ans2]]
        </p>
        [[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            We have the equation \[Q = CO\sqrt{\dfrac{2g\Delta h}{1-O^{2}/P^{2}}}\] First divide by \(CO\) \[\implies \dfrac{Q}{CO} = \sqrt{\dfrac{2g\Delta h}{1-O^{2}/P^{2}}}\] Square both sides \[\implies \dfrac{Q^{2}}{C^{2}O^{2}} = \dfrac{2g\Delta h}{1-O^{2}/P^{2}}\]
                                Multiply through by \((1-O^{2}/P^{2})\), expand out the brackets and simplify \[\implies \dfrac{Q^{2}}{C^{2}O^{2}}(1-O^{2}/P^{2}) = 2g\Delta h \implies \dfrac{Q^{2}}{C^{2}O^{2}}-\dfrac{Q^{2}}{C^{2}O^{2}}(\dfrac{O^{2}}{P^{2}}) = 2g\Delta h \implies
                                \dfrac{Q^{2}}{C^{2}O^{2}}-\dfrac{Q^{2}}{C^{2}P^{2}} = 2g\Delta h\] Multiply through by \(C^2P^2\) \[\implies \dfrac{Q^{2}P^{2}}{O^2}-Q^2 = 2g\Delta hC^{2}P^{2}\] Add \(Q^2\) to both sides \[\implies \dfrac{Q^{2}P^{2}}{O^2} = Q^2+2g\Delta hC^{2}P^{2}\]
                                Multiply by \(O^2\) \[\implies Q^{2}P^{2} = O^2(Q^2+2g\Delta hC^{2}P^{2})\] Divide by \((Q^2+2g\Delta hC^{2}P^{2})\) \[\implies O^2 = \dfrac{Q^{2}P^{2}}{Q^2+2g\Delta hC^{2}P^{2}}\] We then take the square root of both sides of the equation \[\implies
                                \sqrt{O^2} = \sqrt{\dfrac{Q^{2}P^{2}}{Q^2+2g\Delta hC^{2}P^{2}}}\] Since Q, O and A take positive values this can be simplified to transpose the equation in terms of \(O\) \[\implies O = \dfrac{QP}{\sqrt{Q^2+2g\Delta hC^{2}P^{2}}}\]
        </span>
        <span class="multilang" lang="de">
            Wir haben die Gleichung \[Q = CO\sqrt{\dfrac{2g\Delta h}{1-O^{2}/P^{2}}}\] Erst dividieren durch \(CO\) \[\implies \dfrac{Q}{CO} = \sqrt{\dfrac{2g\Delta h}{1-O^{2}/P^{2}}}\] Quadrieren auf beiden Seiten \[\implies \dfrac{Q^{2}}{C^{2}O^{2}} = \dfrac{2g\Delta h}{1-O^{2}/P^{2}}\]
                                Multiplizieren mit \((1-O^{2}/P^{2})\), die Klammern erweitern und vereinfachen \[\implies \dfrac{Q^{2}}{C^{2}O^ {2}}(1-O^{2}/P^{2}) = 2g\Delta h \implies \dfrac{Q^{2}}{C^{2}O^{2}}-\dfrac {Q^{2}}{C^{2}O^{2}}(\dfrac{O^{2}}{P^{2}}) = 2g\Delta h \implies
                                \dfrac{Q^{2}}{C^{2}O^{2}}-\dfrac{Q^{2}}{C^{2}P^{2}} = 2g\Delta h\] Multiplizieren mit \(C^2P^2\) \[\implies \dfrac{Q^{2}P^{2}}{O^2}-Q^2 = 2g\Delta hC^{2}P^{2}\] \(Q^2\) auf beiden Seiten addieren \[\implies \dfrac{Q^{2}P^{2}}{O^2} = Q^2+2g\Delta hC^{2}P^{2}\]
                                Multiplizieren mit \(O^2\) \[\implies Q^{2}P^{2} = O^2(Q^2+2g\Delta hC^{2}P^{2})\] Dividieren durch \((Q^2+2g\Delta hC^{2}P^{2})\) \[\implies O^2 = \dfrac{Q^{2}P^{2}}{Q^2+ 2g\Delta hC^{2}P^{2}}\] Wir ziehen dann die Quadratwurzel von beiden Seiten der Gleichung \[\implies
                                \sqrt{O^2} = \sqrt{\dfrac{Q^{2}P^{2}}{Q^2+2g\Delta hC^{2}P^{2}}}\] Da Q, O und A positive Werte annehmen, kann dies vereinfacht werden, um die Gleichung in \(O\) umzustellen \[\implies O = \dfrac{QP}{\sqrt{Q^2+2g\Delta hC^{2}P^{2}}}\]
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            We first need to convert these input values so to SI units.
        </span>
        <span class="multilang" lang="de">
            Wir müssen diese Eingabewerte also zunächst in SI-Einheiten umrechnen.
        </span>
        <br/>
        <p>
            \(Q = {@aq*10^6@} \textrm{ cm}^{3}\textrm{s}^{-1} = {@dispsf(float(aq),1)@} \textrm{ m}^{3}\textrm{s}^{-1}\)
        </p>
        <p>
            <span class="multilang" lang="en">
                diameter of pipe \(= {@diam*100@} \textrm{ cm} = {@float(diam)@} \textrm{ m} \implies P = \pi \dfrac{{@float(diam)@}^2}{4} = {@dispsf(float(ap),4)@}...\textrm{ m}\)
            </span>
            <span class="multilang" lang="de">
                Rohrdurchmesser \(= {@diam*100@} \textrm{ cm} = {@float(diam)@} \textrm{ m} \implies P = \pi \dfrac{{@float(diam)@}^2}{4} = {@dispsf(float(ap),4)@}...\textrm{m}\)
            </span>
        </p>
        <p>
            \(C = {@float(ac)@}\)
        </p>
        <p>
            \(\Delta h = {@ah*1000@} \textrm{ mm} = {@float(ah)@} \textrm{ m}\)
        </p>
        <p>
            \(g = 9.81 \textrm{ m}\textrm{s}^{-2}\)
        </p>
        <br/>
        <span class="multilang" lang="en">
            We can now substitute these values into our equation for \(O\) \[O = \dfrac{QP}{\sqrt{Q^2+2g\Delta hC^{2}P^{2}}} = \dfrac{{@dispsf(float(aq),1)@}\cdot {@dispsf(float(ap),4)@}...}{\sqrt{{@dispsf(float(aq),1)@}^2+2\cdot 9.81\cdot {@float(ah)@}\cdot
                                {@float(ac)@}^{2}\cdot ({@dispsf(float(ap),4)@}...)^{2}}} = {@ta2w@}\] To finally get the orifice diameter from this value we transpose the formula for the area of a circle \(\text{area} = \pi \cdot (\dfrac{diameter}{2})^2\) to get \[\text{orifice
                        diameter} = 2 \sqrt{\dfrac{O}{\pi}} = 2 \sqrt{\dfrac{{@ta2w@}}{\pi}} = {@ta2@}\]
        </span>
        <span class="multilang" lang="de">
            Wir können diese Werte nun in unsere Gleichung einsetzen für \(O\) \[O = \dfrac{QP}{\sqrt{Q^2+2g\Delta hC^ {2}P^{2}}} = \dfrac{{@dispsf(float(aq),1)@}\cdot {@dispsf(float(ap),4)@}...}{\sqrt{ {@dispsf(float(aq),1)@}^2+2\cdot 9.81\cdot {@float(ah)@}\cdot
                                {@float(ac)@}^{2}\cdot ({@dispsf(float(ap),4)@}...)^{2}}} = {@ta2w@}\] Um endlich den Öffnungsdurchmesser von diesem Wert zu erhalten, stellen wir die Formel für die Kreisfläche \(\text{Fläche} = \pi \cdot (\dfrac{Durchmesser}{2})^2\) um, um \[\text{Öffnungsdurchmesser} = 2 \sqrt{\dfrac{O}{\pi}} = 2 \sqrt{\dfrac{{@ta2w@}}{\pi}} = {@ta2@}\] zu erhalten
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>stack_unit_si_declare(true);
eqn: C*O*((2*g*h)/(1-(O^2/P^2)))

ta1: (Q*P)/sqrt(Q^2+2*g*h*C^2*P^2)

aq: rand([1/10000,1/5000])
diam: rand([1/10,1/5])
ap: %pi*((diam/2)^2)
ah: rand([1/5,2/5])
ac: rand([6/10,8/10])

eqnd: sqrt((4/%pi)*((Q*P)/sqrt(Q^2+2*g*h*C^2*P^2)))
ta2: stackunits(dispsf(float(ev(eqnd,Q=aq,P=ap,g=981/1000,h=ah,C=ac)),3), m)
ta2w: stackunits(dispsf(float(ev(ta1,Q=aq,P=ap,g=981/1000,h=ah,C=ac)),3), m^2)</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\( Q ={@aq*10^6@} \textrm{cm}^{3}\textrm{s}^{-1}, \, C =  {@float(ac)@} \) so \( O = {@ta2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>units</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Units</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>[2,-1]</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>Units</answertest>
                <sans>ans2</sans>
                <tans>stackunits(dispsf(float(ev(sqrt((4/%pi)*(ans1)),Q=aq,P=ap,g=981/1000,h=ah,C=ac)),3), m)</tans>
                <testoptions>[2,-1]</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You have correctly substituted the inputs into your equation for part (a) however because your part (a) is wrong this is also not the correct answer.
</span>
<span class="multilang" lang="de">
    Sie haben die Eingaben für Teil (a) richtig in Ihre Gleichung eingesetzt, aber da Ihr Teil (a) falsch ist, ist dies auch nicht die richtige Antwort.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>11819744</deployedseed>
        <deployedseed>78265521</deployedseed>
        <deployedseed>949671228</deployedseed>
        <deployedseed>78442244</deployedseed>
        <deployedseed>1698539413</deployedseed>
        <deployedseed>1889586457</deployedseed>
        <deployedseed>1545749093</deployedseed>
        <deployedseed>10719976</deployedseed>
        <deployedseed>402899688</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>stackunits(dispsf(float(ev(eqnd*100,Q=aq,P=ap,g=981/1000,h=ah,C=ac)),3), cm)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>dispsf(sqrt((4/%pi)*(2)),3)*m</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
