<quiz>
    <question type="stack">
        <name>
            <text>1.5.2.18 Umstellen [Transposition]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Make \(x\) the subject of
    </span>
    <span class="multilang" lang="de">
        Stellen Sie nach \(x\) um
    </span>
</p>
<br/>
<ol class="HELM_parts">
    <li>
        \(y = \dfrac{{@a1@}+x}{1-{@a1@}x}\), \(x =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \(y = \sqrt{\dfrac{x-{@a2@}}{x+1}}\), \(x =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            We start with the equation \[y = \dfrac{{@a1@}+x}{1-{@a1@}x}\] First multiply both sides by \(1-{@a1@}x\) and expand the brackets \[\implies y(1-{@a1@}x) = {@a1@}+x \implies y - {@a1@}xy = {@a1@} + x\] Add \({@a1@}xy\) to both sides of the equation
                                \[\implies y = {@a1@} + x + rxy\] Subtract \({@a1@}\) \[\implies y - {@a1@} = x + {@a1@}xy\] Factorise x out of the right-hand side of the equation \[\implies y - {@a1@} = x(1 + {@a1@}y)\] Divide both sides of the equation by \(1 + {@a1@}y\) to
                        make x the subject of the equation \[\implies x = \dfrac{y-{@a1@}}{1 + {@a1@}y}\]
        </span>
        <span class="multilang" lang="de">
            Wir beginnen mit der Gleichung \[y = \dfrac{{@a1@}+x}{1-{@a1@}x}\] Multiplizieren Sie zuerst beide Seiten mit \(1-{@a1@}x\) und multiplizieren Sie dann aus \[\implies y(1-{@a1@}x) = {@a1@}+x \implies y - {@a1@}xy = {@a1@} + x\] Addieren Sie \({@a1@}xy\) auf beiden Seiten der Gleichung
                                \[\implies y = {@a1@} + x + rxy\] Subtrahieren Sie \({@a1@}\) \[\implies y - {@a1@} = x + {@a1@}xy\] Faktorisieren Sie x aus der rechten Seite der Gleichung \[\implies y - {@a1@} = x(1 + {@a1@}y)\] Dividieren Sie beide Seiten der Gleichung durch \(1 + {@a1 @}y\) um x zu bestimmen \[\implies x = \dfrac{y-{@a1@}}{1 + {@a1@}y}\]
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            We start with the equation \[y = \sqrt{\dfrac{x-{@a2@}}{x+1}}\] First square both sides of the equation \[\implies y^2 = \dfrac{x-{@a2@}}{x+1}\] multiply both sides of the equation by \(x+1\) and expand the brackets \[\implies y^2(x+1) = x-{@a2@}
                                \implies xy^2+y^2 = x-{@a2@}\] Subtract \(xy^2\) from both sides \[y^2 = x-{@a2@}-xy^2\] Add \({@a2@}\) to both sides of the equation \[{@a2@} + y^2 = x - xy^2\] Factorise \(x\) out of the right-hand side of the equation \[{@a2@} + y^2 = x(1 -
                                y^2)\] Divide both sides of this equation by \(1 - y^2\) to make \(x\) the subject of the equation \[x = \dfrac{{@a2@} + y^2}{1 - y^2}\]
        </span>
        <span class="multilang" lang="de">
            Wir beginnen mit der Gleichung \[y = \sqrt{\dfrac{x-{@a2@}}{x+1}}\] Quadrieren Sie zuerst auf beiden Seiten der Gleichung \[\implies y^2 = \dfrac{x-{@a2@}}{x+1}\] multiplizieren Sie beide Seiten der Gleichung mit \(x+1\) und multiplizieren Sie aus \[\implies y ^2(x+1) = x-{@a2@} \implies xy^2+y^2 = x-{@a2@}\] Subtrahieren Sie \(xy^2\) auf beiden Seiten \[y^2 = x-{@a2@}-xy^2\] Addieren Sie \({@a2@}\) auf beiden Seiten der Gleichung \[{@a2@} + y^2 = x - xy^2\] Faktorisieren Sie \(x\) aus der rechten Seite der Gleichung \[{@a2@} + y^2 = x(1 -
                                y^2)\] Dividieren Sie beide Seiten dieser Gleichung durch \(1 - y^2\), um \(x\) durch \[x = \dfrac{{@a2@} + y^2}{1 - y^2}\] zu bestimmen
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand([4,5,6]);
a2: rand([1,2,3]);

eqn1: (a1+x)/(1-a1*x);
eqn2: sqrt((x-a2)/(x+1));
ta1: (y-a1)/(1+y*a1);
ta2: (y^2+a2)/(1-y^2);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(x = \dfrac{y-{@a1@}}{1 + {@a1@}y}, x = \dfrac{{@a2@} + y^2}{1 - y^2}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAssRules</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>[oneMul]</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You have effectively transposed the formula to make \(x\) the subject but it should be fully simplified to: \(\dfrac{y-{@a1@}}{1 + {@a1@}y}\)
</span>
<span class="multilang" lang="de">
    Sie haben die Formel effektiv umgestellt, um \(x\) zu bestimmen, aber sie sollte vollständig vereinfacht werden zu: \(\dfrac{y-{@a1@}}{ 1 + {@a1@}y}\)
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAssRules</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>[oneMul]</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You have effectively transposed the formula to make \(x\) the subject but it should be fully simplified to: \(\dfrac{{@a2@} + y^2}{1 - y^2}\)
</span>
<span class="multilang" lang="de">
    Sie haben die Formel effektiv umgestellt, um \(x\) zu bestimmen, aber sie sollte vollständig vereinfacht werden zu: \(\dfrac{{@a2@} + y^2). {1 - y^2}\)
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>833470693</deployedseed>
        <deployedseed>1683708407</deployedseed>
        <deployedseed>1657001486</deployedseed>
        <deployedseed>129963072</deployedseed>
        <deployedseed>1327837126</deployedseed>
        <deployedseed>1064686271</deployedseed>
        <deployedseed>1061195832</deployedseed>
        <deployedseed>434961595</deployedseed>
        <deployedseed>1615919156</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>(y)/(1+y*a1) - (a1)/(1+y*a1)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>(y^2)/(1-y^2) + (a2)/(1-y^2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>x = (y-a1)/(1+y*a1)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>x = (y^2+a2)/(1-y^2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore></expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore></expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
