<quiz>
    <question type="stack">
        <name>
            <text>1.5.2.5 Substitution</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        If \(y = \sqrt{\dfrac{x}{z}}\), find \(y \text{ when } x = {@dispdp(a1,3)@} \text{ and } z = {@dispdp(a2,3)@}:\)
    </span>
    <span class="multilang" lang="de">
        Wenn \(y = \sqrt{\dfrac{x}{z}}\), bestimmen Sie \(y \text{ wenn } x = {@dispdp(a1, 3)@} \text{ und } z = {@dispdp(a2,3)@}:\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        (give your answer as a decimal to three decimal places, note that trailing zeroes are required i.e 2.3 must be written as 2.300)
    </span>
    <span class="multilang" lang="de">
        (Geben Sie Ihr Ergebnis als Dezimalzahl mit drei Dezimalstellen an. Beachten Sie, dass nachgestellte Nullen erforderlich sind, d. h. 2,3 muss als 2,300 geschrieben werden)
    </span>
</p>
<br/>
<p>
    \(y =\) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<br/>
<p>
    <span class="multilang" lang="en">
        In this example we have \(x = {@dispdp(a1,3)@} \text{ and } y = {@dispdp(a2,3)@}\). Substituting these values into the formula and then using a calculator we find:
    </span>
    <span class="multilang" lang="de">
        In diesem Beispiel haben wir \(x = {@dispdp(a1,3)@} \text{ und } y = {@dispdp(a2,3)@}\). Wenn wir diese Werte in die Formel einsetzen und dann einen Taschenrechner verwenden, erhalten wir:
    </span>
</p>
<p>
    \[y = \sqrt{\dfrac{x}{z}} = \sqrt{\dfrac{{@a1@}}{{@a2@}}} = {@ev(eqs, x=a1,z=a2)@} = {@ta@}_{\text{ 3.d.p}} \]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>eqs: sqrt(x/z);
a1: rand([11.2,13.6,15.8]);
a2: rand([14.2,15.6,16.2]);
ta:dispdp(float(ev(eqs, x=a1,z=a2)),3);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\(y=\sqrt{\dfrac{{@a1@}}{{@a2@}}} = {@ta@}\) 3.d.p</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>sqrt,/</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>mindp:3</options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>1</name>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>0.011</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>3</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your final answer is almost correct, you may have made a mistake when rounding to three decimal places.
</span>
<span class="multilang" lang="de">
    Ihre endgültige Antwort ist fast richtig, möglicherweise haben Sie beim Runden auf drei Dezimalstellen einen Fehler gemacht.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2083412336</deployedseed>
        <deployedseed>1913237114</deployedseed>
        <deployedseed>410858090</deployedseed>
        <deployedseed>1492207349</deployedseed>
        <deployedseed>1634917600</deployedseed>
        <deployedseed>761589406</deployedseed>
        <deployedseed>490277940</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>dispdp(float(ev(eqs, x=a1,z=a2)),2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore></expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-3-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>dispdp(float(ev(eqs+0.1, x=a1,z=a2)),3)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>dispdp(float(ev(eqs+0.01, x=a1,z=a2)),3)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
