<quiz>
    <question type="stack">
        <name>
            <text>6.3.1.3 Log of a product</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Let \( \log_ab = {@N@}\) and \( \log_ac = {@M@}\). Express \( \log_a(bc) \) in terms of \({@N@}\) and \({@M@}\).
    </span>
    <span class="multilang" lang="de">
        Seien \( \log_ab = {@N@}\) und \( \log_ac = {@M@}\). Drücken Sie \( \log_a(bc) \) durch \({@N@}\) und \({@M@}\) aus.
    </span>
</p>
<p>
    \( \log_a(bc) \) = [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Since \({@N@} = \log_ab\), \(b = a^{{@N@}}\). Similarly, \({@M@} = \log_ac\) implies \(c = a^{{@M@}}\).
    </span>
    <span class="multilang" lang="de">
        Da \({@N@} = \log_ab\), gilt \(b = a^{{@N@}}\). Ebenso impliziert \({@M@} = \log_ac\) \(c = a^{{@M@}}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Hence, \(bc = a^{{@N@}}a^{{@M@}} = a^{{@N@}+{@M@}}\), which implies \(\log_a(bc) = {@N@}+{@M@}\).
    </span>
    <span class="multilang" lang="de">
        Daher ist \(bc = a^{{@N@}}a^{{@M@}} = a^{{@N@}+{@M@}}\), was \(\log_a(bc) = {@N@}+{@M@}\) impliziert.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore, using the property of indices, we have shown that \[\log_a(bc)={@N@}+{@M@}=\log_ab+\log_ac.\]
    </span>
    <span class="multilang" lang="de">
        Daher haben wir unter Verwendung der Eigenschaft von Potenzen gezeigt, dass \[\log_a(bc)={@N@}+{@M@}=\log_ab+\log_ac.\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>N:rand([x,y,z]);
M:rand([u,v,w]);
ta:N+M;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\( \log_a(a^{{@N@}}a^{{@M@}}) = {@N@} + {@M@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>log,lg,ln</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>634001404</deployedseed>
        <deployedseed>475238151</deployedseed>
        <deployedseed>299050156</deployedseed>
        <deployedseed>119881246</deployedseed>
        <deployedseed>1931509851</deployedseed>
        <deployedseed>1185888086</deployedseed>
        <deployedseed>993727916</deployedseed>
        <deployedseed>1917326522</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>N+M</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>N-M</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
