<quiz>
    <question type="stack">
        <name>
            <text>6.3.4.2 Computing logs arbitrary base</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Re-express \(\log_{{@aa@}} {@bb@}\) using base \(10\) and base \(e\).
    </span>
    <span class="multilang" lang="de">
        Drücken Sie \(\log_{{@aa@}} {@bb@}\) erneut aus, indem Sie die Basis \(10\) und die Basis \(e\) verwenden.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Syntax hint: Enter lg for logs base \(10\), and ln for natural logs (base \(e\)), e.g. enter lg(2) for \(\log_{10}2\), enter ln(2) for \(\log_e2\).
    </span>
    <span class="multilang" lang="de">
        Syntaxhinweis: Geben Sie lg für Logarithmen zur Basis \(10\) und ln für natürliche Logarithmen (Basis \(e\)) ein, z.B. geben Sie lg(2) für \(\log_{10}2\) ein, geben Sie ln(2) für \(\log_e2\) ein.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Base \(10\)
            </span>
            <span class="multilang" lang="de">
                Basis \(10\)
            </span>
        </p>
        <p>
            \(\log_{{@aa@}} {@bb@} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Base \(e\)
            </span>
            <span class="multilang" lang="de">
                Basis \(e\)
            </span>
        </p>
        <p>
            \( \log_{{@aa@}} {@bb@} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Now, use your calculator to determine the value of \(\log_{{@aa@}} {@bb@}\).
            </span>
            <span class="multilang" lang="de">
                Benutzen Sie jetzt Ihren Taschenrechner, um den Wert von \(\log_{{@aa@}} {@bb@}\) zu bestimmen.
            </span>
        </p>
        <p>
            \(\log_{{@aa@}} {@bb@} = \) [[input:ans3]] [[validation:ans3]]
            <span class="multilang" lang="en">
                (Round up and enter \(4\) decimal places. If the answer has less than \(4\) decimal places, then enter zero(s) at the end.)
            </span>
            <span class="multilang" lang="de">
                (Runden Sie auf und geben Sie \(4\) Dezimalstellen ein. Wenn die Antwort weniger als \(4\) Dezimalstellen hat, dann geben Sie Null(en) am Ende ein. )
            </span>
            [[feedback:prt3]]
        </p>
    </li>
</ol>
<p>
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Recall the formula for converting logarithms from one base to another \[ \log_a b = \dfrac{\log_p b}{\log_p a}. \]
    </span>
    <span class="multilang" lang="de">
        Erinnern Sie sich an die Formel zur Umwandlung von Logarithmen von einer Basis in eine andere \[ \log_a b = \dfrac{\log_p b}{\log_p a}. \]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        (a) Applying the above formula for \(a = {@aa@}, b ={@bb@} \) and \(p = 10\), we get \[ \log_{{@aa@}} {@bb@} = \dfrac{\log_{10} {@bb@}}{\log_{10}{@aa@}}.\]
    </span>
    <span class="multilang" lang="de">
        (a) Wenden wir die obige Formel für \(a = {@aa@}, b ={@bb@} \) und \(p = 10\) an,erhalten wir \[ \log_{{@aa@}} {@bb@} = \dfrac{\log_{10} {@bb@}}{\log_{10}{@aa@}}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        (b) Applying the same formula for \(a = {@aa@}, b={@bb@}\) and \(p = e\), we get \[ \log_{{@aa@}} {@bb@} = \dfrac{\ln {@bb@}}{\ln{@aa@}}.\]
    </span>
    <span class="multilang" lang="de">
        (b) Wenden wir die gleiche Formel für \(a = {@aa@}, b={@bb@}\) und \(p = e\) an, erhalten wir \[ \log_{{@aa@}} {@bb@} = \dfrac{\ln {@bb@}}{\ln{@aa@}}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        (c) Note that \(\log_{{@aa@}}{@bb@}\) cannot be determined directly on your calculator since logs to base \({@aa@}\) are not available, but it can be found by changing to base \(10\) or base \(e\).
    </span>
    <span class="multilang" lang="de">
        (c) Beachten Sie, dass \(\log_{{@aa@}}{@bb@}\) nicht direkt auf Ihrem Taschenrechner bestimmt werden kann, da Logarithmen zur Basis \({@aa@}\) nicht verfügbar sind, sie können aber gefunden werden, indem Sie zur Basis \(10\) oder zur Basis \(e\) wechseln.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Using logs base \(10\) on your calculator, we get \[ \log_{{@aa@}} {@bb@} = \dfrac{\log_{10} {@bb@}}{\log_{10}{@aa@}} = {@ta3@}.\]
    </span>
    <span class="multilang" lang="de">
        Wenn Sie Logarithmen zur Basis \(10\) auf Ihrem Taschenrechner verwenden, erhalten wir \[ \log_{{@aa@}} {@bb@} = \dfrac{\log_{ 10} {@bb@}}{\log_{10}{@aa@}} = {@ta3@}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Alternatively, using logs base \(e\) on your calculator, we get \[ \log_{{@aa@}} {@bb@} = \dfrac{\ln {@bb@}}{\ln{@aa@}} = {@ta3@}.\]
    </span>
    <span class="multilang" lang="de">
        Alternativ erhalten wir unter Verwendung der Logarithmen zur Basis \(e\) auf Ihrem Taschenrechner \[ \log_{{@aa@}} {@bb@} = \dfrac{\ln {@bb@}}{\ln{@aa@}} = {@ta3@}.\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa:rand(5)+20;
bb:rand(6)+2;
ta3:decimalplaces(float(ln(bb)/ln(aa)),4);
simp:false;
ta1:lg(bb)/lg(aa);
ta2:ln(bb)/ln(aa);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(log_{{@aa@}} {@bb@} = {@ta1@} = {@ta2@} = {@ta3@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>[[BASIC-ALGEBRA]]</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>[[BASIC-ALGEBRA]]</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>dispdp(ta3,4)</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>[[BASIC-ALGEBRA]],log,ln,lg</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>maxdp:4,mindp:4</options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1785931336</deployedseed>
        <deployedseed>874712443</deployedseed>
        <deployedseed>1358879553</deployedseed>
        <deployedseed>322163396</deployedseed>
        <deployedseed>1605127270</deployedseed>
        <deployedseed>216329314</deployedseed>
        <deployedseed>461944937</deployedseed>
        <deployedseed>777861504</deployedseed>
        <deployedseed>1802178369</deployedseed>
        <deployedseed>1199985716</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>dispdp(ta3,4)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>dispdp(ev(ta3+1,simp),4)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
