<quiz>
    <question type="stack">
        <name>
            <text>9.4.3.2 Evaluate the vector product</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Use the key point of this section to find the vector product of
        <br/>
        \(\underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j}\) and \(\underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}\).
    </span>
    <span class="multilang" lang="de">
        Verwenden Sie den Schlüsselpunkt dieses Abschnitts, um das Vektorprodukt von
        <br/>
        \(\underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j}\) und \(\underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}\) zu finden.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        (Note that in this case there are no \(\underline{k}\) components so \(a_3\) and \(b_3\) are both zero).
    </span>
    <span class="multilang" lang="de">
        (Beachten Sie, dass es in diesem Fall keine \(\underline{k}\)-Komponenten gibt, also sind \(a_3\) und \(b_3\) beide
                null).
    </span>
</p>
<p>
</p>
<p>
    <span class="math inline">
        \(\qquad \underline{ p}\times\underline{ q}=\)
    </span>
    [[input:ans1]] [[validation:ans1]] \(\underline{k}\) [[feedback:prt1]]
</p>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are ginen the vectors: \[ \underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j} \, \text{ and } \, \underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}.\]
    </span>
    <span class="multilang" lang="de">
        Wir erhalten die Vektoren: \[ \underline{ p}={@p1@}\underline{ i}+{@p2@}\underline{ j} \, \text{ und } \, \underline{ q}={@q1@}\underline{ i}-{@q2@}\underline{ j}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Identifying \(a_1={@p1@}\), \(a_2={@p2@}\), \(a_3=0\), \(b_1={@q1@}\), \(b_2=-{@q2@}\),
            \(b_3=0\) we find
    </span>
    <span class="multilang" lang="de">
        Durch Identifizieren von \(a_1={@p1@}\), \( a_2={@p2@}\), \(a_3=0\), \(b_1={@ q1@}\),
                \(b_2=-{@q2@}\),
                \(b_3=0\) erhalten wir
    </span>
</p>
<p>
    <span class="math display">
        \[\begin{aligned}
\underline{a}\times \underline{b}&amp;= (({@p2@})(0)-(0)(-{@q2@}))\underline{i}-(({@p1@})(0)-(0)({@q1@}))\underline{j}+(({@p1@})(-{@q2@})-({@p2@})({@q1@}))\underline{k}\\
&amp;= {@ta@}\underline{k}\end{aligned}.\]
    </span>
</p>
<p>
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>p1: 2+rand(3)
p2: 2+rand(4)
q1: 2+rand(5)
q2: 2+rand(6)
ta: -p1*q2 -p2*q1</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(({@p1@}\underline{ i}+{@p2@}\underline{ j}) \times ({@q1@}\underline{ i}-{@q2@}\underline{ j}) = {@ta@}\underline{ k} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2035175511</deployedseed>
        <deployedseed>755940633</deployedseed>
        <deployedseed>121058625</deployedseed>
        <deployedseed>372506200</deployedseed>
        <deployedseed>2145459699</deployedseed>
        <deployedseed>1550917835</deployedseed>
        <deployedseed>809479233</deployedseed>
        <deployedseed>897124644</deployedseed>
        <deployedseed>1076608172</deployedseed>
        <deployedseed>19453885</deployedseed>
        <deployedseed>228579085</deployedseed>
        <deployedseed>2021917302</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
