<quiz>
    <question type="stack">
        <name>
            <text>9.4.4.7 Associativity of the vector product</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For the vectors \(\underline{ a}={@a1@}\underline{ i} +{@a3@}\underline{ k}\), \(\underline{ b}=-3\underline{ j} +{@b3@}\underline{ k}\), and \(\underline{ c}={@c1@}\underline{ i} -2\underline{ j}\):
    </span>
    <span class="multilang" lang="de">
        Für die Vektoren \(\underline{ a}={@a1@}\underline{ i} +{@a3@}\underline{ k}\), \(\underline{ b}=-3\underline{ j} +{@b3@}\underline{ k}\) und
                \(\underline{ c}={@c1@}\underline{ i} -2\underline{ j}\):
    </span>
</p>
<p>
    <br/>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Compute \(\underline{ a}\times(\underline{ b}\times\underline{ c}) = \)
            </span>
            <span class="multilang" lang="de">
                Berechnen Sie \(\underline{ a}\times(\underline{ b}\times\underline{ c}) = \)
            </span>
            [[input:ans1]] [[validation:ans1]] \(\underline{i}\) \(+\) [[input:ans2]] [[validation:ans2]] \(\underline{j}\) \(+\) [[input:ans3]]
                                                [[validation:ans3]] \(\underline{k}\) [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find \((\underline{ a}\times\underline{ b})\times\underline{ c} = \)
            </span>
            <span class="multilang" lang="de">
                Finden Sie \((\underline{ a}\times\underline{ b})\times\underline{ c} = \)
            </span>
            [[input:ans4]] [[validation:ans4]] \(\underline{i}\) \(+\) [[input:ans5]] [[validation:ans5]] \(\underline{j}\) \(+\) [[input:ans6]]
                                                [[validation:ans6]] \(\underline{k}\) [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                In general, is the vector product associative?
            </span>
            <span class="multilang" lang="de">
                Ist das Vektorprodukt im Allgemeinen assoziativ?
            </span>
            [[input:ans7]] [[validation:ans7]]
        </p>
        <p>
            [[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                First we have:
            </span>
            <span class="multilang" lang="de">
                Zuerst erhalten wir:
            </span>
        </p>
        <p>
            \(\begin{aligned} \underline{ b}\times \underline{ c} &amp;= \left| \begin{matrix} \underline{i} &amp; \underline{j} &amp; \underline{k} \\ 0 &amp; -3 &amp; {@b3@} \\ {@c1@} &amp; -2 &amp; 0 \\ \end{matrix} \right| \\ &amp; = \underline{i}((-3)(0)-({@b3@})(-2))
                                                - \underline{j}((0)(0)-({@b3@})({@c1@})) + \underline{k}((0)(-2)-(-3)({@c1@})) \\ &amp; =  {@2*b3@}\underline{i}+{@b3*c1@}\underline{j} +{@3*c1@}\underline{k} \end{aligned}\).
        </p>
        <p>
            <span class="multilang" lang="en">
                Thus:
            </span>
            <span class="multilang" lang="de">
                Also:
            </span>
        </p>
        <p>
            \(\begin{aligned} \underline{ a}\times(\underline{ b}\times\underline{ c}) &amp;= \left| \begin{matrix} \underline{i} &amp; \underline{j} &amp; \underline{k} \\ {@a1@} &amp; 0 &amp; {@a3@} \\ {@2*b3@} &amp; {@b3*c1@} &amp; {@3*c1@} \\ \end{matrix} \right|
                                                \\ &amp; = \underline{i}((0)({@3*c1@})-({@a3@})({@b3*c1@})) - \underline{j}(({@a1@})({@3*c1@})-({@a3@})({@2*b3@})) + \underline{k}(({@a1@})({@b3*c1@})-(0)({@2*b3@})) \\ &amp; = -{@ta_i1@} \underline{i} - {@ta_j1@} \underline{j}
                                                + {@ta_k1@} \underline{k} \end{aligned}\)
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                First we have:
            </span>
            <span class="multilang" lang="de">
                Zuerst erhalten wir:
            </span>
        </p>
        <p>
            \(\begin{aligned} \underline{ a}\times \underline{ b} &amp;= \left| \begin{matrix} \underline{i} &amp; \underline{j} &amp; \underline{k} \\ {@a1@} &amp; 0 &amp; {@a3@} \\ 0 &amp; -3 &amp; {@b3@} \\ \end{matrix} \right| \\ &amp; = \underline{i}((0)({@b3@})-({@a3@})(-3))
                                                - \underline{j}(({@a1@})({@b3@})-({@a3@})(0)) + \underline{k}(({@a1@})(-3)-(0)(0)) \\ &amp; =  {@3*a3@}\underline{i}-{@a1*b3@}\underline{j} -{@3*a1@}\underline{k} \end{aligned}\).
        </p>
        <p>
            <span class="multilang" lang="en">
                Thus:
            </span>
            <span class="multilang" lang="de">
                Also:
            </span>
        </p>
        <p>
            \(\begin{aligned} (\underline{ a}\times\underline{ b})\times\underline{ c} &amp;= \left| \begin{matrix} \underline{i} &amp; \underline{j} &amp; \underline{k} \\ {@3*a3@} &amp; -{@a1*b3@} &amp; -{@3*a1@} \\ {@c1@} &amp; -2 &amp; 0 \\ \end{matrix} \right|
                                                \\ &amp; = \underline{i}((-{@a1*b3@})(0)-(-{@3*a1@})(-2)) - \underline{j}(({@3*a3@})(0)-(-{@3*a1@})({@c1@})) + \underline{k}(({@3*a3@})(-2)-(-{@a1*b3@})({@c1@})) \\ &amp; = -{@ta_i2@} \underline{i} - {@ta_j2@} \underline{j}
                                                + {@ta_k2@} \underline{k} \end{aligned}\)
        </p>
    </li>
    <li>
        <span class="multilang" lang="en">
            In general, the vector product is not associative.
        </span>
        <span class="multilang" lang="de">
            Im Allgemeinen ist das Vektorprodukt nicht assoziativ.
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: 8 + rand(4)
a3: 2 + rand(3)
b1: 8 + rand(4)
b3: 2 + rand(3)
c1: 2 + rand(4)
c3: 2 + rand(3)
ta_i1: a3*b3*c1
ta_j1: 3*a1*c1 - 2*a3*b3
ta_k1: a1*c1*b3
ta_i2: 6*a1
ta_j2: 3*a1*c1
ta_k2: -6*a3 +a1*b3*c1</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[
\underline{ a}\times(\underline{ b}\times\underline{ c}) = -{@ta_i1@} \underline{i} - {@ta_j1@} \underline{j} + {@ta_k1@} \underline{k}
\]
\[
(\underline{ a}\times\underline{ b})\times\underline{ c} = -{@ta_i2@} \underline{i} - {@ta_j2@} \underline{j} - {@ta_k2@} \underline{k}
\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>-ta_i1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>-ta_j1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>ta_k1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>numerical</type>
            <tans>-ta_i2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>numerical</type>
            <tans>-ta_j2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>numerical</type>
            <tans>ta_k2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans7</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>-ta_i1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>-ta_j1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta_k1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>-ta_i2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>-ta_j2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>ta_k2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>PropLogic</answertest>
                <sans>ans7</sans>
                <tans>false</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>752212846</deployedseed>
        <deployedseed>886008112</deployedseed>
        <deployedseed>1562199388</deployedseed>
        <deployedseed>1572791435</deployedseed>
        <deployedseed>1316779509</deployedseed>
        <deployedseed>1322767625</deployedseed>
        <deployedseed>1757028178</deployedseed>
        <deployedseed>476249521</deployedseed>
        <deployedseed>653989543</deployedseed>
        <deployedseed>609354042</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>-ta_i1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>-ta_j1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta_k1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>-ta_i2</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>-ta_j2</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta_k2</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>false</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-3-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-3-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>true</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
