<quiz>
    <question type="stack">
        <name>
            <text>3.2.3.3 Abschluss der Quadratübungen 2 [Completing the square exercises 2]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the equation \(4x^2-{@aa@}x+{@bb@}=0\) by completing the square.
    </span>
    <span class="multilang" lang="no">
        Løs ligningen \(4x^2-{@aa@}x+{@bb@}=0\) ved å fullføre kvadratet.
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die Gleichung \(4x^2-{@aa@}x+{@bb@}=0\), indem Sie die quadratische Ergänzung verwenden.
    </span>
</p>
<p>
    \(x =\) [[input:ans1]] [[validation:ans1]]
    <span class="multilang" lang="en">
        or
    </span>
    <span class="multilang" lang="no">
        eller
    </span>
    <span class="multilang" lang="de">
        oder
    </span>
    \(x =\) [[input:ans2]] [[validation:ans2]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        First examine the two left-most terms in the equation: \(4x^2-{@aa@}x\). Complete the square for these terms:
    </span>
    <span class="multilang" lang="no">
        Undersøk først de to leddene lengst til venstre i ligningen: \(4x^2-{@aa@}x\). Fullfør kvadratet for disse uttrykkene:
    </span>
    <span class="multilang" lang="de">
        Untersuchen Sie zuerst die beiden Terme ganz links in der Gleichung: \(4x^2-{@aa@}x\). Ergänzen Sie das Quadrat für diese Terme:
    </span>
</p>
<p>
    \[4x^2-{@aa@}x = 4\left(x^2-\frac{{@aa@}}{4} x\right)=4\left(x-\frac{{@aa@}}{8}\right)^2-\frac{a^2}{16}.\]
</p>
<p>
    <span class="multilang" lang="en">
        Use the above result to rewrite the equation \(4x^2-{@aa@}x+{@bb@}\) in the form \(2(x-\ ?)^2+?=0\):
    </span>
    <span class="multilang" lang="no">
        Bruk resultatet ovenfor til å omskrive ligningen \(4x^2-{@aa@}x+{@bb@}\) i formen \(2(x-\ ?)^2+?=0\):
    </span>
    <span class="multilang" lang="de">
        Verwenden Sie das obige Ergebnis, um die Gleichung \(4x^2-{@aa@}x+{@bb@}\) in die Form \(2(x-\ ?)^2+?=0\) zu bringen:
    </span>
</p>
<p>
    \[4x^2-{@aa@}x+{@bb@} = 4\left(x-\frac{{@aa@}}{8}\right)^2-\frac{{@a@}^2}{16}+1 =0.\]
</p>
<p>
    <span class="multilang" lang="en">
        From this now obtain the roots:
    </span>
    <span class="multilang" lang="no">
        Fra dette får vi røttene:
    </span>
    <span class="multilang" lang="de">
        Hieraus erhalten Sie nun die Lösungen:
    </span>
</p>
<p>
    \[\left(x-\frac{{@aa@}}{8}\right)^2=\frac{1}{4}\left(\frac{{@a@}^2}{16}-1\right). \text{
    <span class="multilang" lang="en">
        Therefore
    </span>
    <span class="multilang" lang="no">
        Dermed
    </span>
    <span class="multilang" lang="de">
        Daher
    </span>
    }\quad x={@aa@}/8\pm \frac{\sqrt{{@aa@}^2/16-{@bb@}}}{2}.\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa: 6+rand(5);
bb: 1
ta1: aa/8+sqrt(aa^2/16-bb)/2
ta2: aa/8-sqrt(aa^2/16-bb)/2</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\(4x^2-{@aa@}x+{@bb@}=0,\) roots \( x = {@ta1@}, x = {@ta2@} , \, a={@aa@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>{ratsimp(ans1),ratsimp(ans2)}</sans>
                <tans>{ratsimp(ta1),ratsimp(ta2)}</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>695849901</deployedseed>
        <deployedseed>1686896826</deployedseed>
        <deployedseed>1391515015</deployedseed>
        <deployedseed>1319430899</deployedseed>
        <deployedseed>898342988</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
