<quiz>
    <question type="stack">
        <name>
            <text>3.3.9.4 Zerlege Viertel mit wiederholten Nullstellen x=+-a [Factorise quartic with repeated roots x=+-a]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the equation {@eqn@}:
    </span>
    <span class="multilang" lang="no">
        Løs likningen {@eqn@}:
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die Gleichung {@eqn@}:
    </span>
</p>
\( x = \) [[input:ans1]] [[validation:ans1]]
<p>
</p>
<p>
    <b>
        <span class="multilang" lang="en">
            Give your answer as a set of precise numbers. Only provide the
            <u>
                distinct
            </u>
            solutions.
        </span>
        <span class="multilang" lang="no">
            Gi svaret ditt som et sett av presise tall. Gi bare de
            <u>
                distinkte
            </u>
            løsningene.
        </span>
        <span class="multilang" lang="de">
            Geben Sie Ihre Antwort als eine Menge konkreter Zahlen an. Geben Sie nur die
            <u>
                unterschiedlichen
            </u>
            Lösungen an.
        </span>
    </b>
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Try viewing this as a quadratic function of \(x^2.\) We solve for \(x^2\), from which we can obtain \(x\).
    </span>
    <span class="multilang" lang="no">
        Prøv å se dette som en kvadratisk funksjon av \(x^2.\) Vi løser for \(x^2\), hvorfra vi kan få \(x\).
    </span>
    <span class="multilang" lang="de">
        Versuchen Sie, dies als quadratische Funktion von \(x^2.\) zu betrachten. Wir lösen nach \(x^2\) auf, woraus wir \(x\) erhalten können.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We are given the equation {@eqn@}. Let \(y = x^2\). Then the equation becomes \[y^2 - {@2*aa^2*y@} + {@aa^4@} = 0.\] Recognise that this is equivalent to \((y - {@aa^2@})^2 = 0.\)
    </span>
    <span class="multilang" lang="no">
        Vi får ligningen {@eqn@}. La \(y = x^2\). Deretter blir ligningen \[y^2 - {@2*aa^2*y@} + {@aa^4@} = 0.\] Gjenkjenne at dette tilsvarer \((y - {@aa^2@})^2 = 0.\)
    </span>
    <span class="multilang" lang="de">
        Uns wird die Gleichung {@eqn@} gegeben. Sei \(y = x^2\). Dann wird die Gleichung zu \[y^2 - {@2*aa^2*y@} + {@aa^4@} = 0.\] Beachten Sie, dass dies äquivalent zu \((y - {@aa^2 @})^2 = 0\) ist.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore \(y = {@aa^2@}\) which implies \(x = \pm{@aa@}.\)
    </span>
    <span class="multilang" lang="no">
        Derfor er \(y = {@aa^2@}\) som innebærer at \(x = \pm{@aa@}.\)
    </span>
    <span class="multilang" lang="de">
        Daher gilt \(y = {@aa^2@}\), was \(x = \pm{@aa@}.\) impliziert
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand(3)+1;
ta: setify([aa,-aa]);

simp:true;
factor: x + aa;
%eqn: (x + aa)^2*(x - aa)^2;
eqn: x^4 - 2*aa^2*x^2 + aa^4 = 0;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\({@eqn@}, \quad sol: {@ta@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{a,b,c,d}</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0.1</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A useful property to notice is that this quadratic can be viewed as a quadratic polynomial in \(x^2\), so we can solve {@eqn@} for \(x^2\) (by factorising or using the quadratic formula) which will give us two expressions for \(x^2\), and this expression
        gives us the values of \(x\) as desired.
    </span>
    <span class="multilang" lang="de">
        Eine nützliche Eigenschaft ist, dass dieses Polynom als quadratisches Polynom in \(x^2\) angesehen werden kann, sodass wir {@eqn@} nach \(x^2\) (durch Faktorisieren oder Verwenden der quadratischen Formel) lösen können, was uns zwei Ausdrücke für \(x^2\) gibt, und dieser Ausdruck gibt uns die gewünschten Werte von \(x\) an.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>967783658</deployedseed>
        <deployedseed>2059533060</deployedseed>
        <deployedseed>778014411</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>{aa}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>{-aa}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
