<quiz>
    <question type="stack">
        <name>
            <text>3.3.9.3 Faktorisiere einen gegebenen kubischen Faktor mit wiederholter Nullstelle [Factorise cubic given factor with repeated root]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Given that \(x={@aa@}\) is a root of \({@eqn@} = 0,\) determine the other roots. Hence record all roots.
    </span>
    <span class="multilang" lang="no">
        Gitt at \(x={@aa@}\) er en rot av \({@eqn@} = 0,\) bestem røttene.
    </span>
    <span class="multilang" lang="de">
        Angenommen, dass \(x={@aa@}\) eine Nullstelle von \({@eqn@} = 0\) ist, bestimmen Sie die anderen Nullstellen. Notieren Sie damit alle Nullstellen.
    </span>
</p>
[[input:ans1]] [[validation:ans1]]
<p>
</p>
<p>
    <b>
        <span class="multilang" lang="en">
            Give your answer as a set of precise numbers.
        </span>
        <span class="multilang" lang="no">
            Gi svaret ditt som et set av presise tall.
        </span>
        <span class="multilang" lang="de">
            Geben Sie Ihre Antwort als eine Menge konkreter Zahlen an.
        </span>
    </b>
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        By expanding \({@eqn@}=({@factor@})({x^2+bx+c})\) and equating coefficients we find \({@eqn@}=({@factor@})({@ratsimp(eqn/factor)@})\).
    </span>
    <span class="multilang" lang="no">
        Ved å utvide \({@eqn@}=({@factor@})({x^2+bx+c})\) og likestille koeffisientene finner vi \({@eqn@}=({@factor@} )({@ratsimp(eqn/factor)@})\).
    </span>
    <span class="multilang" lang="de">
        Durch Erweitern von \({@eqn@}=({@factor@})({x^2+bx+c})\) und Gleichsetzen von Koeffizienten erhalten wir \({@eqn@}=({@factor@})({@ratsimp(eqn/factor)@})\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        From here we can solve the quadratic and find that the roots of the cubic are {@ta@}.
    </span>
    <span class="multilang" lang="no">
        Herfra kan vi bruke abc-formelen og finne at røttene til andregradsuttrykket er {@ta@}.
    </span>
    <span class="multilang" lang="de">
        Von hier aus können wir die quadratische Gleichung lösen und feststellen, dass die Nullstellen der kubischen Gleichung {@ta@} sind.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand(5)+1;
bb: rand(4)-6;
cc: bb;
ta: setify([aa,bb,cc]);

simp:true;
factor: x - aa;
%eqn: (x - cc)*(x - bb)*(x - aa);
eqn: x^3 - (aa + bb + cc)*x^2 + (aa*bb + bb*cc + cc*aa)*x - aa*bb*cc;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\({@eqn@}=0,\) sol: [{@aa@},{@bb@},{@cc@}]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{a,b,c}</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0.1</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Try taking the factor of \( (x - {@aa@}) \) out of {@eqn@} in order to be left with a quadratic function which you can factorise yourself or solve using the quadratic formula.
    </span>
    <span class="multilang" lang="de">
        Versuchen Sie, den Faktor von \( (x - {@aa@}) \) aus {@eqn@} herauszurechnen, um eine quadratische Funktion zu erhalten, die Sie selbst faktorisieren oder mit der quadratischen Formel lösen können.
    </span>
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1270450540</deployedseed>
        <deployedseed>1846018860</deployedseed>
        <deployedseed>698463200</deployedseed>
        <deployedseed>726162301</deployedseed>
        <deployedseed>713289394</deployedseed>
        <deployedseed>472262832</deployedseed>
        <deployedseed>521056434</deployedseed>
        <deployedseed>1229741939</deployedseed>
        <deployedseed>1747365201</deployedseed>
        <deployedseed>1722426606</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>{aa,bb}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>{bb}</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
