<quiz>
    <question type="stack">
        <name>
            <text>3.3.1.3 Aufgabe: Grad des Koeffizienten von Linearen [Task: Degree of coefficient of linear]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            If \[{@aa@} x^2 + {@d1@} x + {@d2@} = (x + {@dd@}) \times p(x),\] the degree of the polynomial \(p(x)\) is
        </span>
        <span class="multilang" lang="no">
            Hvis \[{@aa@} x^2 + {@d1@} x + {@d2@} = (x + {@dd@}) \times p(x),\] er graden til polynomet \(p(x)\)
        </span>
        <span class="multilang" lang="de">
            Wenn \[{@aa@} x^2 + {@d1@} x + {@d2@} = (x + {@dd@}) \times p (x),\] ist der Grad des Polynoms \(p(x)\)
        </span>
        [[input:ans1]] [[validation:ans1]].
    </li>
    <li>
        <span class="multilang" lang="en">
            What is the coefficient of \(x\) in the polynomial \(p(x)\)?
        </span>
        <span class="multilang" lang="no">
            Hva er koeffisienten til \(x\) i polynomet \(p(x)\)?
        </span>
        <span class="multilang" lang="de">
            Was ist der Koeffizient von \(x\) im Polynom \(p(x)\)?
        </span>
        [[input:ans2]] [[validation:ans2]] [[feedback:prt1]]
        <br/>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            A quadratic is the product of two linear polynomials.
        </span>
        <span class="multilang" lang="no">
            En andregradspolynom er produktet av to lineære polynomer.
        </span>
        <span class="multilang" lang="de">
            Ein Quadrat ist das Produkt zweier linearer Polynome.
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            The coefficient must be \({@aa@}\) in order to generate the term \({@aa@}x^2\) when the linear factors are multiplied together.
        </span>
        <span class="multilang" lang="no">
            Koeffisienten må være \({@aa@}\) for å generere leddet \({@aa@}x^2\) når de lineære faktorene multipliseres sammen.
        </span>
        <span class="multilang" lang="de">
            Der Koeffizient muss \({@aa@}\) sein, um den Term \({@aa@}x^2\) zu erzeugen, wenn die linearen Faktoren zusammen multipliziert werden.
        </span>
    </li>
</ol>
<p>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>dd: rand(3)+2;
aa: rand(3)+2;
bb: rand(10)+5;

d1: aa*dd + bb;
d2: bb*dd;
pol: aa*x^2 +d1*x +d2;
ta: 1;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[{@pol@} = (x + {@dd@}) \times \hbox{ (a polynomial)},\]
degree = {@ta@}, coeff = {@aa@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>aa</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>aa</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2099408468</deployedseed>
        <deployedseed>1727824709</deployedseed>
        <deployedseed>1843517274</deployedseed>
        <deployedseed>510329735</deployedseed>
        <deployedseed>2047281865</deployedseed>
        <deployedseed>1749471452</deployedseed>
        <deployedseed>1333538751</deployedseed>
        <deployedseed>1732879817</deployedseed>
        <deployedseed>174323735</deployedseed>
        <deployedseed>319153438</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>aa</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1, simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>aa</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(aa+1, simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
