<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>1.1.4.T4 BODMAS</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Trening
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Evaluate the expression below
    </span>
    <span class="multilang" lang="no">
        Evaluere uttrykket nedenfor
    </span>
</p>
<p>
    <span class="multilang" lang="en">
    </span>
    <span class="multilang" lang="no">
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\dfrac{{@a1@}-{@a2@}}{9-{@a3@}} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </span>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The dividing line implies that brackets are present around the numerator and around the denominator. Thus
    </span>
    <span class="multilang" lang="no">
        Skillelinjen inneb&#230;rer at parentes er til stede rundt telleren og rundt nevneren.Dermed
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\qquad \dfrac{{@a1@}-{@a2@}}{9-{@a3@}} = \dfrac{({@a1@}-{@a2@})}{(9-{@a3@})}\)
    </span>
    <span class="multilang" lang="no">
        \(\qquad \dfrac{{@a1@}-{@a2@}}{9-{@a3@}} = \dfrac{({@a1@}-{@a2@})}{(9-{@a3@})}\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We then perform the subtraction within the brackets to give the final answer
    </span>
    <span class="multilang" lang="no">
        Vi utf&#248;rer deretter subtraksjonen i parentesene for &#229; gi det endelige svaret
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\qquad \dfrac{({@a1@}-{@a2@})}{(9-{@a3@})} = \dfrac{{@a1-a2@}}{{@9-a3@}}\)
    </span>
    <span class="multilang" lang="no">
        \(\qquad \dfrac{({@a1@}-{@a2@})}{(9-{@a3@})} = \dfrac{{@a1-a2@}}{{@9-a3@}}\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand([4,5])
a2: rand([1,2,3])
a3: rand([2,4])

ta: (a1-a2)/(9-a3)</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote>
            <text>\(\dfrac{{@a1@}-{@a2@}}{9-{@a3@}} = \dfrac{{@a1-a2@}}{{@9-a3@}}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Correct answer, well done.
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partially correct.
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Incorrect answer.
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>+,-,*</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumRelative</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>0</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>265763055</deployedseed>
        <deployedseed>893931160</deployedseed>
        <deployedseed>51965423</deployedseed>
        <deployedseed>2023064728</deployedseed>
        <deployedseed>290737095</deployedseed>
        <deployedseed>74012104</deployedseed>
        <deployedseed>164005195</deployedseed>
        <deployedseed>2049378770</deployedseed>
        <deployedseed>720197030</deployedseed>
        <deployedseed>1648862304</deployedseed>
        <deployedseed>253164936</deployedseed>
        <deployedseed>291804035</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>(a1-a2)*a3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>