<quiz>
    <question type="stack">
        <name>
            <text>8.2.2.2 Exercise 2x2 linear system infinitely many solutions</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the following linear system of equations:
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie das folgende lineare Gleichungssystem:
    </span>
    <span class="math display">
        \[\begin{matrix}
                {@eqq1@} \\
                {@eqq2@} \\
\end{matrix}\]
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Consider the matrix form \(AX = B\) of the system and identify \(A\)
            </span>
            <span class="multilang" lang="de">
                Betrachten Sie die Matrixform \(AX = B\) des Systems und identifizieren Sie \(A\)
            </span>
        </p>
    </li>
    <div>
        \(A=\) [[input:ans1]] [[validation:ans1]]
    </div>
    <div>
        <br/>
    </div>
    <li>
        <p>
            <span class="multilang" lang="en">
                Find the determinant of \(A\):
            </span>
            <span class="multilang" lang="de">
                Finden Sie die Determinante von \(A\):
            </span>
        </p>
        <p>
            <span class="math inline">
                \(|A| =\)
            </span>
            [[input:ans2]] [[validation:ans2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Determine if the system has no solution, a unique solution, or infinitely many solutions.
            </span>
            <span class="multilang" lang="de">
                Bestimmen Sie, ob das System keine Lösung, eine eindeutige Lösung oder unendlich viele Lösungen hat.
            </span>
        </p>
    </li>
    <p>
        [[input:ans3]] [[validation:ans3]]
    </p>
    <li>
        <p>
            <span class="multilang" lang="en">
                If it has a unique solution or infinitely many solutions, these are:
            </span>
            <span class="multilang" lang="de">
                Wenn es eine eindeutige Lösung oder unendlich viele Lösungen gibt, sind dies:
            </span>
        </p>
    </li>
    <p>
        \(x =\) [[input:ans4]] [[validation:ans4]]
    </p>
    <p>
        \(y =\) [[input:ans5]] [[validation:ans5]]
    </p>
    <p>
        <span class="multilang" lang="en">
            Give all the solutions if there is more than one. If there are no solutions, fill the gaps for \(x \) and \(y \) with the word
            <b>
                null
            </b>
            .
        </span>
        <span class="multilang" lang="de">
            Geben Sie alle Lösungen an, wenn es mehr als eine gibt. Wenn es keine Lösungen gibt, füllen Sie die Lücken für \(x\) und \(y\) mit dem Wort
            <b>
                null
            </b>
            .
        </span>
    </p>
</ol>
[[feedback:prt1]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                In matrix form \(AX = B\), the matrix \(A \) is \[ A = {@AA@}.\]
            </span>
            <span class="multilang" lang="de">
                In der Matrixform \(AX = B\) ist die Matrix \(A \) \[ A = {@AA@}.\]
            </span>
        </p>
    </li>
    <li>
        <p>
            \( \det A = {@determinant(AA)@} \).
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Since the determinant of \(A \) is \( 0 \), we already know that the system has either no solution or infinitely many solutions. Note that if we multiply the first equation by \({@kk@} \) we obtain \[ {@kk@} \cdot ({@lhs(eqq1)@}) = {@kk@} \cdot
                                                ({@rhs(eqq1)@}) \Leftrightarrow {@expand(kk * lhs(eqq1))@} = {@kk * rhs(eqq1)@} \] which is the second equation. Therefore, the system has infinitely many solutions.
            </span>
            <span class="multilang" lang="de">
                Da die Determinante von \(A \) \( 0 \) ist, wissen wir bereits, dass das System entweder keine Lösung oder unendlich viele Lösungen hat. Beachten Sie, dass wir, wenn wir die erste Gleichung mit \({@kk@} \) multiplizieren, \[ {@kk@} \cdot ({@lhs(eqq1)@}) = {@kk@} \cdot 
                                                ({@rhs(eqq1)@}) \Leftrightarrow {@expand(kk * lhs(eqq1))@} = {@kk * rhs(eqq1)@} \] erhalten, was die zweite Gleichung ist. Daher hat das System unendlich viele Lösungen.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Since both equations are equivalent, we can ignore one of them, say the second one. From the first equation, if we set \(y = t \), where \( t \) can be any real number, we obtain all the solutions to the system \[ x = {@xx1@}, \qquad y = {@xx2@}.
                                                \]
            </span>
            <span class="multilang" lang="de">
                Da beide Gleichungen äquivalent sind, können wir eine davon ignorieren, beispielsweise die zweite. Wenn wir aus der ersten Gleichung \(y = t \) setzen, wobei \( t \) eine beliebige reelle Zahl sein kann, erhalten wir alle Lösungen des Systems \[ x = {@xx1@}, \qquad y = {@xx2@}.
                                                \]
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/* Set an order in the variables */
ordergreat(x, y);
/* Generate non-zero random numbers. They will be the coefficients of the first equation */
aa : rand(6) * 2 - 5 ;
bb : rand(6) * 2 - 5  ;
cc : rand(6) * 2 - 5  ;
/* Generate a random number different than 0, 1. It will be the opposite coefficient by which we multiply the second equation*/
kk : rand(4) * 3 - 4;
/* Define equations */
eqq1 : aa * x + bb * y = cc;
eqq2: expand(kk * eqq1);

AA : matrix([aa, bb], [kk * aa, kk * bb]);
BB : matrix([cc], [kk * cc]);
detAA : determinant(AA);

/* Variables for dropdown list */
ta3:[["No solution",false],["One unique solution",false],["Infinitely many solutions",true]];
ta3:random_permutation(ta3);
ta3c: "Infinitely many solutions";

/* Determine solution to the system */
eq1 : aa * x + bb * t = cc;
eq2 : kk * eq1;
ta4 : algsys([eq1, eq2], [x]);
xx1 : rhs(ta4[1][1]);
xx2 : t;
XX : matrix([xx1], [xx2]);

nullMat : matrix([null], [null]);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[\begin{matrix} {@eqq1@} \\ {@eqq2@} \\ \end{matrix}\]
\( A = {@AA@}, \quad |A| = {@detAA@},\) {@ta3c@}, \( x={@xx1@}, y = {@xx2@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>AA</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>detAA</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>xx1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>xx2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* Create the vector containing the answers */
vecAns : matrix([ans4], [ans5]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>AA</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.2</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (a) Matrix \(A\) correctly identified.
    </span>
    <span class="multilang" lang="de">
        (a) Matrix \(A\) korrekt identifiziert.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (a) Be careful, matrix \(A\) was not introduced correctly. Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="de">
        (a) Seien Sie vorsichtig, Matrix \(A\) wurde nicht korrekt eingeführt. Leider hat sich durch diesen Fehler der Schwierigkeitsgrad der Aufgabe verändert und es können keine Punkte mehr vergeben werden.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>detAA</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2</truescore>
                <truepenalty></truepenalty>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (b) Determinant of \(A\) correctly computed.
    </span>
    <span class="multilang" lang="de">
        (b) Determinante von \(A\) korrekt berechnet.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (b) Be careful, the determinant of \(A\) was not computed correctly. Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="de">
        (b) Seien Sie vorsichtig, die Determinante von \(A\) wurde nicht korrekt berechnet. Leider hat sich durch diesen Fehler der Schwierigkeitsgrad der Aufgabe geändert und es können keine Punkte mehr vergeben werden.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3c</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2</truescore>
                <truepenalty></truepenalty>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (c) The dimension of the solution space was identified correctly.
    </span>
    <span class="multilang" lang="de">
        (c) Die Dimension des Lösungsraums wurde korrekt identifiziert.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (c) The dimension of the solution space was not identified correctly. What happens if we multiply the first equation by \( {@kk@} \)? Unfortunately, this error changed the difficulty of the problem and no more points can be awarded.
    </span>
    <span class="multilang" lang="de">
        (c) Die Dimension des Lösungsraums wurde nicht korrekt identifiziert. Was passiert, wenn wir die erste Gleichung mit \( {@kk@} \) multiplizieren? Leider hat sich durch diesen Fehler der Schwierigkeitsgrad der Aufgabe geändert und es können keine Punkte mehr vergeben werden.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>AA.vecAns</sans>
                <tans>BB</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2</truescore>
                <truepenalty></truepenalty>
                <truenextnode>4</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <answertest>AlgEquiv</answertest>
                <sans>length(listofvars(ans4))</sans>
                <tans>1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (d) You've identified all the solutions correctly
    </span>
    <span class="multilang" lang="de">
        (d) Sie haben alle Lösungen richtig identifiziert
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (d) Your answer satisfies the equation, but is not fully general. You need to add in a constant to represent the general solution.
    </span>
    <span class="multilang" lang="de">
        (d) Ihre Antwort erfüllt die Gleichung, ist aber nicht ganz allgemein. Sie müssen eine Konstante hinzufügen, um die allgemeine Lösung darzustellen.
    </span>
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <answertest>AlgEquiv</answertest>
                <sans>xx1</sans>
                <tans>null</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-6-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        (d) Your answer does not solve the equation \[ {@AA@}{@vecAns@} = {@BB@}.\]
    </span>
    <span class="multilang" lang="de">
        (d) Ihre Antwort löst die Gleichung  \[ {@AA@}{@vecAns@} = {@BB@}\] nicht.
    </span>
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1599695625</deployedseed>
        <deployedseed>2129295706</deployedseed>
        <deployedseed>671571537</deployedseed>
        <deployedseed>872771527</deployedseed>
        <deployedseed>896876049</deployedseed>
        <deployedseed>1009348279</deployedseed>
        <deployedseed>1312654757</deployedseed>
        <deployedseed>131991284</deployedseed>
        <deployedseed>768028467</deployedseed>
        <deployedseed>2073326638</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>AA</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>detAA</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3c</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>xx1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>xx2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-5-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(2*AA,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(detAA+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["No solution"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>null</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>null</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>AA</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>detAA</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["No solution"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>null</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>null</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.4000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
