<quiz>
    <question type="stack">
        <name>
            <text>9.5.1.4 Direction ratios and cosines, Exercise</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Points \(P\) and \(Q\) have position vectors \({@stackvector(a)@} ={@A1@}\), and \({@stackvector(b)@} = {@B1@}\) respectively. Find
    </span>
    <span class="multilang" lang="de">
        Die Punkte \(P\) und \(Q\) haben die Ortsvektoren \({@stackvector(a)@} ={@A1@}\) und \({@stackvector(b)@} = {@B1@}\). Finden Sie
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                The direction ratio of \(\overrightarrow{PQ}\):
            </span>
            <span class="multilang" lang="de">
                Das Richtungsverhältnis von \(\overrightarrow{PQ}\) ist:
            </span>
            [[input:dr1]] [[validation:dr1]] \(:\) [[input:dr2]] [[validation:dr2]] [[feedback:DR]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                The direction cosines of \(\overrightarrow{PQ}\), \((\ell,m)=\)
            </span>
            <span class="multilang" lang="de">
                Der Richtungskosinus von \(\overrightarrow{PQ}\), ist \((\ell,m)=\)
            </span>
            [[input:dc]]. [[validation:dc]] [[feedback:DC]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The vector \(\overrightarrow{PQ}\), where \(P\) and \(Q\) have position vectors \(\underline{a} = {@A1@}\) and \(\underline{b} = {@B1@}\) respectively, is given by
    </span>
    <span class="multilang" lang="de">
        Der Vektor \(\overrightarrow{PQ}\), wobei \(P\) und \(Q\) die Ortsvektoren \(\underline{a} = {@ A1@}\) bzw. \(\underline{b} = {@B1@}\) haben, ist gegeben durch
    </span>
</p>
\[\overrightarrow{PQ} = \underline{b} - \underline{a} = {@R1@}.\]
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                The direction ratios of \(\overrightarrow{PQ}\) are \(\displaystyle{@r1[1,1]@}: \displaystyle{@r1[2,1]@}\).
            </span>
            <span class="multilang" lang="de">
                Die Richtungsverhältnisse von \(\overrightarrow{PQ}\) sind \(\displaystyle{@r1[1,1]@}: \displaystyle{@r1[2, 1]@}\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                To find the direction cosines, we must first find the magnitude of the vector. In this case,
            </span>
            <span class="multilang" lang="de">
                Um die Richtungskosinusse zu finden, müssen wir zuerst die Größe des Vektors finden. In diesem Fall gilt
            </span>
        </p>
        \[\lvert \overrightarrow{PQ} \rvert = \sqrt{({@r1[1][1]@})^2 + ({@r1[2][1]@})^2} = {@magr1@}.\]
        <p>
            <span class="multilang" lang="en">
                Then, the direction cosines of \(\overrightarrow{PQ}\) are
            </span>
            <span class="multilang" lang="de">
                Dann sind die Richtungskosinuswerte von \(\overrightarrow{PQ}\)
            </span>
        </p>
        \[(\ell,m) = \displaystyle{@tadc@}.\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
a1: matrix([7],[3]) - rand(matrix([7],[3]));
b1: matrix([7],[8]) - rand(matrix([7],[2]));
r1:  b1-a1;
magr1: sqrt(dotproduct(r1,r1));
u1: matrix([i],[j]);
A1: texboldatoms(dotproduct(a1,u1)+0*k);
B1: texboldatoms(dotproduct(b1,u1)+0*k);
R1: texboldatoms(dotproduct(r1,u1));
l1: (r1[1,1])/(magr1);
m1: (r1[2,1])/(magr1);
tadc: matrix([l1,m1]);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>The direction ratios for the vector \(\underline{r} = {@R1@}\) are \({@r1[1,1]@}\,:\,{@r1[2,1]@}\), and its direction cosines are \(\ell = {@l1@}\) and \(m = {@m1@}\).</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Correct answer, well done.
    </span>
    <span class="multilang" lang="de">
        Richtige Antwort, gut gemacht.
    </span>
</p>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is partially correct.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist teilweise richtig.
    </span>
</p>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Incorrect answer.
    </span>
    <span class="multilang" lang="de">
        Falsche Antwort.
    </span>
</p>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>symi</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>dc</name>
            <type>matrix</type>
            <tans>tadc</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>dr1</name>
            <type>algebraic</type>
            <tans>r1[1,1]</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>dr2</name>
            <type>algebraic</type>
            <tans>r1[2,1]</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>DC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>dc1</sans>
                <tans>l1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>VEC-DC-X_COR</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-DC-X_INCOR</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer for the direction cosine in the \(x\)-direction appears to be incorrect.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort für den Richtungskosinus in der \(x\)-Richtung scheint falsch zu sein.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.
    </span>
    <span class="multilang" lang="de">
        Merken Sie sich die Formel zur Berechnung des Richtungskosinus (zeichnen Sie eine Skizze, wenn das hilft!), und versuchen Sie es erneut.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>dc2</sans>
                <tans>m1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-DC-Y_COR</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-DC-Y_INCOR</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer for the direction cosine in the \(y\)-direction appears to be incorrect.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort für den Richtungskosinus in der \(y\)-Richtung scheint falsch zu sein.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Remember the formula for calculating the direction cosine (draw a diagram if this helps!), and try again.
    </span>
    <span class="multilang" lang="de">
        Merken Sie sich die Formel zur Berechnung des Richtungskosinus (zeichnen Sie eine Skizze, wenn das hilft!), und versuchen Sie es erneut.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>dc</sans>
                <tans>tadc</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-DC-CORRECT</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>0</falsenextnode>
                <falseanswernote>VEC-DC-Z_INCOR</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>DR</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>student_r: matrix([dr1],[dr2]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>student_r</sans>
                <tans>r1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-DR-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-DR-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Remember how we obtain the direction ratios for a three-dimensional vector from its expression in terms of \(\mathbf{i}\) and \(\mathbf{j}\) components, and try again.
    </span>
    <span class="multilang" lang="de">
        Denken Sie daran, wie wir die Richtungsverhältnisse für einen dreidimensionalen Vektor aus seinem Ausdruck in Form von \(\mathbf{i}\) und \(\mathbf{j)}\) Komponenten erhalten und versuchen Sie es erneut.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>405915637</deployedseed>
        <deployedseed>1199752509</deployedseed>
        <deployedseed>1931035035</deployedseed>
        <deployedseed>111862704</deployedseed>
        <deployedseed>866605738</deployedseed>
        <deployedseed>1701042024</deployedseed>
        <deployedseed>2044316482</deployedseed>
        <deployedseed>1413350108</deployedseed>
        <deployedseed>1309609691</deployedseed>
        <deployedseed>870524120</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>dc</name>
                <value>tadc</value>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]</value>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DC-CORRECT</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>dc</name>
                <value>ev(2*tadc,simp)</value>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>ev(r1[1,1]+1,simp)</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>ev(r1[2,1]+1,simp)</value>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>VEC-DC-X_INCOR</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>VEC-DR-FALSE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>dc</name>
                <value>matrix([r1[1,1],r1[2,1]])</value>
            </testinput>
            <testinput>
                <name>dr1</name>
                <value>r1[1,1]</value>
            </testinput>
            <testinput>
                <name>dr2</name>
                <value>r1[2,1]</value>
            </testinput>
            <expected>
                <name>DC</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>VEC-DC-X_INCOR</expectedanswernote>
            </expected>
            <expected>
                <name>DR</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-DR-TRUE</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
