<quiz>
    <question type="stack">
        <name>
            <text>9.5.3.8 Vector and cartesian equation of a line, Exercise</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Let the point \(A\) have position vector \({@stackvector(a)@} = {@A1@}\), and the point \(B\) have position vector \({@stackvector(b)@} = {@B1@}\).
    </span>
    <span class="multilang" lang="de">
        Der Punkt \(A\) habe den Ortsvektor \({@stackvector(a)@} = {@A1@}\) und der Punkt \(B\) habe den Ortsvektor \({@stackvector(b)@} = {@B1@}\).
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Using this information, give the
                <i>
                    canonical
                </i>
                vector equation of the line passing through \(A\) and \(B\). Input your answer, in column vector notation, in the boxes provided below. The line through \(A\) and \(B\) is given by the following
                                        equation, where \(t\) is a parameter which may take any real number as its value.
            </span>
            <span class="multilang" lang="de">
                Geben Sie anhand dieser Informationen die
                <i>
                    kanonische
                </i>
                Vektorgleichung der Geraden an, die durch \(A\) und \(B\) verläuft. Geben Sie Ihre Antwort in Spaltenvektorschreibweise in die unten bereitgestellten Felder ein. Die Gerade durch \(A\) und \(B\) ist durch die folgende
                                        Gleichung gegeben, wobei \(t\) ein Parameter ist, der jede reelle Zahl als Wert annehmen kann.
            </span>
        </p>
        <div>
            \(\underline{r} = \) [[input:ans1]] [[validation:ans1]] \({}+t\ \) [[input:ans2]] [[validation:ans2]]
        </div>
        [[feedback:VecEq]]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Now give the equation of the line in its Cartesian form. Enter your answers as fractions, e.g.
                <code>
                    (k-5)/18
                </code>
                .
            </span>
            <span class="multilang" lang="de">
                Geben Sie nun die Gleichung der Geraden in ihrer kartesischen Form an. Geben Sie Ihre Antworten als Brüche ein, z.B.
                <code>
                    (k-5)/18
                </code>
                .
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                The Cartesian equation of the line through \(A\) and \(B\) is:
            </span>
            <span class="multilang" lang="de">
                Die kartesische Gleichung der Geraden durch \(A\) und \(B\) lautet:
            </span>
        </p>
        <div>
            [[input:ans3]] [[validation:ans3]] \({}={}\) [[input:ans4]] [[validation:ans4]] \({}={}\) [[input:ans5]] [[validation:ans5]]
        </div>
        [[feedback:CartEq]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                The
                <i>
                    canonical
                </i>
                vector equation of the line passing through \(A\) and \(B\) is given by the formula
            </span>
            <span class="multilang" lang="de">
                Die
                <i>
                    kanonische
                </i>
                Vektorgleichung der Geraden, die durch \(A\) und \(B\) verläuft, wird durch die folgende Formel gegeben
            </span>
        </p>
        \[\underline{r} = \underline{a} + t(\underline{b} - \underline{a}).\]
        <p>
            <span class="multilang" lang="en">
                This is the
                <i>
                    canonical
                </i>
                form because we have \(\underline{r}(t = 0) = \underline{a}\) and \(\underline{r}(t = 1) = \underline{b}\).
            </span>
            <span class="multilang" lang="de">
                Dies ist die
                <i>
                    kanonische
                </i>
                Form, weil wir \(\underline{r}(t = 0) = \underline{a}\) und \(\underline{r}(t = 1) = \underline{b}\) haben.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                In our case, the values are
            </span>
            <span class="multilang" lang="de">
                In unserem Fall sind die Werte
            </span>
        </p>
        \[\underline{r} = {@a1@} + t\, {@BMA@}.\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Recall that the Cartesian equation of a line given by \(\underline{r} = \underline{a} + t(\underline{b} - \underline{a})\) is given by
            </span>
            <span class="multilang" lang="de">
                Denken Sie daran, dass die kartesische Gleichung einer Geraden, die durch \(\underline{r} = \underline{a} + t(\underline{b} - \underline{a} )\) gegeben ist, durch
            </span>
        </p>
        \[\frac{x - a_1}{b_1 - a_1} = \frac{y - a_2}{b_2 - a_2} = \frac{z - a_3}{b_3 - a_3},\]
        <p>
            <span class="multilang" lang="en">
                where
            </span>
            <span class="multilang" lang="de">
                gegeben ist, wobei
            </span>
        </p>
        <span class="multilang" lang="en">
            \[\underline{a} = \begin{bmatrix} a_1 \\ a_2 \\ a_3 \end{bmatrix} \quad \text{and} \quad \underline{b} = \begin{bmatrix} b_1 \\ b_2 \\b_3 \end{bmatrix}.\]
        </span>
        <span class="multilang" lang="de">
            \[\underline{a} = \begin{bmatrix} a_1 \\ a_2 \\ a_3 \end{bmatrix} \quad \text{und} \quad \underline{b} = \begin{bmatrix} b_1 \\ b_2 \\b_3 \end{bmatrix}.\]
        </span>
        <p>
            <span class="multilang" lang="en">
                Applying this formula in our case gives the equation:
            </span>
            <span class="multilang" lang="de">
                Die Anwendung dieser Formel ergibt in unserem Fall die Gleichung:
            </span>
        </p>
        \[{@canonical_mat[1,1]@} = {@canonical_mat[2,1]@} = {@canonical_mat[3,1]@}\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
a1: matrix([16],[16],[16]) - rand(matrix([33],[33],[33]));
b1: matrix(
  [rand_with_prohib(-16, 16, [a1[1,1]])],
  [rand_with_prohib(-16, 16, [a1[2,1]])],
  [rand_with_prohib(-16, 16, [a1[3,1]])]);
u1: matrix([i],[j],[k]); 
A1: texboldatoms(dotproduct(a1,u1));
B1: texboldatoms(dotproduct(b1,u1));

BMA: b1 - a1;
ta11: a1;
ta12: BMA;
canonical_mat: matrix([(x-a1[1,1])/(BMA[1,1])],[(y-a1[2,1])/(BMA[2,1])],[(z-a1[3,1])/(BMA[3,1])]);
ta21: canonical_mat[1,1];
ta22: canonical_mat[2,1];
ta23: canonical_mat[3,1];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>The vector equation of the line through the points {@A1@} and {@B1@} is
\[\mathbf{r} = {@a1@} + t\,{@BMA@},\]
and the Cartesian form of the equation is
\[{@ev((x - a1[1,1])/BMA[1,1],simp)@} = {@ev((y - a1[2,1])/BMA[2,1],simp)@} = {@ev((z - a1[3,1])/BMA[3,1],simp)@}.\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Correct answer, well done.
    </span>
    <span class="multilang" lang="de">
        Richtige Antwort, gut gemacht.
    </span>
</p>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is partially correct.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist teilweise richtig.
    </span>
</p>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Incorrect answer.
    </span>
    <span class="multilang" lang="de">
        Falsche Antwort.
    </span>
</p>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>symi</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta11</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>matrix</type>
            <tans>ta12</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta21</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta22</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>ta23</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>CartEq</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>answer_mat: matrix([ans3],[ans4],[ans5]);
ECFBMA: ans2;
ECF_mat: matrix([(x-ans1[1,1])/(ECFBMA[1,1])],[(y-ans1[2,1])/(ECFBMA[2,1])],[(z-ans1[3,1])/(ECFBMA[3,1])]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>answer_mat</sans>
                <tans>canonical_mat</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-CARTEQ-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>VEC-CARTEQ-FALSE-0</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>answer_mat</sans>
                <tans>ECF_mat</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-CARTEQ-FALSE_BUT_ECF</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Although your answers for this part of the question do give the Cartesian equation of the line from part (a), it appears that the equation of the line you gave in part (a) is invalid – either incorrect outright, or not in canonical form.
    </span>
    <span class="multilang" lang="de">
        Obwohl Ihre Antworten für diesen Teil der Frage die kartesische Gleichung der Geraden aus Teil (a) angeben, scheint es, dass die Gleichung der Geraden, die Sie in Teil (a) angegeben haben ungültig ist – entweder völlig falsch oder nicht in kanonischer Form.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please correct your answers to part (a), then amend your Cartesian equation of the line, and input the correct values in the boxes for part (b).
    </span>
    <span class="multilang" lang="de">
        Bitte korrigieren Sie Ihre Antworten zu Teil (a), ändern Sie dann Ihre kartesische Gleichung der Geraden und geben Sie die richtigen Werte in die Felder für Teil (b) ein.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-CARTEQ-FALSE-1</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        It appears your answer for the Cartesian equation of the line through \(A\) and \(B\) is incorrect.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort für die kartesische Gleichung der Geraden durch \(A\) und \(B\) scheint falsch zu sein.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Refer back to how we can move from the vector equation of a line to the Cartesian equation of a line, and try again. Make sure your answer to part (a) is correct and in the canonical form, too!
    </span>
    <span class="multilang" lang="de">
        Sehen Sie sich noch einmal an, wie wir von der Vektorgleichung einer Geraden zur kartesischen Gleichung einer Geraden übergehen können, und versuchen Sie es erneut. Stellen Sie sicher, dass Ihre Antwort für Teil (a) korrekt und auch in der kanonischen Form ist!
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>VecEq</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta11</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>VEC-VECEQ-TRUE-0</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-VECEQ-FALSE-0</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The first vector you have entered as part of the vector equation of the line is invalid – it is either incorrect outright, or not the canonical choice of vector. Please try again.
    </span>
    <span class="multilang" lang="de">
        Der erste Vektor, den Sie als Teil der Vektorgleichung der Geraden eingegeben haben, ist ungültig – entweder ist er völlig falsch oder entspricht nicht der kanonischen Wahl des Vektors. Bitte versuchen Sie es erneut.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta12</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>VEC-VECEQ-TRUE-1</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>VEC-VECEQ-FALSE-1</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The second vector you have entered as part of the vector equation of the line is invalid – it is either incorrect outright, or not the canonical choice of vector. Please try again.
    </span>
    <span class="multilang" lang="de">
        Der zweite Vektor, den Sie als Teil der Vektorgleichung der Geraden eingegeben haben, ist ungültig – entweder ist er völlig falsch oder entspricht nicht der kanonischen Wahl des Vektors. Bitte versuchen Sie es erneut.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1692160676</deployedseed>
        <deployedseed>274327739</deployedseed>
        <deployedseed>13451817</deployedseed>
        <deployedseed>774589802</deployedseed>
        <deployedseed>1248931864</deployedseed>
        <deployedseed>1572694177</deployedseed>
        <deployedseed>641284927</deployedseed>
        <deployedseed>1286376242</deployedseed>
        <deployedseed>994504584</deployedseed>
        <deployedseed>1826193064</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta11</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta12</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta21</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta22</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta23</value>
            </testinput>
            <expected>
                <name>CartEq</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-CARTEQ-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>VecEq</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-VECEQ-TRUE-1</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>b1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta12</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta21</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta22</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta23</value>
            </testinput>
            <expected>
                <name>CartEq</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-CARTEQ-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>VecEq</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-VECEQ-FALSE-0</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta11</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta12</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>0</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>0</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>0</value>
            </testinput>
            <expected>
                <name>CartEq</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>VEC-CARTEQ-FALSE-1</expectedanswernote>
            </expected>
            <expected>
                <name>VecEq</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-VECEQ-TRUE-1</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta11</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(-ta12,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta21</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta22</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta23</value>
            </testinput>
            <expected>
                <name>CartEq</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-CARTEQ-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>VecEq</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>VEC-VECEQ-FALSE-1</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
