<quiz>
    <question type="stack">
        <name>
            <text>9.5.2.3 The angle between a vector and the axes, Exercise</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the vector \(\overrightarrow{OP}\), where \(P\) is the point with coordinates \({@p1@}\).
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie den Vektor \(\overrightarrow{OP}\), wobei \(P\) der Punkt mit den Koordinaten \({@p1@}\) ist.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Find the direction ratios \(l : m : n\) of \(\overrightarrow{OP}\).
            </span>
            <span class="multilang" lang="de">
                Finden Sie die Richtungsverhältnisse \(l : m : n\) von \(\overrightarrow{OP}\).
            </span>
        </p>
        <div>
            \((l,m,n) = \) [[input:ansRat]] [[validation:ansRat]] [[feedback:ratio]]
        </div>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Now find the cosines of the angles between the vector \(\overrightarrow{OP}\) and the \(x\),\(y\), and \(z\)-axes. Let these angles be \(\theta_x\), \(\theta_y\) and \(\theta_z\), respectively. Give exact answers, using square roots.
            </span>
            <span class="multilang" lang="de">
                Bestimmen Sie nun die Kosinuswerte der Winkel zwischen dem Vektor \(\overrightarrow{OP}\) und den \(x\),\(y\) und \(z \)-Achsen. Diese Winkel seien jeweils \(\theta_x\), \(\theta_y\) und \(\theta_z\). Geben Sie genaue Antworten, indem Sie Quadratwurzeln verwenden.
            </span>
        </p>
        <p>
            \(\cos \theta_x =\) [[input:ans4x]] [[validation:ans4x]] [[feedback:Part3ax]]
        </p>
        <p>
            \(\cos \theta_y =\) [[input:ans4y]] [[validation:ans4y]] [[feedback:Part3ay]]
        </p>
        <p>
            \(\cos \theta_z =\) [[input:ans4z]] [[validation:ans4z]] [[feedback:Part3az]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Lastly, find the angles between the vector \(\overrightarrow{OP}\) and the \(x\),\(y\), and \(z\)-axes. Give your answer in degrees, to three significant figures.
            </span>
            <span class="multilang" lang="de">
                Bestimmen Sie abschließend die Winkel zwischen dem Vektor \(\overrightarrow{OP}\) und den \(x\),\(y\) und \(z\)-Achsen. Geben Sie Ihre Antwort in Grad mit drei signifikanten Stellen an.
            </span>
        </p>
        <p>
            \( \theta_x =\) [[input:ans5x]] [[validation:ans5x]]
            <span class="multilang" lang="en">
                degrees.
            </span>
            <span class="multilang" lang="de">
                Grad.
            </span>
            [[feedback:Part3bx]]
        </p>
        <p>
            \( \theta_y =\) [[input:ans5y]] [[validation:ans5y]]
            <span class="multilang" lang="en">
                degrees.
            </span>
            <span class="multilang" lang="de">
                Grad.
            </span>
            [[feedback:Part3by]]
        </p>
        <p>
            \( \theta_z =\) [[input:ans5z]] [[validation:ans5z]]
            <span class="multilang" lang="en">
                degrees.
            </span>
            <span class="multilang" lang="de">
                Grad.
            </span>
            [[feedback:Part3bz]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We can just read off the direction rations \(l : m : n\) from the vector.
            </span>
            <span class="multilang" lang="de">
                Wir können einfach die Richtungsverhältnisse \(l : m : n\) aus dem Vektor ablesen.
            </span>
        </p>
        \[(l,m,n) = \begin{bmatrix} {@p1[1,1]@} &amp; {@p1[2,1]@} &amp; {@p1[3,1]@} \end{bmatrix}\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Next, we are asked for the cosines of the angles between \(\overrightarrow{OP}\) and the \(x\), \(y\) and \(z\)-axes. These are simply the direction cosines. In order to find them, we must first compute the magnitude of \(\overrightarrow{OP}\).
            </span>
            <span class="multilang" lang="de">
                Als nächstes werden wir nach den Kosinuswerten der Winkel zwischen \(\overrightarrow{OP}\) und den \(x\), \(y\) und \( z\)-Achsen gefragt. Dies sind einfach die Richtungskosinusse. Um sie zu bestimmen, müssen wir zunächst die Größe von \(\overrightarrow{OP}\) berechnen.
            </span>
        </p>
        \[\lvert\overrightarrow{OP}\rvert = {@mag_p1@}\]
        <p>
            <span class="multilang" lang="en">
                Then we have
            </span>
            <span class="multilang" lang="de">
                Dann erhalten wir
            </span>
        </p>
        \[\begin{aligned} \cos \theta_x &amp;= {@cos_thetax@} \\ \cos \theta_y &amp;= {@cos_thetay@} \\ \cos \theta_z &amp;= {@cos_thetaz@} \end{aligned}\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Lastly, we are asked to find what each of these angles is. We can do this using the arccosine function in our calculator. Remember we must write our answers in degress and to three significant figures.
            </span>
            <span class="multilang" lang="de">
                Zum Schluss sollen wir herausfinden, wie groß jeder dieser Winkel ist. Dazu können wir die Funktion Arkosinus in unserem Taschenrechner verwenden. Denken Sie daran, dass wir unsere Antworten in Grad und auf drei signifikante Zahlen runden müssen.
            </span>
        </p>
        \[\begin{aligned} \theta_x &amp;= {@r_thetax_d@} \\ \theta_y &amp;= {@r_thetay_d@} \\ \theta_z &amp;= {@r_thetaz_d@} \end{aligned}\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>5.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
 /*Sets up relevant vectors to be used in the question.*/
 p1: matrix([9],[9],[8]) - rand(matrix([18],[18],[7]));
drRowVec: transpose(p1);
 u1: matrix([i],[j],[k]);
 /*Presents the vectors p in a suitable form to be displayed in the question text / note.*/ 
 P1: texboldatoms(dotproduct(p1,u1)); 
 /*Computes the vectors' respective magnitudes.*/
 mag_p1: sqrt(dotproduct(p1,p1));
 /*Evaluates cos theta in terms of the scalar product and the magnitudes.*/
 cos_thetax: p1[1,1]/(mag_p1);
 cos_thetay: p1[2,1]/(mag_p1);
 cos_thetaz: p1[3,1]/(mag_p1);
 /*Finds the value of theta (truncated to three sig fig) in both radians and degrees.*/
 r_thetax_r: significantfigures(ev(acos(cos_thetax),float),3);
 r_thetax_d: significantfigures(ev(acos(cos_thetax)*(180/pi),float),3);
 r_thetay_r: significantfigures(ev(acos(cos_thetay),float),3);
 r_thetay_d: significantfigures(ev(acos(cos_thetay)*(180/pi),float),3);
 r_thetaz_r: significantfigures(ev(acos(cos_thetaz),float),3);
 r_thetaz_d: significantfigures(ev(acos(cos_thetaz)*(180/pi),float),3);
 /*Students may possibly sum the components of the vector product as one of their answers. This item may be used as part of the PRT, to give relevant feedback about their error.*/
 /*q1 is basis vector for x axis, should probably include y and z*/
q1: matrix([1],[0],[0]);
 vector_sum:(p1[2,1]*q1[3,1] - p[3,1]*q1[2,1]) + (p1[3,1]*q1[1,1] - p1[1,1]*q1[3,1]) + (p1[1,1]*q1[2,1] - p1[2,1]*q1[1,1]);
 /*entri that wont occur for test of ECF*/
ecfTest: matrix([10,20,30]);
ecfTestCosx: 1/sqrt(14);
ecfTestCosy: 2/sqrt(14);
ecfTestCosz: 3/sqrt(14);
 r_ecfthetax_d: significantfigures(ev(acos(ecfTestCosx)*(180/pi),float),3);
 r_ecfthetay_d: significantfigures(ev(acos(ecfTestCosy)*(180/pi),float),3);
 r_ecfthetaz_d: significantfigures(ev(acos(ecfTestCosz)*(180/pi),float),3);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ (\ell,m,n)= {@drRowVec@}\]
\[\cos \theta_x = {@cos_thetax@}, \theta_x = {@r_thetax_d@}\  \text{degrees}\]
\[\cos \theta_y = {@cos_thetay@}, \theta_y = {@r_thetay_d@}\  \text{degrees}\]
\[\cos \theta_z = {@cos_thetaz@}, \theta_z = {@r_thetaz_d@}\  \text{degrees}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Correct answer, well done.
    </span>
    <span class="multilang" lang="de">
        Richtige Antwort, gut gemacht.
    </span>
</p>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is partially correct.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist teilweise richtig.
    </span>
</p>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Incorrect answer.
    </span>
    <span class="multilang" lang="de">
        Falsche Antwort.
    </span>
</p>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>symi</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans4x</name>
            <type>algebraic</type>
            <tans>cos_thetax</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4y</name>
            <type>algebraic</type>
            <tans>cos_thetay</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4z</name>
            <type>algebraic</type>
            <tans>cos_thetaz</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5x</name>
            <type>numerical</type>
            <tans>r_thetax_d</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5y</name>
            <type>numerical</type>
            <tans>r_thetay_d</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5z</name>
            <type>numerical</type>
            <tans>r_thetaz_d</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ansRat</name>
            <type>matrix</type>
            <tans>drRowVec</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>Part3ax</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4x</sans>
                <tans>cos_thetax</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>COSx-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>COSx-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You have incorrectly calculated \(\cos \theta_x\) here.
    </span>
    <span class="multilang" lang="de">
        Sie haben \(\cos \theta_x\) hier falsch berechnet.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Part3ay</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4y</sans>
                <tans>cos_thetay</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>COSy-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>COSy-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You have incorrectly calculated \(\cos \theta_y\) here.
    </span>
    <span class="multilang" lang="de">
        Sie haben \(\cos \theta_y\) hier falsch berechnet.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Part3az</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4z</sans>
                <tans> cos_thetaz</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>COSz-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>COSz-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You have incorrectly calculated \(\cos \theta_z\) here.
    </span>
    <span class="multilang" lang="de">
        Sie haben \(\cos \theta_z\) hier falsch berechnet.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Part3bx</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>s_thetax_deg_ECF: significantfigures(ev(acos(ans4x)*(180/pi),float),3);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5x</sans>
                <tans>r_thetax_d</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGx-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ANGx-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5x</sans>
                <tans>r_thetax_r</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGx-INRADIANS</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Your (rounded) answer appears to have been given in radians. Please enter the appropriate answer in degrees. Remember, only round your answer at the very final stage of the calculation!
    </span>
    <span class="multilang" lang="de">
        Ihre (gerundete) Antwort scheint im Bogenmaß angegeben worden zu sein. Bitte geben Sie die entsprechende Antwort in Grad ein. Denken Sie daran, Ihre Antwort erst im allerletzten Schritt der Berechnung zu runden!
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ANGx-FALSE1</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5x</sans>
                <tans>s_thetax_deg_ECF</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGx-FALSE-TRUEwECF</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You appear to have used the correct formula to find \(\theta_x\) in degrees (and to round it to the appropriate number of significant figures!), but it appears that your value of \(\cos \theta_x\) from earlier is incorrect; this means your answer here
        is not correct, either.
    </span>
    <span class="multilang" lang="de">
        Sie scheinen die richtige Formel verwendet zu haben, um \(\theta_x\) in Grad zu ermitteln (und es auf die entsprechende Anzahl signifikanter Stellen zu runden!), aber es sieht so aus, dass Ihr Wert von \(\cos \theta_x\) von früher falsch ist; das bedeutet Ihre Antwort hier
        ist auch nicht korrekt.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please amend your answer to the previous part of the question, and attempt this part again as well using the correct value of \(\cos \theta_x\).
    </span>
    <span class="multilang" lang="de">
        Ändern Sie bitte Ihre Antwort zum vorherigen Teil der Frage und versuchen Sie auch diesen Teil erneut, indem Sie den richtigen Wert von \(\cos \theta_x\) verwenden.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ANGx-FALSE2</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You do not appear to have entered the correct answer. Perhaps you are using some incorrect units? Perhaps you have not rounded correctly? Perhaps you are using an incorrect formula?
    </span>
    <span class="multilang" lang="de">
        Sie scheinen nicht die richtige Antwort eingegeben zu haben. Vielleicht verwenden Sie falsche Einheiten? Vielleicht haben Sie nicht richtig gerundet? Vielleicht verwenden Sie eine falsche Formel?
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Remind yourself of how to correctly calculate the angle between the two vectors, and please try again.
    </span>
    <span class="multilang" lang="de">
        Erinnern Sie sich daran, wie Sie den Winkel zwischen den beiden Vektoren richtig berechnen, und versuchen Sie es bitte erneut.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Part3by</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>s_thetay_deg_ECF: significantfigures(ev(acos(ans4y)*(180/pi),float),3);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5y</sans>
                <tans>r_thetay_d</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGy-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ANGy-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5y</sans>
                <tans>r_thetay_r</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGy-INRADIANS</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Your (rounded) answer appears to have been given in radians. Please enter the appropriate answer in degrees. Remember, only round your answer at the very final stage of the calculation!
    </span>
    <span class="multilang" lang="de">
        Ihre (gerundete) Antwort scheint im Bogenmaß angegeben worden zu sein. Bitte geben Sie die entsprechende Antwort in Grad ein. Denken Sie daran, Ihre Antwort erst in der allerletzten Phase der Berechnung zu runden!
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ANGy-FALSE1</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5y</sans>
                <tans>s_thetay_deg_ECF</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGy-FALSE-TRUEwECF</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You appear to have used the correct formula to find \(\theta_y\) in degrees (and to round it to the appropriate number of significant figures!), but it appears that your value of \(\cos \theta_y\) from earlier is incorrect; this means your answer here
        is not correct, either.
    </span>
    <span class="multilang" lang="de">
        Sie scheinen die richtige Formel verwendet zu haben, um \(\theta_y\) in Grad zu ermitteln (und es auf die entsprechende Anzahl signifikanter Stellen zu runden!), aber es sieht so aus, dass Ihr Wert von \(\cos \theta_y\) von früher falsch ist; das bedeutet Ihre Antwort hier
        ist auch nicht korrekt.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please amend your answer to the previous part of the question, and attempt this part again as well using the correct value of \(\cos \theta_y\).
    </span>
    <span class="multilang" lang="de">
        Bitte ändern Sie Ihre Antwort zum vorherigen Teil der Frage und versuchen Sie auch diesen Teil noch einmal mit dem korrekten Wert von \(\cos \theta_y\).
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ANGy-FALSE2</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You do not appear to have entered the correct answer. Perhaps you are using some incorrect units? Perhaps you have not rounded correctly? Perhaps you are using an incorrect formula?
    </span>
    <span class="multilang" lang="de">
        Sie scheinen nicht die richtige Antwort eingegeben zu haben. Möglicherweise verwenden Sie falsche Einheiten? Vielleicht haben Sie nicht richtig gerundet? Möglicherweise verwenden Sie eine falsche Formel?
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Part3bz</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>s_thetaz_deg_ECF: significantfigures(ev(acos(ans4z)*(180/pi),float),3);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5z</sans>
                <tans>r_thetaz_d</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGz-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ANGz-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5z</sans>
                <tans>r_thetaz_r</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGz-INRADIANS</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your (rounded) answer appears to have been given in radians. Please enter the appropriate answer in degrees. Remember, only round your answer at the very final stage of the calculation!
    </span>
    <span class="multilang" lang="de">
        Ihre (gerundete) Antwort scheint im Bogenmaß angegeben worden zu sein. Bitte geben Sie die entsprechende Antwort in Grad ein. Denken Sie daran, Ihre Antwort erst im allerletzten Schritt der Berechnung zu runden!
    </span>
</p>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ANGz-FALSE1</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5z</sans>
                <tans>s_thetaz_deg_ECF</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ANGz-FALSE-TRUEwECF</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You appear to have used the correct formula to find \(\theta_z\) in degrees (and to round it to the appropriate number of significant figures!), but it appears that your value of \(\cos \theta_z\) from earlier is incorrect; this means your answer here
        is not correct, either.
    </span>
    <span class="multilang" lang="de">
        Sie scheinen die richtige Formel verwendet zu haben, um \(\theta_z\) in Grad zu ermitteln (und es auf die entsprechende Anzahl signifikanter Stellen zu runden!), aber es sieht so aus, dass Ihr Wert von \(\cos \theta_z\) von früher falsch ist; das bedeutet Ihre Antwort hier
        ist auch nicht korrekt.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please amend your answer to the previous part of the question, and attempt this part again as well using the correct value of \(\cos \theta_z\).
    </span>
    <span class="multilang" lang="de">
        Bitte ändern Sie Ihre Antwort zum vorherigen Teil der Frage und versuchen Sie auch diesen Teil noch einmal mit dem korrekten Wert von \(\cos \theta_z\).
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ANGz-FALSE2</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You do not appear to have entered the correct answer. Perhaps you are using some incorrect units? Perhaps you have not rounded correctly? Perhaps you are using an incorrect formula?
    </span>
    <span class="multilang" lang="de">
        Sie scheinen nicht die richtige Antwort eingegeben zu haben. Vielleicht verwenden Sie falsche Einheiten? Vielleicht haben Sie nicht richtig gerundet? Vielleicht verwenden Sie eine falsche Formel?
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ratio</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>drRowVec</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ansRat</sans>
                <tans>drRowVec</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>DIR-RATIO-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>DIR-RATIO-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Remember how we obtain the direction ratios for a three-dimensional vector from its expression in terms of \(\mathbf{i}\) and \(\mathbf{j}\) components, and try again.
    </span>
    <span class="multilang" lang="de">
        Denken Sie daran, wie wir die Richtungsverhältnisse für einen dreidimensionalen Vektor aus seinem Ausdruck in Form von \(\mathbf{i}\) und \(\mathbf{j) erhalten }\) Komponenten und versuchen Sie es erneut.
    </span>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1427186675</deployedseed>
        <deployedseed>958945957</deployedseed>
        <deployedseed>1087728891</deployedseed>
        <deployedseed>298331475</deployedseed>
        <deployedseed>1567412295</deployedseed>
        <deployedseed>191135283</deployedseed>
        <deployedseed>5280898</deployedseed>
        <deployedseed>261526416</deployedseed>
        <deployedseed>1164644081</deployedseed>
        <deployedseed>727148584</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans4x</name>
                <value>cos_thetax</value>
            </testinput>
            <testinput>
                <name>ans4y</name>
                <value>cos_thetay</value>
            </testinput>
            <testinput>
                <name>ans4z</name>
                <value>cos_thetaz</value>
            </testinput>
            <testinput>
                <name>ans5x</name>
                <value>r_thetax_d</value>
            </testinput>
            <testinput>
                <name>ans5y</name>
                <value>r_thetay_d</value>
            </testinput>
            <testinput>
                <name>ans5z</name>
                <value>r_thetaz_d</value>
            </testinput>
            <testinput>
                <name>ansRat</name>
                <value>drRowVec</value>
            </testinput>
            <expected>
                <name>Part3ax</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSx-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3ay</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSy-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3az</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSz-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bx</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ANGx-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3by</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ANGy-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bz</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ANGz-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>ratio</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>DIR-RATIO-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans4x</name>
                <value>cos_thetax</value>
            </testinput>
            <testinput>
                <name>ans4y</name>
                <value>cos_thetay</value>
            </testinput>
            <testinput>
                <name>ans4z</name>
                <value>cos_thetaz</value>
            </testinput>
            <testinput>
                <name>ans5x</name>
                <value>r_thetax_r</value>
            </testinput>
            <testinput>
                <name>ans5y</name>
                <value>r_thetay_r</value>
            </testinput>
            <testinput>
                <name>ans5z</name>
                <value>r_thetaz_r</value>
            </testinput>
            <testinput>
                <name>ansRat</name>
                <value>drRowVec</value>
            </testinput>
            <expected>
                <name>Part3ax</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSx-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3ay</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSy-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3az</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>COSz-TRUE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bx</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>ANGx-INRADIANS</expectedanswernote>
            </expected>
            <expected>
                <name>Part3by</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>ANGy-INRADIANS</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bz</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>ANGz-INRADIANS</expectedanswernote>
            </expected>
            <expected>
                <name>ratio</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>DIR-RATIO-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans4x</name>
                <value>ecfTestCosx</value>
            </testinput>
            <testinput>
                <name>ans4y</name>
                <value>ecfTestCosy</value>
            </testinput>
            <testinput>
                <name>ans4z</name>
                <value>ecfTestCosz</value>
            </testinput>
            <testinput>
                <name>ans5x</name>
                <value> r_ecfthetax_d: significantfigures(ev(acos(ecfTestCosx)*(180/pi),float),3);</value>
            </testinput>
            <testinput>
                <name>ans5y</name>
                <value> r_ecfthetay_d: significantfigures(ev(acos(ecfTestCosy)*(180/pi),float),3);</value>
            </testinput>
            <testinput>
                <name>ans5z</name>
                <value> r_ecfthetaz_d: significantfigures(ev(acos(ecfTestCosz)*(180/pi),float),3);</value>
            </testinput>
            <testinput>
                <name>ansRat</name>
                <value>ecfTest</value>
            </testinput>
            <expected>
                <name>Part3ax</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>COSx-FALSE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3ay</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>COSy-FALSE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3az</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>COSz-FALSE</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bx</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>ANGx-FALSE-TRUEwECF</expectedanswernote>
            </expected>
            <expected>
                <name>Part3by</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>ANGy-FALSE-TRUEwECF</expectedanswernote>
            </expected>
            <expected>
                <name>Part3bz</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>ANGz-FALSE-TRUEwECF</expectedanswernote>
            </expected>
            <expected>
                <name>ratio</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>DIR-RATIO-FALSE</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
