<quiz>
    <question type="stack">
        <name>
            <text>9.5.4.7 Perpendicular distance between plane and origin</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        What is the perpendicular distance from the origin to of the plane given by \(\underline{r}\cdot{@a1@}={@c1@}\)?
    </span>
    <span class="multilang" lang="de">
        Wie groß ist der senkrechte Abstand vom Ursprung zur Ebene, die durch \(\underline{r}\cdot{@a1@}={@c1@}\) gegeben ist?
    </span>
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
[[feedback:prt1]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        By dividing through by \(\lvert\underline{n}\rvert = {@maga1@}\), we can express the equation of the plane as \({@stackvector(r)@}\cdot \underline{\hat{n}} =c\) where \(\underline{\hat{n}}\) is a unit vector parallel to \(\underline{n}\) and \(c\) is
        a constant.
    </span>
    <span class="multilang" lang="de">
        Durch Division durch \(\lvert\underline{n}\rvert = {@maga1@}\) können wir die Gleichung der Ebene als \({@stackvector(r)@}\cdot \underline{\hat{n}} =c\) ausdrücken, wobei \(\underline{\hat{n}}\) ein Einheitsvektor parallel zu \(\underline{n}\) ist und \(c\)
        eine Konstante ist.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Here \(\underline{\hat{n}} = \displaystyle{\frac{1}{{@maga1@}}{@a1@}}\) and \(c=\displaystyle{\frac{{@c1@}}{{@maga1@}}}\).
    </span>
    <span class="multilang" lang="de">
        Hier \(\underline{\hat{n}} = \displaystyle{\frac{1}{{@maga1@}}{@a1@}}\) und \(c=\displaystyle{\frac{{@c1@}}{{@maga1@}}}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The distance between the plane and the origin is thus given by \(\displaystyle{{@h1@}}\).
    </span>
    <span class="multilang" lang="de">
        Der Abstand zwischen der Ebene und dem Ursprung ist somit gegeben durch \(\displaystyle{{@h1@}}\).
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stackvectortex(ex):= block(sconcat("{ \\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
a1: matrix([7],[7],[5]) - rand(matrix([3],[4],[2]));
maga1: sqrt(dotproduct(a1,a1));
c1: 7-rand(5);
h1: c1/maga1;]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>The distance from the plane to the origin is given by {@h1@};</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>h1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>h1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1578399687</deployedseed>
        <deployedseed>707307225</deployedseed>
        <deployedseed>2064314929</deployedseed>
        <deployedseed>661465317</deployedseed>
        <deployedseed>1128504441</deployedseed>
        <deployedseed>2049153565</deployedseed>
        <deployedseed>339768357</deployedseed>
        <deployedseed>854742259</deployedseed>
        <deployedseed>434622819</deployedseed>
        <deployedseed>1001742774</deployedseed>
        <deployedseed>907463674</deployedseed>
        <deployedseed>1607721965</deployedseed>
        <deployedseed>915053910</deployedseed>
        <deployedseed>1132488088</deployedseed>
        <deployedseed>1177899607</deployedseed>
        <deployedseed>992291957</deployedseed>
        <deployedseed>76477694</deployedseed>
        <deployedseed>1961864795</deployedseed>
        <deployedseed>522586823</deployedseed>
        <deployedseed>1840609713</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>h1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>0</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
