<quiz>
    <question type="stack">
        <name>
            <text>6.1.1.2 Vereinfachen Sie den Ausdruck 1 [Simplify the expression 1]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Simplify the expression \(\displaystyle{\dfrac{b^{@aa@} \,b^{@bb@}}{b^{@cc@}}}.\)
    </span>
    <span class="multilang" lang="de">
        Vereinfachen Sie den Ausdruck \(\displaystyle{\dfrac{b^{@aa@} \,b^{@bb@}}{b^{@cc@}} }.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We can do this in two steps.
    </span>
    <span class="multilang" lang="de">
        Wir können dies in zwei Schritten tun.
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        First simplify the numerator:
    </span>
    <span class="multilang" lang="de">
        Vereinfachen Sie zuerst den Zähler:
    </span>
</p>
<p>
    \(\displaystyle{b^{@aa@}b^{@bb@}}=\) \(b\)^[[input:ans1]][[validation:ans1]]
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Now include the denominator:
    </span>
    <span class="multilang" lang="de">
        Fügen Sie jetzt den Nenner hinzu:
    </span>
</p>
<p>
    \(\displaystyle{\dfrac{b^{@aa@}b^{@bb@}}{b^{@cc@}}}\)=\(b\)^[[input:ans2]] [[validation:ans2]]
</p>
<br/>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We have:
    </span>
    <span class="multilang" lang="de">
        Wir erhalten:
    </span>
</p>
<p>
    \(\displaystyle \dfrac{b^{@aa@} \,b^{@bb@}}{b^{@cc@}} = \dfrac{b^{({@aa@})+{@bb@}}}{b^{@cc@}} = \dfrac{b^{@ta1@}}{b^{@cc@}} = b^{{@ta1@}-{@cc@}} = b^{@ta2@}.\)
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ordergreat(m,n);
aa1: rand(3)+1;
aa2: rand(4)+1;
aa: (aa1*m)-(aa2*n);
aa3: rand(5)+1;
bb: rand(7)+1;
cc: aa3*m;
ta1:  ((aa1*m)+bb-(aa2*n));
ta2: ((aa1*m)+bb-(aa2*n)-(aa3*m));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<p>
    [[feedback:prt1]]
</p>
<p>
    [[feedback:prt2]]
    <br/>
</p>
]]></text>
        </specificfeedback>
        <questionnote>
            <text>\[ \displaystyle{b^{@aa@} \, b^{@bb@}}= {@ta1@}\]
\[ \displaystyle{\dfrac{b^{@aa@} \, b^{@bb@}}{b^{@cc@}}} = {@ta2@}.\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your answer has the correct value but you need to simplify it.
</span>
<span class="multilang" lang="de">
    Ihre Antwort hat den richtigen Wert, aber Sie müssen ihn vereinfachen.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer has the correct value but you need to simplify it.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort hat den richtigen Wert, aber Sie müssen ihn vereinfachen.
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1332212346</deployedseed>
        <deployedseed>1868150579</deployedseed>
        <deployedseed>868525191</deployedseed>
        <deployedseed>367333791</deployedseed>
        <deployedseed>1083164968</deployedseed>
        <deployedseed>384686307</deployedseed>
        <deployedseed>388599967</deployedseed>
        <deployedseed>330983371</deployedseed>
        <deployedseed>1918799835</deployedseed>
        <deployedseed>121687614</deployedseed>
        <deployedseed>293972492</deployedseed>
        <deployedseed>709728796</deployedseed>
        <deployedseed>119880551</deployedseed>
        <deployedseed>1385245951</deployedseed>
        <deployedseed>1157173526</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>b^((aa1*m)-bb-(aa2*n))</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>b^((aa1*m)-bb-(aa2*n)+(aa3*m))</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
