<quiz>
    <question type="stack">
        <name>
            <text>6.4.4.2 tanhx = b</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Find the solution for \(x\) if \(\tanh x=\displaystyle{@bb@}\). Give a precise answer.
    </span>
    <span class="multilang" lang="de">
        Finden Sie die Lösung für \(x\), wenn \(\tanh x=\displaystyle{@bb@}\). Geben Sie eine genaue Antwort.
    </span>
</p>
<p>
    \( x = \) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    First re-write \(\tanh x\) in terms of exponentials: \[\tanh x=\dfrac{e^x-e^{-x}}{e^x+e^{-x}}=\dfrac{e^{2x}-1}{e^{2x}+1}\] Now substitute into \(\tanh x={@bb@}\): \(\tanh x={@bb@}\) implies \[\dfrac{e^{2x}-1}{e^{2x}+1}={@bb@}\] so, on simplifying, \[\dfrac{e^{2x}-1}{e^{2x}+1}={@bb@}
\text{ implies } (e^{2x}-1)=\dfrac{1}{{@aa@}}(e^{2x}+1)\] so \(\dfrac{({@aa@}-1)e^{2x}}{{@aa@}}=\dfrac{{@aa@}+1}{{@aa@}}\) so, finally, \(e^{2x}=\frac{{@aa@}+1}{{@aa@}-1}\). Now complete your solution by finding \(x\): \[e^{2x}=\frac{{@aa@}+1}{{@aa@}-1}
\text{ so } x= \dfrac{1}{2}\ln \left(\frac{{@aa@}+1}{{@aa@}-1}\right) = \dfrac{1}{2}\ln \left(\frac{{@aa+1@}}{{@aa-1@}}\right).\] Alternatively, many calculators can directly calculate the inverse function \(\tanh^{-1}\). If you have such a calculator
then you can use the fact that \[\tanh x={@bb@} \qquad \text{implies} \qquad x=\tanh^{-1}{@bb@}\] to obtain \(x\) directly
</span>
<span class="multilang" lang="de">
    Schreiben Sie zunächst \(\tanh x\) in Exponentialzahlen um: \[\tanh x=\dfrac{e^x-e^{-x}}{e^ x+e^{-x}}=\dfrac{e^{2x}-1}{e^{2x}+1}\] Substituieren Sie nun in \(\tanh x={@bb@}\): \(\tanh x={@bb@}\), was \[\dfrac{e^{2x}-1}{e^{2x}+1}={@bb@}\] impliziert, also vereinfacht \[ \dfrac{e^{2x}-1}{e^{2x}+1}={@bb@}
\text{ impliziert } (e^{2x}-1)=\dfrac{1}{{@aa@}}(e^{2x}+1)\] also \(\dfrac{({@aa@} -1)e^{2x}}{{@aa@}}=\dfrac{{@aa@}+1}{{@aa@}}\) also schließlich \(e^{2x}=\frac{{@aa@}+1}{{@aa@}-1}\). Vervollständigen Sie nun Ihre Lösung, indem Sie \(x\) finden: \[e^{2x}=\frac{{@aa@}+1}{{@aa@}-1}
\text{ also } x= \dfrac{1}{2}\ln \left(\frac{{@aa@}+1}{{@aa@}-1}\right) = \dfrac{1}{ 2}\ln \left(\frac{{@aa+1@}}{{@aa-1@}}\right).\] Alternativ können viele Taschenrechner direkt die Umkehrfunktion \(\tanh^{- 1}\) berechnen. Wenn Sie einen solchen Rechner haben
dann können Sie die Tatsache verwenden, dass \[\tanh x={@bb@} \qquad \text{impliziert} \qquad x=\tanh^{-1}{@bb@},\] um \(x\) direkt zu erhalten
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa : 2+rand(5)
bb : 1/aa
exp : ((aa-1)*e^(2*x))*bb = (aa+1)*bb
ta : ln((aa+1)/(aa-1))/2</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[\tanh x = {@bb@} \Rightarrow x = {@ta@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>tanh</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        You need to give a precise answer, you should expect an expression involving logs.
    </span>
    <span class="multilang" lang="de">
        Sie müssen eine genaue Antwort geben, Sie sollten einen Ausdruck verwenden, der Logarithmen enthält.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1193163592</deployedseed>
        <deployedseed>243744400</deployedseed>
        <deployedseed>1072245417</deployedseed>
        <deployedseed>499832638</deployedseed>
        <deployedseed>24763017</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
