<quiz>
    <question type="stack">
        <name>
            <text>6.4.2.2 Solve e^ax=b</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the equation \({@exp@}={@n1@}\). To write your answer as a natural logarithm please write \(\log(n)\).
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die Gleichung \({@exp@}={@n1@}\). Um Ihre Antwort als natürlichen Logarithmus zu schreiben, schreiben Sie bitte \(\log(n)\).
    </span>
    <br/>
</p>
<p>
    \( x= \) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    We have that \({@exp@}={@n1@}\) implies \(e^x = \sqrt{@n1@}\) where we take the positive square root as exponential functions are
    <b>
        always positive
    </b>
    . Taking logarithms of both sides gives \[x = \ln(\sqrt{@n1@}) = \frac{\ln({@n1@})}{2}.\]
</span>
<span class="multilang" lang="de">
    Wir wissen, dass \({@exp@}={@n1@}\) \(e^x = \sqrt{@n1@}\) impliziert, wobei wir die positive Quadratwurzel nehmen, da Exponentialfunktionen
    <b>
        immer positiv
    </b>
    sind. Logarithmus auf beiden Seiten nehmen ergibt \[x = \ln(\sqrt{@n1@}) = \frac{\ln({@n1@})}{2}.\]
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>n1:10+rand(40);
ta: ln(n1)/2;
simp:false
exp: (e^x)^2;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:Result]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[ {@exp@} ={@n1@} \Rightarrow x= {@ta@}.\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>solve</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>Result</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>Result-0-T </trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>Result-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>log(-sqrt(n1))</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>Result-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    The logarithm of a negative number is not defined in the \( \mathbb{R}.\)
</span>
<span class="multilang" lang="de">
    Der Logarithmus einer negativen Zahl ist in \( \mathbb{R}\) nicht definiert.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>Result-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The value of \((e^x)^2\) for your input is \({@ev(e^(2*ans1),simp)@}\) which is not equal to {@n1@}.
    </span>
    <span class="multilang" lang="de">
        Der Wert von \((e^x)^2\) für Ihre Eingabe ist \({@ev(e^(2*ans1),simp)@}\ ), was nicht gleich {@n1@} ist.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>225486370</deployedseed>
        <deployedseed>1877197567</deployedseed>
        <deployedseed>1591925146</deployedseed>
        <deployedseed>1890734474</deployedseed>
        <deployedseed>1785313027</deployedseed>
        <deployedseed>724448384</deployedseed>
        <deployedseed>329290905</deployedseed>
        <deployedseed>612838917</deployedseed>
        <deployedseed>1376033034</deployedseed>
        <deployedseed>910068254</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>Result-0-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>Result-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>log(-sqrt(n1))</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>Result-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
