<quiz>
    <question type="stack">
        <name>
            <text>6.4.3.2 Chemical Reaction 1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The reaction
    </span>
    <span class="multilang" lang="de">
        Die Reaktion
    </span>
</p>
<p>
    2NO
    <span class="math inline">
        \(_2\)
    </span>
    (g)
    <span class="math inline">
        \(\longrightarrow\)
    </span>
    2NO(g) + O
    <span class="math inline">
        \(_2\)
    </span>
    (g)
</p>
<span class="multilang" lang="en">
    <nowrap>
        has a reaction rate constant of \(1.0 \times 10 ^{-10}\) \(\text{s}^{-1}\) at \(300 \text{K}\) and activation energy of \(111 \text{kJ mol}^{-1} = 111 000\text{ J mol}^{-1}\). Use Arrhenius' law to find the reaction rate constant at a temperature of {@Tt@}\(\text{K}\).
        Please give your answer to \(4\) s.f.
    </nowrap>
</span>
<span class="multilang" lang="de">
    <nowrap>
        hat eine Reaktionsgeschwindigkeitskonstante von \(1.0 \times 10 ^{-10}\) \(\text{s}^{-1}\) bei \(300 \text{K}\) und einer Aktivierungsenergie von \(111 \text{kJ mol}^{-1} = 111 000 \text{ J mol}^{-1}\). Verwenden Sie das Gesetz von Arrhenius, um die Reaktionsgeschwindigkeitskonstante bei einer Temperatur von {@Tt@}\(\text{K}\) zu finden.
            Bitte geben Sie Ihre Antwort mit \(4\) signifikanten Stellen ein.
    </nowrap>
</span>
<p>
    \( k = \) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    Rearranging Arrhenius' equation gives \[k_0 = k \, \exp\left(\frac{E_a}{R\,T}\right).\] Substituting the values gives \[k_0= 2.126 \times 10^{9} \mbox{ s}^{-1}.\] Now we use this value of \(k_0\) with \(E_a\) in Arrhenius' equation (1) to find \(k\) at
        \(T\) = {@Tt@} K \[k= k_0 \, \exp \left( \frac{-E_a}{R\,T} \right)= {@tar@} \mbox{ s}^{-1}.\]
</span>
<span class="multilang" lang="de">
    Das Umstellen der Arrhenius-Gleichung ergibt \[k_0 = k \, \exp\left(\frac{E_a}{R\,T}\right).\] Ersetzen der Werte ergibt \[k_0= 2.126 \times 10^{9} \mbox{ s}^{-1}.\] Nun verwenden wir diesen Wert von \(k_0\) mit \(E_a\) in Arrhenius' Gleichung (1) um \(k\) bei 
            \(T\) = {@Tt@} K zu finden \[k= k_0 \, \exp \left( \frac{-E_a}{R\,T} \right)= {@tar@} \mbox{ s} ^{-1}.\]
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>k0 : e^(111000/(8.314*300))*10^(-10);
Tt : 273;
ta : k0*e^(-111000/(8.314*Tt))
tar: significantfigures(ta,4);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>{@Tt@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>tar</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans1</sans>
                <tans>tar</tans>
                <testoptions>4</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>tar</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>0</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
