<quiz>
    <question type="stack">
        <name>
            <text>6.4.4.5 Exercises b)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the equations
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die Gleichungen
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \(e^{2x}={@d1@}e^x\)
        </p>
        <p>
            [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
        </p>
    </li>
    <li>
        <p>
            \({@exp2@}=0\)
        </p>
        <p>
            [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
        </p>
    </li>
    <li>
        <p>
            \(\cosh x={@f1@}\)
        </p>
        <p>
            \(x=\) [[input:ans6]][[validation:ans6]]
        </p>
        <p>
            \(x=\)[[input:ans7]][[validation:ans7]][[feedback:prt6]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            \(e^{2x}={@d1@}e^x\)
        </p>
        <p>
            <span class="multilang" lang="en">
                Take logs of base \(e\) on both sides:
            </span>
            <span class="multilang" lang="de">
                Logarithmen zur Basis \(e\) auf beiden Seiten nehmen:
            </span>
        </p>
        <p>
            \(\ln(e^{2x}) = \ln({@d1@}e^x) \Rightarrow\)
        </p>
        <p>
            \(2x=\ln {@d1@}+x \Rightarrow x=\ln {@d1@}.\)
        </p>
    </li>
    <li>
        <p>
            \({@exp2@}=0\)
        </p>
        <p>
            <span class="multilang" lang="en">
                Let \(y=e^x,\) then the equation becomes \({@expiny@}=0.\)
            </span>
            <span class="multilang" lang="de">
                Setzen Sie \(y=e^x,\) dann wird die Gleichung zu \({@expiny@}=0.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                By solving the quadratic we find: \(y={@e1@}\) or \(y={@e2@}\) (we cannot take the negative root since exponentials can never be negative).
            </span>
            <span class="multilang" lang="de">
                Durch Lösen der quadratischen Gleichung erhalten wir: \(y={@e1@}\) oder \(y={@e2@}\) (wir können keine negative Wurzel ziehen da Exponentiale niemals negativ sein können).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Thus \(e^x= {@e1@} \Rightarrow x=\ln ({@e1@}).\)
            </span>
            <span class="multilang" lang="de">
                Also \(e^x= {@e1@} \Rightarrow x=\ln ({@e1@}).\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \(e^x+e^{-x}={@f12@}\) can be written as \({@exp3@}=0.\)
            </span>
            <span class="multilang" lang="de">
                \(e^x+e^{-x}={@f12@}\) kann geschrieben werden als \({@exp3@}=0.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Let \(y=e^x,\) then the equation becomes \({@y^2-2*f1*y+1@} =0. \)
            </span>
            <span class="multilang" lang="de">
                Setzen Sie \(y=e^x,\) dann wird die Gleichung zu \({@y^2-2*f1*y+1@} =0. \)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                By solving the quadratic we find: \(y = e^x = \dfrac{{@f12@} \pm \sqrt{({@f12@})^2-4}}{2} = {@f1@} \pm \sqrt{{@f1@}^2-1}.\)
            </span>
            <span class="multilang" lang="de">
                Durch Lösen der quadratischen Gleichung erhalten wir: \(y = e^x = \dfrac{{@f12@} \pm \sqrt{({@f12@})^2 -4}}{2} = {@f1@} \pm \sqrt{{@f1@}^2-1}.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                We accept both roots since both are positive.
            </span>
            <span class="multilang" lang="de">
                Wir akzeptieren beide Wurzeln, da beide positiv sind.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Finally we have \(x={@tf1@}\) or \(x={@tf2@}.\)
            </span>
            <span class="multilang" lang="de">
                Schließlich erhalten wir \(x={@tf1@}\) oder \(x={@tf2@}.\)
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>d1 : 17 + rand(5)
tad : log(d1)

e1: 2+ rand(5)
e2: -3 + rand(3)
tae: log(e1)
exp2 : expand((e^x-e1)*(e^x-e2))
expiny : expand((y-e1)*(y-e2))

f1 : 3 + rand(5)
f12:2*f1
exp3 : e^(2*x)-2*f1*e^(x)+1
tf1 : log(f1+sqrt(f1^2-1))
tf2 : log(f1-sqrt(f1^2-1))</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[e^{2x}={@d1@}e^x \Rightarrow x = {@tad@}\]
\[e^{2x}={@d1@}e^x \Rightarrow x = {@tae@}\]
\[e^{2x}={@d1@}e^x \Rightarrow x = {@tf1@}, \, {@tf2@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>tad</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>tae</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>tf1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans7</name>
            <type>algebraic</type>
            <tans>tf2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>tad</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>tae</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>{ans6,ans7}</sans>
                <tans>{tf1,tf2}</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>871867844</deployedseed>
        <deployedseed>1235455587</deployedseed>
        <deployedseed>377909047</deployedseed>
        <deployedseed>699015544</deployedseed>
        <deployedseed>1403802325</deployedseed>
        <deployedseed>1803304232</deployedseed>
        <deployedseed>546466030</deployedseed>
        <deployedseed>1666292854</deployedseed>
        <deployedseed>1212179119</deployedseed>
        <deployedseed>531184388</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans4</name>
                <value>tad</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>tae</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>tf1</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>tf2</value>
            </testinput>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans4</name>
                <value>tad+1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>tae+2</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>tf1+2</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>tf2+4</value>
            </testinput>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
