<quiz>
    <question type="stack">
        <name>
            <text>6.4.3.3 Chemical Reaction 2</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<span class="multilang" lang="en">
    For a chemical reaction with frequency factor \(k_0 = 0.5 \, \text{s}^{-1}\) and ratio \(E_a/R={@Tt@} \, \text{K}\), use Arrhenius' law to find the temperature at which the reaction rate constant would be equal to \(0.1 \, \text{s}^{-1}\). Please give
        your answer to \(3\) s.f.
</span>
<span class="multilang" lang="de">
    Für eine chemische Reaktion mit Häufigkeitsfaktor \(k_0 = 0.5 \, \text{s}^{-1}\) und Verhältnis \(E_a/R={@Tt @} \, \text{K}\), verwenden Sie das Gesetz von Arrhenius, um die Temperatur zu finden, bei der die Reaktionsgeschwindigkeitskonstante gleich \(0.1 \, \text{s}^{-1}\) wäre. Bitte geben Sie
            Ihre Antwort in \(3\) signifikanten Stellen an
</span>
<p>
    \( T = \) [[input:ans1]] [[validation:ans1]] \( \text{K}\)
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    Rearranging Equation (1) \[\frac{k}{k_0} = \exp \left( \frac{-E_a}{R\,T} \right).\] Taking the natural logarithm of both sides \[\ln \left( \frac{k}{k_0} \right) = \frac{-E_a}{R\,T}\] so that \[T = \frac{-E_a}{R\,\ln \left( k/k_0 \right)}= \frac{E_a}{R\,\ln
        \left( k_0/k \right)}.\] Substituting the values gives \(T={@ta@} \mbox{ K}.\)
</span>
<span class="multilang" lang="de">
    Gleichung (1) umstellen \[\frac{k}{k_0} = \exp \left( \frac{-E_a}{R\,T} \right). \] Den natürlichen Logarithmus beider Seiten nehmen \[\ln \left( \frac{k}{k_0} \right) = \frac{-E_a}{R\,T}\] sodass \[T = \frac {-E_a}{R\,\ln \left( k/k_0 \right)}= \frac{E_a}{R\,\ln
            \left( k_0/k \right)}.\] Einsetzen der Werte ergibt \(T={@ta@} \mbox{ K}.\)
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>Tt: 780+rand(20);
ta : significantfigures(Tt/ln(0.5/0.1),3);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\( E_a/R={@Tt@} \, \text{K} \rightarrow T = {@ta@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumSigFigs</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions>3</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>384728890</deployedseed>
        <deployedseed>2012963661</deployedseed>
        <deployedseed>703262219</deployedseed>
        <deployedseed>1193452761</deployedseed>
        <deployedseed>894926972</deployedseed>
        <deployedseed>538789899</deployedseed>
        <deployedseed>1152128249</deployedseed>
        <deployedseed>36147130</deployedseed>
        <deployedseed>4251478</deployedseed>
        <deployedseed>1879417725</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
