<quiz>
    <question type="stack">
        <name>
            <text>6.4.2.5 Solve ae^{2x} − be^x + c = 0</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the following equation for \(x\).
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die folgende Gleichung nach \(x\) auf.
    </span>
</p>
<p>
    \[\displaystyle{{@exp2@}=0}\]
</p>
<p>
    \(x=\) [[input:ans1]] [[validation:ans1]]
</p>
<p>
    \(x=\) [[input:ans2]] [[validation:ans2]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    First write this equation as a quadratic in the variable \(y=e^x\) remembering that \(e^{2x}\equiv (e^x)^2\): \[{@exp2@} \, \text{ becomes } \, {@expiny@}=0.\] Now solve the quadratic for \(y\): \[{@expiny@}=0 \, \text{ implies } \, (y-{@aa@})(y-{@bb@})=0.\]
        Therefore \(y={@aa@}\) or \(y={@bb@}.\)
</span>
<span class="multilang" lang="de">
    Schreiben Sie diese Gleichung zuerst als quadratische Gleichung der Variablen \(y=e^x\) und denken Sie daran, dass \(e^{2x}\equiv (e^x)^2 \): \[{@exp2@} \, \text{ wird zu } \, {@expiny@}=0.\] Lösen Sie nun die Quadrate nach \(y\) auf: \[{@expiny@}=0 \ , \text{ impliziert } \, (y-{@aa@})(y-{@bb@})=0.\]
        Also \(y={@aa@}\) oder \(y={@bb@}.\)
</span>
<p>
    <span class="multilang" lang="en">
        Finally, for each of your values of \(y\), find \(x\):
    </span>
    <span class="multilang" lang="de">
        Schließlich finden Sie \(x\) für jeden Ihrer Werte von \(y\):
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        <b>
            Case 1
        </b>
        If \(y={@aa@}\) then \({@aa@}=e^x\) implies \(x=\ln {@aa@}\);
    </span>
    <span class="multilang" lang="de">
        <b>
            Fall 1
        </b>
        Wenn \(y={@aa@}\) dann impliziert \({@aa@}=e^x\) \(x=\ln {@aa@}\);
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        <b>
            Case 2
        </b>
        If \(y={@bb@}\) then \({@bb@}=e^x\) implies \(x=\ln {@bb@}.\)
    </span>
    <span class="multilang" lang="de">
        <b>
            Fall 2
        </b>
        Wenn \(y={@bb@}\) dann impliziert \({@bb@}=e^x\) \(x=\ln {@bb@}.\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa : 2+rand(3)
bb : 5+rand(5)
ta1 : log(aa)
ta2 : log(bb)
exp2 : expand((e^x-aa)*(e^x-bb))
expiny : expand((y-aa)*(y-bb))</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\[{@exp2@}=0, \, x = {@{ta1,ta2}@}.\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>{ans1,ans2}</sans>
                <tans>{ta1,ta2}</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>280994670</deployedseed>
        <deployedseed>1903939603</deployedseed>
        <deployedseed>944805284</deployedseed>
        <deployedseed>1840076342</deployedseed>
        <deployedseed>695245441</deployedseed>
        <deployedseed>144695415</deployedseed>
        <deployedseed>1032625739</deployedseed>
        <deployedseed>1261242567</deployedseed>
        <deployedseed>1411295672</deployedseed>
        <deployedseed>1293924990</deployedseed>
        <deployedseed>1931047277</deployedseed>
        <deployedseed>793654513</deployedseed>
        <deployedseed>760809054</deployedseed>
        <deployedseed>1876305588</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2+12</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
