<quiz>
    <question type="stack">
        <name>
            <text>2.1.2.2 Funktionswert auswerten [Evaluate value of function]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Given \(f\left(x\right) = {@fx@}\) evaluate:
    </span>
    <span class="multilang" lang="no">
        Gitt \(f\left(x\right) = {@fx@}\) evaluer:
    </span>
    <span class="multilang" lang="de">
        Gegeben sei \(f\left(x\right) = {@fx@}\). Werten Sie aus:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \(f\left({@xvalue1@}\right)=\) [[input:ans1]]. [[validation:ans1]] [[feedback:Result]]
        </p>
    </li>
    <li>
        <p>
            \(f\left({@xvalue2@}\right)=\) [[input:ans2]]. [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We must substitute {@xvalue1@} in place of \(x\) in the definition of the function:
            </span>
            <span class="multilang" lang="no">
                Vi må substituere  \(x\) med {@xvalue1@} i definisjonen til funksjonen:
            </span>
            <span class="multilang" lang="de">
                Wir müssen {@xvalue1@} anstelle von \(x\) in der Definition der Funktion einsetzen:
            </span>
        </p>
        \[\begin{array}{rcl}\displaystyle f\left({@xvalue1@}\right) &amp;\displaystyle =&amp;\displaystyle {@a@} \cdot ({@xvalue1@})^2 + ({@b@}) \cdot ({@xvalue1@}) + ({@c@})\\ &amp;\displaystyle =&amp;\displaystyle {@fxvalue@}\mbox{.} \end{array}\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We must substitute {@xvalue2@} in place of \(x\) in the definition of the function:
            </span>
            <span class="multilang" lang="no">
                Vi må substituere \(x\) med {@xvalue2@} i definisjonen til funksjonen:
            </span>
            <span class="multilang" lang="de">
                Wir müssen {@xvalue2@} anstelle von \(x\) in der Definition der Funktion einsetzen:
            </span>
        </p>
        \[\begin{array}{rcl}\displaystyle f\left({@xvalue2@}\right) &amp;\displaystyle =&amp;\displaystyle {@a@} \cdot ({@xvalue2@})^2 + ({@b@}) \cdot ({@xvalue2@}) + ({@c@})\\ &amp;\displaystyle =&amp;\displaystyle {@fxvalue_a2@}\mbox{.} \end{array}\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a_magnitude:rand(8)+2;
 a_sign:2*rand(2)-1;
 a:a_magnitude*a_sign;
 b_magnitude:rand(8)+2;
 b_sign:2*rand(2)-1;
 b:b_magnitude*b_sign;
 c_magnitude:rand(9)+1;
 c_sign:2*rand(2)-1;
 c:c_magnitude*c_sign;
 xvalue1:rand(19)-9;
 xvalue2:0;
 fx:a*x^2+b*x+c;
 fxvalue:a*xvalue1^2+b*xvalue1+c;
 ta1: fxvalue;
ta2:a*xvalue2^2+b*xvalue2+c;
 fxvalue_a2:a^2*xvalue2^2+b*xvalue2+c;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>Given \(f\left(x\right) = {@fx@}\) then \(f({@xvalue1@})={@ta1@}\) and \(f({@xvalue2@})={@ta2@}\).</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>*</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>*</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>Result</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>FNC-EVL-TRUE</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>FNC-EVL-FALSE</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>fxvalue_a2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ART-RUL-FALSE </trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        You should look at order of operations when it comes to multiplication and powers.
    </span>
    <span class="multilang" lang="de">
        Sie sollten sich die Reihenfolge der Operationen ansehen, wenn es um Multiplikation und Potenzen geht.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>Result-1-F </falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>522143659</deployedseed>
        <deployedseed>1172463441</deployedseed>
        <deployedseed>1956421563</deployedseed>
        <deployedseed>1513275847</deployedseed>
        <deployedseed>1958928738</deployedseed>
        <deployedseed>1228527132</deployedseed>
        <deployedseed>1648564778</deployedseed>
        <deployedseed>1995697291</deployedseed>
        <deployedseed>2093631138</deployedseed>
        <deployedseed>211873478</deployedseed>
        <deployedseed>2009479218</deployedseed>
        <deployedseed>1598660664</deployedseed>
        <deployedseed>332114259</deployedseed>
        <deployedseed>395369297</deployedseed>
        <deployedseed>1134648556</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>fxvalue</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>Result</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>FNC-EVL-TRUE</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>fxvalue-1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2-1</value>
            </testinput>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>Result-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
