<quiz>
    <question type="stack">
        <name>
            <text>1.3.3.2 Entfernen von 2 Sätzen von Klammern [Removing 2 sets of brackets]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Remove the brackets from each of the following expressions:
    </span>
    <span class="multilang" lang="de">
        Entfernen Sie die Klammern in jedem der folgenden Ausdrücke:
    </span>
    <br/>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \( ({@a1@}+a)({@a2@}+b)=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \( (x + {@b1@})(x + {@b2@})=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            \((x + {@c1@})(x + {@c2@})=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
            \((x + {@d1@})(x - {@d2@})=\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            \[ ({@a1@}+a) ({@a2@}+b) =({@a1@}+a) ({@a2@}) + ({@a1@}+a)(b) \]
        </p>
        <p>
            \[ = ({@a1@})({@a2@}) + (a)({@a2@}) + ({@a1@})(b) + (a)(b) \]
        </p>
        <p>
            \[ = {@a1*a2@} + {@a2@}a + {@a1@}b + ab. \]
        </p>
    </li>
    <li>
        <p>
            \[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) + (x + {@b1@})({@b2@}) \]
        </p>
        <p>
            \[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]
        </p>
        <p>
            \[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]
        </p>
    </li>
    <li>
        <p>
            \[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]
        </p>
        <p>
            \[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]
        </p>
        <p>
            \[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]
        </p>
    </li>
    <li>
        <p>
            \[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]
        </p>
        <p>
            \[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]
        </p>
        <p>
            \[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand(2)+3;
a2:a1+1;
ta1: a1*a2 + b*a1 + a*a2 + a*b;

b1: rand(2)+2;
b2:b1+1;
ta2: x^2 + (b1+b2)*x + b1*b2;
ta2f: x^2 + (b1*x) + b2*x + b1*b2;

c1: rand(3)+3;
c2:c1;
ta3: x^2 + (c1+c2)*x + c1*c2;
ta3f:x^2 + c1*x + c2*x + c1*c2;

d1: rand(2)+6;
d2:d1-2;
ta4: x^2 + (d1 - d2)*x - d1*d2;
ta4f:x^2 + d1*x -d2*x -d1*d2;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\({@ta1@}, \, {@ta2@}, \, {@ta3@}, \,  {@ta4@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(, )</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(, )</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(, )</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>(, )</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    \[ ({@a1@}+a) ({@a2@}+b) = ({@a1@}+a) ({@a2@}) + ({@a1@}+a)(b) \]
</p>
<p>
    \[ = ({@a1@})({@a2@}) + (a)({@a2@}) + ({@a1@})(b) + (a)(b) \]
</p>
<p>
    \[ = {@a1*a2@} + {@a2@}a + {@a1@}b + ab. \]
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <br/>
</p>
<br/>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        You have removed the brackets correctly, but you have not simplified your answer fully.
    </span>
    <span class="multilang" lang="de">
        Sie haben die Klammern korrekt entfernt, aber Ihre Antwort nicht vollständig vereinfacht.
    </span>
</p>
<p>
</p>
<p>
    \[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) + (x + {@b1@}) ({@b2@}) \]
</p>
<p>
    \[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]
</p>
<p>
    \[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]
</p>
<br/>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    \[ (x + {@b1@}) (x + {@b2@})= (x + {@b1@}) (x) + (x + {@b1@}) ({@b2@}) \]
</p>
<p>
    \[ = (x)(x) + ({@b1@})(x) + ({@b2@})(x) + ({@b1@})({@b2@}) \]
</p>
<p>
    \[ = x^2 + {@b1 + b2@}x + {@b1*b2@}. \]
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You have removed the brackets correctly, but you have not simplified your answer fully.
    </span>
    <span class="multilang" lang="de">
        Sie haben die Klammern korrekt entfernt, aber Ihre Antwort nicht vollständig vereinfacht.
    </span>
    <br/>
</p>
<p>
</p>
<p>
    \[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]
</p>
<p>
    \[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]
</p>
<p>
    \[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]
</p>
<p>
</p>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    \[ (x + {@c1@}) (x + {@c2@})= (x + {@c1@}) (x) + (x + {@c1@})({@c2@}) \]
    <br/>
</p>
<p>
</p>
<p>
    \[ = (x)(x) + ({@c1@})(x) + ({@c2@})(x) + ({@c1@})({@c2@}) \]
</p>
<p>
    \[ = x^2 + {@c1 + c2@}x + {@c1*c2@}. \]
</p>
<br/>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <br/>
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    <span class="multilang" lang="en">
        You have removed the brackets correctly, but you have not simplified your answer fully.
    </span>
    <span class="multilang" lang="de">
        Sie haben die Klammern korrekt entfernt, aber Ihre Antwort nicht vollständig vereinfacht.
    </span>
    <br/>
</p>
<p>
</p>
<p>
    \[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]
</p>
<p>
    \[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]
</p>
<p>
    \[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]
</p>
<br/>
<p>
</p>
<p>
</p>
<p>
    <br/>
</p>
<br/>
<br/>
<p>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
</p>
<p>
    \[ (x + {@d1@}) (x - {@d2@})= (x + {@d1@}) (x) - (x + {@d1@})({@d2@}) \]
</p>
<p>
    \[ = (x)(x) + ({@d1@})(x) - ({@d2@})(x) - ({@d1@})({@d2@}) \]
</p>
<p>
    \[ = x^2 + {@d1 - d2@}x - {@d1*d2@}. \]
</p>
<br/>
<p>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>792256749</deployedseed>
        <deployedseed>595157284</deployedseed>
        <deployedseed>159550764</deployedseed>
        <deployedseed>396959405</deployedseed>
        <deployedseed>346237507</deployedseed>
        <deployedseed>1875898046</deployedseed>
        <deployedseed>1111892330</deployedseed>
        <deployedseed>272590190</deployedseed>
        <deployedseed>1410360394</deployedseed>
        <deployedseed>1777923988</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta4</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value> x^2 + (b1*x) + b2*x + b1*b2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>x^2 + c1*x + c2*x + c1*c2</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>x^2 + d1*x -d2*x -d1*d2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
