<quiz>
    <question type="stack">
        <name>
            <text>5.2.4.2 Gleichung einer Parabel finden[Finding the equation of a parabola]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Find the equation of a parabola with vertex at \(({@bb@}, {@cc@} )\) passing through the point \(({@p1@}, {@p2@})\).
    </span>
    <span class="multilang" lang="de">
        Finden Sie die Gleichung einer Parabel mit Scheitelpunkt bei \(({@bb@}, {@cc@} )\) und die durch den Punkt \(({@p1@}, {@p2@} )\) verläuft.
    </span>
</p>
<p>
    \(y = \)[[input:ans1]] [[validation:ans1]]
    <span style="font-size: calc(0.90375rem + 0.045vw);">
        [[feedback:prt1]]
    </span>
</p>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The general form of the parabola is: \[y-C= A(x-B)^2\]
    </span>
    <span class="multilang" lang="de">
        Die allgemeine Form der Parabel ist: \[y-C= A(x-B)^2\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The vertex is at \(({@bb@},{@cc@} )\) so \(B = {@bb@} \) and \( C= {@cc@} .\) Substituting into the previous equation we get: \[ y-{@cc@} = A(x-{@bb@})^2 \Rightarrow y-{@cc@} = Ax^2 \]
    </span>
    <span class="multilang" lang="de">
        Der Scheitelpunkt ist bei \(({@bb@},{@cc@} )\), also \(B = {@bb@} \) und \( C = {@cc@} .\) Durch Einsetzen in die vorherige Gleichung erhalten wir: \[ y-{@cc@} = A(x-{@bb@})^2 \Rightarrow y-{@cc@} = Ax ^2 \]
    </span>
    <span class="multilang" lang="de">
        Der Scheitelpunkt ist bei \(({@bb@},{@cc@} )\), also \(B = {@bb@} \) und \( C = {@cc@} .\) Durch Einsetzen in die vorherige Gleichung erhalten wir: \[ y-{@cc@} = A(x-{@bb@})^2 \Rightarrow y-{@cc@} = Ax ^2 \]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Then using the point \(({@p1@},{@p2@} )\) we get: \[ {@p2@}- {@cc@} = {@p1@}^2A \Rightarrow A= \displaystyle {@an2@}\]
    </span>
    <span class="multilang" lang="de">
        Dann verwenden wir den Punkt \(({@p1@},{@p2@} )\) und erhalten: \[ {@p2@}- {@cc@} = {@p1@}^2A \Rightarrow A= \displaystyle {@an2@}\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The required equation is \[y={@ta1@}\]
    </span>
    <span class="multilang" lang="de">
        Die erforderliche Gleichung ist \[y={@ta1@}\]
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>/* Vertex (B,C) */
bb: 0
cc: rand([1,2,3]);

/*Point (P1, P2) */
p1: rand(5)+4;
p2: p1;

an1:cc;
an2:(p2-cc)/(p1)^2;
ta1: an1+an2*x^2</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(V=({@bb@}, {@cc@} ), \, P=({@p1@}, {@p2@} ), \, y= {@ta1@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>np</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>79711560</deployedseed>
        <deployedseed>308310331</deployedseed>
        <deployedseed>133203837</deployedseed>
        <deployedseed>1556033659</deployedseed>
        <deployedseed>605066843</deployedseed>
        <deployedseed>506192277</deployedseed>
        <deployedseed>948203759</deployedseed>
        <deployedseed>2105197912</deployedseed>
        <deployedseed>498325772</deployedseed>
        <deployedseed>1946695957</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
