<quiz>
    <question type="stack">
        <name>
            <text>2.2.1.5 Definitionsmenge, Wertebereich und Graph einer Funktion [Domain, Range and Graph of function]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First, a quick note on entering your answers. Remember 'and/or' here are logical 'and/or', so 'and' means the values satisfying both inequalities, 'or' means the values satisfying at least one of the inequalities.
    </span>
    <span class="multilang" lang="no">
        Først en kort kommentar om hvordan du skriver inn svarene dine. Husk 'and/or' her er logiske 'and/or', så 'and' betyr verdiene som tilfredsstiller begge ulikhetene, 'or' betyr verdiene som tilfredsstiller minst en av ulikhetene.
    </span>
    <span class="multilang" lang="de">
        Zunächst eine kurze Anmerkung zur Eingabe Ihrer Antworten. Denken Sie daran, dass „and/or“ hier logisches „and/or“ sind, also bezeichnet „and“ die Werte, die beide Ungleichungen erfüllen, „or“ bezeichnet die Werte, die mindestens eine der Ungleichungen erfüllen.
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the function given by \(f(x) = {@f@},\quad {@leftend@} \leq x \leq {@rightend@}\)
    </span>
    <span class="multilang" lang="no">
        Vurder funksjonen gitt ved \(f(x) = {@f@},\quad {@leftend@} \leq x \leq {@rightend@}\)
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie die Funktion \(f(x) = {@f@},\quad {@leftend@} \leq x \leq {@rightend@}\)
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                State the domain of the function. Give your answer as an inequality in \(x\), e.g. "\( \alpha \leq x\) and/or \(x \leq \beta\)".
            </span>
            <span class="multilang" lang="no">
                Oppgi definisjonsmengden til funksjonen. Gi svaret ditt som en ulikhet i \(x\), f.eks. "\( \alpha \leq x\) and/or \(x \leq \beta\)".
            </span>
            <span class="multilang" lang="de">
                Geben Sie den Definitionsbereich der Funktion an. Geben Sie Ihre Antwort als Ungleichung in \(x\) an, z.B. "\( \alpha \leq x\) and/or \(x \leq \beta\)".
            </span>
            [[input:ans6]] [[validation:ans6]][[feedback:prt6]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Complete the table of values for the function \( f(x) = {@f@} \) for \(x\) between \(-3\) and \(3\).
            </span>
            <span class="multilang" lang="no">
                Fullfør tabellen av verdier for funksjonen \( f(x) = {@f@} \) for \(x\) mellom \(-3\) and \(3\).
            </span>
            <span class="multilang" lang="de">
                Vervollständigen Sie die Wertetabelle für die Funktion \( f(x) = {@f@} \) für \(x\) zwischen \(-3\) und \(3\).
            </span>
        </p>
        <table class="HELM_table">
            <tbody>
                <tr>
                    <th>
                        input, \(x\)
                    </th>
                    <th>
                        output, \(f(x)\)
                    </th>
                </tr>
                <tr>
                    <td>
                        \(-3\)
                    </td>
                    <td>
                        {@fatneg3@}
                    </td>
                </tr>
                <tr>
                    <td>
                        \(-2\)
                    </td>
                    <td>
                        {@fatneg2@}
                    </td>
                </tr>
                <tr>
                    <td>
                        \(-1\)
                    </td>
                    <td>
                        [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
                    </td>
                </tr>
                <tr>
                    <td>
                        \(0\)
                    </td>
                    <td>
                        [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
                    </td>
                </tr>
                <tr>
                    <td>
                        \(1\)
                    </td>
                    <td>
                        [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
                    </td>
                </tr>
                <tr>
                    <td>
                        \(2\)
                    </td>
                    <td>
                        [[input:ans5]] [[validation:ans5]][[feedback:prt5]]
                    </td>
                </tr>
                <tr>
                    <td>
                        \(3\)
                    </td>
                    <td>
                        {@fat3@}
                    </td>
                </tr>
            </tbody>
        </table>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                From the table of values, what is the shape of the function \( f(x) = {@f@} \)?
            </span>
            <span class="multilang" lang="no">
                Hva er formen på funksjonen fra verditabellen \( f(x) = {@f@} \)?
            </span>
            <span class="multilang" lang="de">
                Welche Form hat die Funktion \( f(x) = {@f@} \) aus der Wertetabelle?
            </span>
        </p>
        <p>
            [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
            <br/>
        </p>
        <p>
            <span class="multilang" lang="en">
                Choose from the graphs below:
            </span>
            <span class="multilang" lang="no">
                Velg fra grafene under:
            </span>
            <span class="multilang" lang="de">
                Wählen Sie aus den folgenden Grafiken:
            </span>
        </p>
        <ul>
        </ul>
        <table>
            <tbody>
                <tr>
                    <td style="text-align: center;">
                        <b>
                            <span class="multilang" lang="en">
                                Graph A
                            </span>
                            <span class="multilang" lang="no">
                                Graf A
                            </span>
                            <span class="multilang" lang="de">
                                Graph A
                            </span>
                        </b>
                    </td>
                    <td style="text-align: center;">
                        <b>
                            <span class="multilang" lang="en">
                                Graph B
                            </span>
                            <span class="multilang" lang="no">
                                Graf B
                            </span>
                            <span class="multilang" lang="de">
                                Graph B
                            </span>
                        </b>
                    </td>
                </tr>
                <tr>
                    <td style="text-align: center;">
                        <span class="multilang" lang="en">
                            {@thumbnailplot(funcstoplot[1],"Graph A")@}
                        </span>
                        <span class="multilang" lang="no">
                            {@thumbnailplot(funcstoplot[1],"Graf A")@}
                        </span>
                        <span class="multilang" lang="de">
                            {@thumbnailplot(funcstoplot[1],"Graph A")@}
                        </span>
                    </td>
                    <td style="text-align: center;">
                        <span class="multilang" lang="en">
                            {@thumbnailplot(funcstoplot[2],"Graph B")@}
                        </span>
                        <span class="multilang" lang="no">
                            {@thumbnailplot(funcstoplot[2],"Graf B")@}
                        </span>
                        <span class="multilang" lang="de">
                            {@thumbnailplot(funcstoplot[2],"Graph B")@}
                        </span>
                    </td>
                </tr>
                <tr>
                    <td style="text-align: center;">
                    </td>
                    <td style="text-align: center;">
                    </td>
                </tr>
                <tr>
                    <td style="text-align: center;">
                        <b>
                            <span class="multilang" lang="en">
                                Graph C
                            </span>
                            <span class="multilang" lang="no">
                                Graf C
                            </span>
                            <span class="multilang" lang="de">
                                Graph C
                            </span>
                        </b>
                    </td>
                    <td style="text-align: center;">
                        <b>
                            <span class="multilang" lang="en">
                                Graph D
                            </span>
                            <span class="multilang" lang="no">
                                Graf D
                            </span>
                            <span class="multilang" lang="de">
                                Graph D
                            </span>
                        </b>
                    </td>
                </tr>
                <tr>
                    <td style="text-align: center;">
                        <span class="multilang" lang="en">
                            {@thumbnailplot(funcstoplot[3],"Graph C")@}
                        </span>
                        <span class="multilang" lang="no">
                            {@thumbnailplot(funcstoplot[3],"Graf C")@}
                        </span>
                        <span class="multilang" lang="de">
                            {@thumbnailplot(funcstoplot[3],"Graph C")@}
                        </span>
                    </td>
                    <td style="text-align: center;">
                        <span class="multilang" lang="en">
                            {@thumbnailplot(funcstoplot[4],"Graph D")@}
                        </span>
                        <span class="multilang" lang="no">
                            {@thumbnailplot(funcstoplot[4],"Graf D")@}
                        </span>
                        <span class="multilang" lang="de">
                            {@thumbnailplot(funcstoplot[4],"Graph D")@}
                        </span>
                    </td>
                </tr>
            </tbody>
        </table>
    </li>
    <li>
        <span class="multilang" lang="en">
            Deduce the range of the function by inspecting the graph and table of values. Give your answer as an inequality in \(y\), e.g. "\( \alpha \leq y\) and \(y \leq \beta\)".
        </span>
        <span class="multilang" lang="no">
            Finn verdiområdet til funksjonen ved å inspisere grafen og verditabellen. Gi svaret ditt som en ulikhet i \(y\), f.eks. "\( \alpha \leq y\) and \(y \leq \beta\)".
        </span>
        <span class="multilang" lang="de">
            Leiten Sie den Bereich der Funktion ab, indem Sie das Diagramm und die Wertetabelle untersuchen. Geben Sie Ihre Antwort als Ungleichung in \(y\) an, z.B. "\( \alpha \leq y\) and \(y \leq \beta\)".
        </span>
        [[input:ans7]][[validation:ans7]][[feedback:prt7]]
    </li>
</ol>
<p>
    <br/>
</p>
<p>
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Since \( {@f@}\) can take any value of \(x\) in the given interval \({@leftend@} \leq x\leq {@rightend@}\) , the domain remains as \({@leftend@} \leq x\leq {@rightend@}\).
            </span>
            <span class="multilang" lang="no">
                Siden \( {@f@}\) kan ta hvilken som helst \(x\) verdi  i det gitte intervallet \({@leftend@} \leq x\leq {@rightend@}\) , er definisjonsmengden fortsatt \({@leftend@} \leq x\leq {@rightend@}\).
            </span>
            <span class="multilang" lang="de">
                Da \( {@f@}\) im gegebenen Intervall \({@leftend@} \leq x\leq {@rightend@}\) jeden Wert von \(x\) annehmen kann, bleibt die Definitionsmenge als \( {@leftend@} \leq x\leq {@rightend@}\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We fill in the values in the table by substituting the value of \(x\) into the function expression \(f(x) = {@f@}\). Hence:
            </span>
            <span class="multilang" lang="no">
                Vi fyller ut verdiene i tabellen ved å substituere \(x\) i funksjonsuttrykket \(f(x) = {@f@}\) med de gitte verdiene. Altså:
            </span>
            <span class="multilang" lang="de">
                Wir tragen die Werte in die Tabelle ein, indem wir den Wert von \(x\) in die Funktion \(f(x) = {@f@}\) einsetzen. Somit:
            </span>
        </p>
        <p>
            \(f(-1) = {@a@}(-1)^3 = {@ta2@}\),
        </p>
        <p>
            \(f(0) = {@a@}(0)^3 = {@ta3@}\),
        </p>
        <p>
            \(f(1) = {@a@}(1)^3 = {@ta4@}\)
            <span class="multilang" lang="en">
                and
            </span>
            <span class="multilang" lang="de">
                und
            </span>
        </p>
        <p>
            \(f(2) = {@a@}(2)^3 = {@ta5@}\).
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                To find the graph of \(f(x) = {@f@}\), you plot the \(x\) and \(f(x)\) points on a graph and pick the graph that is close to your sketch.
            </span>
            <span class="multilang" lang="no">
                For å finne grafen til \(f(x) = {@f@} \), plotter du punktene \(x \) og \(f(x) \) på en graf og velger diagrammet som er nærmest din skisse.
            </span>
            <span class="multilang" lang="de">
                Um den Graphen von \(f(x) = {@f@}\) zu finden, zeichnen Sie die Punkte \(x\) und \(f(x)\) in einen Graphen und wählen den Graphen aus, der am besten mit Ihrer Skizze übereinstimmt.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                From table and graph, the smallest value of \(f(x)\) occurs at \(x = 0\) which is \(f(0) = {@ta3@}\) and the biggest value at \(x=-3 \text{ or }3\), \(f(3) = {@fat3@}\), hence the range is \({@ta3@} \leq x\leq {@fat3@}\)
            </span>
            <span class="multilang" lang="no">
                Fra tabell og graf forekommer den minste verdien av \(f(x) \) ved \(x = 0\) som er \(f(0) = {@ta3@} \) og den største verdien forekommer ved \(x = -3 \text{or} 3 \), \(f (3) = {@fat3@} \), derfor er verdiområdet \({@ta3@} \leq x \leq {@fat3@} \)
            </span>
            <span class="multilang" lang="de">
                Aus Tabelle und Grafik ergibt sich der kleinste Wert von \(f(x)\) bei \(x = 0\), was \(f(0) = {@ta3@}\) ist, und der größte Wert bei \(x =-3 \text{ oder }3\), \(f(3) = {@fat3@}\), also ist der Bereich \({@ta3@} \leq x\leq {@fat3@}\)
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a: rand(5)+2
b:rand_with_prohib(2,6,[a])
f: expand(a*x^2+b)
f1: fullratsimp(f)
f(x):=a*x^2 +b
/* domain of function*/
leftend:-3;
rightend:3;
ta6: -3<= x nounand x<= 3;

fatneg3:ev(f,x=-3);
fatneg2: ev(f,x=-2);
ta2: ev(f,x=-1);
ta3: ev(f,x=0);
ta4: ev(f,x=1);
ta5: ev(f,x=2);
fat3: ev(f,x=3);

/* range */
ta7: ta3<=y nounand y<=fat3;

max: f(a/2);
min: f(a/3);

c_funcs: [f];
w_funcs: [-f, f-max, -f+max];
/* Choose functions to plot */
funcs: append(c_funcs, rand_selection(w_funcs,3)) 
/* permute the graphs but keep note of which ones are right! */
reordering:random_permutation([1,2,3,4]);
fntoletter:zip_with("[",reordering,[A,B,C,D]);
/* apply the reordering to produce the shuffled list of options */
funcstoplot: makelist(funcs[reordering[i]],i,1,4);
corr:maplist(lambda([ex],assoc(ex,fntoletter)), [1,2,3,4]);
ta1:maplist(lambda([ex],[ex, is(ex=assoc(1,fntoletter))]), [A,B,C,D]);
thumbnailplot:lambda([ex1,ex2],plot(ex1,[x,-5,5],[y,-30, 30],[box, false],[yx_ratio, 0.7],[xtics, -5, 1, 5],[ytics, -30, 15, 30],[axes, solid],[size,240,240],[plottags,false],[alt,ex2]));]]></text>
        </questionvariables>
        <specificfeedback format="moodle_auto_format">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>The domain of {@f@} is {@ta6@}, the missing table values are {@ta2@},{@ta3@},{@ta4@},{@ta5@},  the multiple-choice of graphs is {@funcstoplot@}, and the range is {@ta7@}.</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>ta5</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>ta6</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>solve</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans7</name>
            <type>algebraic</type>
            <tans>ta7</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>solve</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>corr[1]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0.1</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>ta6</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt7</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans7</sans>
                <tans>ta7</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt7-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1875490280</deployedseed>
        <deployedseed>1719349947</deployedseed>
        <deployedseed>93998742</deployedseed>
        <deployedseed>430777532</deployedseed>
        <deployedseed>1263631391</deployedseed>
        <deployedseed>1544963262</deployedseed>
        <deployedseed>646791225</deployedseed>
        <deployedseed>1933431540</deployedseed>
        <deployedseed>173438256</deployedseed>
        <deployedseed>1975031979</deployedseed>
        <deployedseed>468778757</deployedseed>
        <deployedseed>1557119625</deployedseed>
        <deployedseed>464609026</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>corr[1]</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta6</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>ta7</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-0-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt7-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>corr[2]</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2-1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3-1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4-1</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5-1</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value><![CDATA[leftend <=x]]></value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value><![CDATA[x<=fat3]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-0-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt7-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
