<quiz>
    <question type="stack">
        <name>
            <text>1.4.1.14 Vereinfachen von algebraischen Brüchen [Simplifying algebraic fractions]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Leave the expression in its most simplified form
    </span>
    <span class="multilang" lang="de">
        Bringen Sie den Ausdruck in seine einfachste Form
    </span>
</p>
<ol class="HELM_parts_inline">
    <li>
        \(\displaystyle{{@a1@}}\)
    </li>
    <li>
        \(\displaystyle{{@a2@}}\)
    </li>
    <li>
        \(\displaystyle{{@a3@}}\)
    </li>
    <li>
        \(\displaystyle{{@a4@}}\)
    </li>
</ol>
<p>
    <span class="multilang" lang="en">
        When simplifying remember only common factors can be cancelled.
    </span>
    <span class="multilang" lang="de">
        Beachten Sie beim Vereinfachen, dass nur gemeinsame Faktoren gestrichen werden können.
    </span>
</p>
<h4 class="HELM_solution">
    <span class="multilang" lang="en">
        Solution
    </span>
    <span class="multilang" lang="de">
        Lösung
    </span>
</h4>
<ol class="HELM_parts">
    <li>
        [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
    <li>
        [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            We have that \((x+1)\) is a common factor between the numerator and the denominator of \(\displaystyle{{@a1@}}\) and so we have:
        </span>
        <span class="multilang" lang="de">
            Wir haben festgestellt, dass \((x+1)\) ein gemeinsamer Faktor von dem Zähler und dem Nenner von \(\displaystyle{{@a1@}}\) ist, also erhalten wir:
        </span>
        <p>
            \(\displaystyle{{@a1@}} = \displaystyle{{@ta1@}}\)
        </p>
        <p>
            <span class="multilang" lang="en">
                There are no more common factors between the numerator and denominator and so this is the final answer
            </span>
            <span class="multilang" lang="de">
                Es gibt keine weiteren gemeinsamen Faktoren von Zähler und Nenner und daher ist dies die endgültige Antwort
            </span>
        </p>
    </li>
    <li>
        <span class="multilang" lang="en">
            If we first factorise the numerator and the denominator of \(\displaystyle{{@a2@}}\) we see that \((x+1)\) is again a common factor and so can be cancelled to give:
        </span>
        <span class="multilang" lang="de">
            Wenn wir zuerst den Zähler und den Nenner von \(\displaystyle{{@a2@}}\) faktorisieren, sehen wir, dass \((x+1)\) wieder ein gemeinsamer Faktor ist, und kann daher gestrichen werden zu:
        </span>
        <p>
            \(\displaystyle{{@a2@}} = \displaystyle{{@ta2@}}\)
        </p>
    </li>
    <li>
        <span class="multilang" lang="en">
            Here we can see that \({@ev(7-2*n3,simp)@}\) divides every term in the expression \(\displaystyle{{@a3@}}\) and so dividing the numerator and the denominator by \({@ev(7-2*n3,simp)@}\) gives us:
        </span>
        <span class="multilang" lang="de">
            Hier können wir sehen, dass \({@ev(7-2*n3,simp)@}\) jeden Term im Ausdruck \(\displaystyle{{@a3 @}}\) teilt und so dividieren wir Zähler und Nenner durch \({@ev(7-2*n3,simp)@}\) und erhalten:
        </span>
        <p>
            \(\displaystyle{{@a3@}} = \displaystyle{{@ta3@}}\)
        </p>
        <p>
            <span class="multilang" lang="en">
                Just as before this cannot be simplified further
            </span>
            <span class="multilang" lang="de">
                Das lässt sich wie zuvor nicht weiter vereinfachen
            </span>
        </p>
    </li>
    <li>
        <span class="multilang" lang="en">
            By factorising the numerator of \(\displaystyle{{@a4@}}\) we can see that the term \((x-{@ev(n4+2,simp)@})\) is a common factor between the numerator and the denominator and hence we have:
        </span>
        <span class="multilang" lang="de">
            Indem wir den Zähler von \(\displaystyle{{@a4@}}\) faktorisieren, sehen wir, dass der Term \((x-{@ev(n4+2, simp)@})\) ein gemeinsamer Teiler von Zähler und Nenner ist, und daher erhalten wir:
        </span>
        <p>
            \(\displaystyle{{@a4@}} = \displaystyle{{@ta4@}}\)
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>simp: true
n1: rand(2)+1
n2: rand(2)+1
n3: rand(2)+1
n4: rand(2)+1
simp:false

a1: [(x+1)/(2*(x+1)), (2*(x+1))/(x+1)][n1];
a2: [(x+1)/(2*x+2), (3*x+3)/(x+1)][n2];
a3: [(5*x-15)/5, (3*x+18)/3][n3];
a4: [(5*x-15)/(x-3),(4*x-16)/(x-4)][n4];

ta1: [1/2,2][n1];
ta2: [1/2,3][n2];
ta3: [x-3,x+6][n3];
ta4: [5,4][n4];</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(\displaystyle{{@a1@}} = {@ta1@}, \quad
\displaystyle{{@a2@}} = {@ta2@}, \quad
\displaystyle{{@a3@}} = {@ta3@}, \quad
\displaystyle{{@a4@}} = {@ta4@}\)</text>
        </questionnote>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>548062658</deployedseed>
        <deployedseed>1715528910</deployedseed>
        <deployedseed>1084160980</deployedseed>
        <deployedseed>445636837</deployedseed>
        <deployedseed>43140579</deployedseed>
        <deployedseed>1334302778</deployedseed>
        <deployedseed>2115030355</deployedseed>
        <deployedseed>809516604</deployedseed>
        <deployedseed>1131630747</deployedseed>
        <deployedseed>41005081</deployedseed>
        <deployedseed>123303805</deployedseed>
        <deployedseed>1191060345</deployedseed>
        <deployedseed>1867622767</deployedseed>
        <deployedseed>677025468</deployedseed>
        <deployedseed>1556459471</deployedseed>
        <deployedseed>1498153512</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1*1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2*1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
