<quiz>
    <question type="stack">
        <name>
            <text>4.2.5.4 Passende Wellengraphen [Wave graphs matching]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        For each of the functions, identify its graph.
    </span>
    <span class="multilang" lang="no">
        Identifiser grafen til hver av funksjonene.
    </span>
    <span class="multilang" lang="de">
        Identifizieren Sie für jede der Funktionen ihren Graphen.
    </span>
</p>
<p>
</p>
<p>
    \(A: y = \displaystyle{@pr_funcs[1]@}\)
</p>
<p>
    \(B: y = \displaystyle{@pr_funcs[2]@} \)
    <br/>
</p>
<p>
    \(C: y = \displaystyle{@pr_funcs[3]@} \)
    <br/>
</p>
<p>
    \(D: y = \displaystyle{@pr_funcs[4]@} \)
</p>
<br/>
<p>
</p>
<table>
    <tbody>
        <tr>
            <td>
                {@thumbnailplot(funcstoplot[1],"Function 1")@}
            </td>
            <td>
                {@thumbnailplot(funcstoplot[2],"Function 2")@}
            </td>
            <td>
                {@thumbnailplot(funcstoplot[3],"Function 3")@}
            </td>
            <td>
                {@thumbnailplot(funcstoplot[4],"Function 4")@}
            </td>
        </tr>
        <tr>
            <td>
                [[input:ans1]] [[validation:ans1]]
            </td>
            <td>
                [[input:ans2]] [[validation:ans2]]
            </td>
            <td>
                [[input:ans3]] [[validation:ans3]]
            </td>
            <td>
                [[input:ans4]] [[validation:ans4]]
            </td>
        </tr>
    </tbody>
</table>
<p>
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The general form of the sinusoidal function is \( y = A \sin(\omega t +\alpha), \omega&gt;0. \)
    </span>
    <span class="multilang" lang="no">
        Den generelle formen for den sinusformede funksjonen er \( y = A \sin(\omega t +\alpha), \omega&gt;0. \)
    </span>
    <span class="multilang" lang="de">
        Die allgemeine Form der Sinusfunktion ist \( y = A \sin(\omega t +\alpha), \omega&gt;0. \)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We are given the functions:
    </span>
    <span class="multilang" lang="no">
        Vi får funksjonene:
    </span>
    <span class="multilang" lang="de">
        Uns werden die Funktionen gegeben:
    </span>
</p>
<p>
</p>
<p>
    \(A: y = {@pr_funcs[1]@}\)
</p>
<p>
    \(B: y = {@pr_funcs[2]@} \)
    <br/>
</p>
<p>
    \(C: y = {@pr_funcs[3]@} \)
    <br/>
</p>
<p>
    \(D: y ={@pr_funcs[4]@} \)
</p>
<span class="multilang" lang="en">
    We start with the function \( y = {@a*sin(w*t)@} \) which has amplitude \( A = {@a@} \), \( \omega = {@w@}\) and period \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@}. \)
</span>
<span class="multilang" lang="no">
    Vi starter med funksjonen \( y = {@a*sin(w*t)@} \) som har amplitude \( A = {@a@} \), \( \omega = {@w@}\) og periode \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@}. \)
</span>
<span class="multilang" lang="de">
    Wir beginnen mit der Funktion \( y = {@a*sin(w*t)@} \) mit der Amplitude \( A = {@a@} \) , \( \omega = {@w@}\) und Periode \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@} . \)
</span>
<p>
    <span class="multilang" lang="en">
        The graph of \( y = {@a*sin(w*t+phi)@}\) is obtained by moving the graph of \( y = {@a*sin(w*t)@} \) horizontally to the left by {@phi@} units.
        <br/>
        The function \( y = {@(a/2)*sin(w*t)@} \) has \( A = {@a/2@} \), \( \omega = {@w@}\) and period \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@}. \)
    </span>
    <span class="multilang" lang="no">
        Grafen til \( y = {@a*sin(w*t+phi)@}\) fås ved å flytte grafen til \( y = {@a*sin(w*t)@} \) horisontalt til venstre med {@phi@} enheter.
        <br/>
        Funksjonen \( y = {@(a/2)*sin(w*t)@} \) har \( A = {@a/2@} \), \( \omega = {@w @}\) og periode \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@}. \)
    </span>
    <span class="multilang" lang="de">
        Der Graph von \( y = {@a*sin(w*t+phi)@}\) wird durch Verschieben des Graphen von \( y = {@a *sin(w*t)@} \) horizontal um {@phi@} Einheiten nach links erhalten.
        <br/>
        Die Funktion \( y = {@(a/2)*sin(w*t)@} \) hat \( A = {@a/2@} \), \( \omega = {@w @}\) und Periode \( T = \frac{2 \pi}{\omega} = {@T@} = {@dispdp(float(T),2)@}. \)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Finally, the function \( y = {@a*sin(w/2*t)@} \) has \( A = {@a@} \), \( \omega = {@w/2@}\) and period \( T = \frac{2 \pi}{\omega} = {@2*T@} = {@dispdp(float(2*T),2)@}. \)
    </span>
    <span class="multilang" lang="no">
        Til slutt har funksjonen \( y = {@a*sin(w/2*t)@} \) \( A = {@a@} \), \( \omega = {@w/2@}\) og periode \( T = \frac{2 \pi}{\omega} = {@2*T@} = {@dispdp(float(2*T),2)@}. \)
    </span>
    <span class="multilang" lang="de">
        Schließlich hat die Funktion \( y = {@a*sin(w/2*t)@} \) \( A = {@a@} \), \( \omega = {@w/2@}\) und Periode \( T = \frac{2 \pi}{\omega} = {@2*T@} = {@dispdp(float(2*T) ,2)@}. \)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a: rand([2,4,6,8])
w: rand(5)+2
phi: rand([%pi/4, %pi/3])
T: 2*%pi/w

funcs:[a*sin(w*t), a*sin(w*t+phi), a*sin(w/2*t), (a/2)*sin(w*t)];
/* Choose functions to plot */
funcstoplot:rand_selection(funcs,4);
/* permute the graphs but keep note of which ones are right! */
reordering:random_permutation([1,2,3,4]);
fntoletter:zip_with("[",reordering,[A,B,C,D]);
/* apply the reordering to produce the shuffled list of options */
pr_funcs: makelist(funcstoplot[reordering[i]],i,1,4);

corr:maplist(lambda([ex],assoc(ex,fntoletter)), [1,2,3,4]);
ta1:maplist(lambda([ex],[ex, is(ex=assoc(1,fntoletter))]), [A,B,C,D]);
ta2:maplist(lambda([ex],[ex, is(ex=assoc(2,fntoletter))]), [A,B,C,D]);
ta3:maplist(lambda([ex],[ex, is(ex=assoc(3,fntoletter))]), [A,B,C,D]);
ta4:maplist(lambda([ex],[ex, is(ex=assoc(4,fntoletter))]), [A,B,C,D]);
thumbnailplot:lambda([ex1,ex2],plot(ex1,[t,-2*T,2*T],[y,-a,a],[box, false],[yx_ratio, 1],[xtics, -T, 2*T, T],[ytics, -a, 2*a, a],[axes, solid],[size,170,170],[plottags,false],[alt,ex2]));]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]][[feedback:prt2]][[feedback:prt3]][[feedback:prt4]]</text>
        </specificfeedback>
        <questionnote>
            <text>{@funcstoplot@} {#fntoletter#}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>dropdown</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>dropdown</type>
            <tans>ta4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>corr[1]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0.1</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>corr[2]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-0-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0.1</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>corr[3]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>corr[4]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>120007461</deployedseed>
        <deployedseed>546535141</deployedseed>
        <deployedseed>494950851</deployedseed>
        <deployedseed>1008786367</deployedseed>
        <deployedseed>500393716</deployedseed>
        <deployedseed>1177638710</deployedseed>
        <deployedseed>277951245</deployedseed>
        <deployedseed>907812347</deployedseed>
        <deployedseed>1402846964</deployedseed>
        <deployedseed>474150818</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>corr[1]</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>corr[2]</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>corr[3]</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>corr[4]</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-0-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-0-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>first(mcq_incorrect(ta1))</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>first(mcq_incorrect(ta2))</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>first(mcq_incorrect(ta3))</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>first(mcq_incorrect(ta4))</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-0-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-0-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
