<quiz>
    <question type="stack">
        <name>
            <text>4.3.5.4 Multiwinkelformel [multi angle formula]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                Express \(\sin(3t)\) in terms of \(\sin t\).
            </span>
            <span class="multilang" lang="no">
                Uttrykk \(\sin(3t)\) i form av \(\sin t\).
            </span>
            <span class="multilang" lang="de">
                Drücken Sie \(\sin(3t)\) in Form von \(\sin t\) aus.
            </span>
        </p>
        <p>
            \(\sin(3t) =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Express \(\cos(3t)\) in terms of \(\cos t\).
            </span>
            <span class="multilang" lang="no">
                Uttrykk \(\cos(3t)\) i form av \(\cos t\).
            </span>
            <span class="multilang" lang="de">
                Drücken Sie \(\cos(3t)\) in Form von \(\cos t\) aus.
            </span>
        </p>
        <p>
            \(\cos(3t)=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                We use the sum and double angle formulas.
            </span>
            <span class="multilang" lang="no">
                Vi bruker sum- og dobbelvinkelformlene.
            </span>
            <span class="multilang" lang="de">
                Wir verwenden die Summen- und Doppelwinkelformeln.
            </span>
        </p>
        \[\begin{aligned} \sin 3\theta &amp;\equiv \sin (2\theta+\theta)=\sin 2\theta\cos\theta+\cos 2\theta\sin\theta \\ &amp;\equiv 2\sin\theta\cos^2\theta+(\cos^2\theta-\sin^2\theta)\sin\theta\\ &amp;\equiv 3\sin\theta\cos^2\theta-\sin^3\theta\\ &amp;\equiv
                                3\sin\theta(1-\sin^2\theta)-\sin^3\theta\equiv 3\sin\theta-4\sin^3\theta \end{aligned}\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We proceed simiarly.
            </span>
            <span class="multilang" lang="no">
                Vi fortsetter på samme måte.
            </span>
            <span class="multilang" lang="de">
                Wir gehen ähnlich vor.
            </span>
        </p>
        \[\begin{aligned} \cos 3\theta &amp;\equiv \cos (2\theta+\theta)\equiv \cos 2\theta\cos\theta-\sin 2\theta\sin\theta\\ &amp;\equiv (\cos^2\theta-\sin^2\theta)\cos\theta-2\sin\theta\cos\theta\sin\theta\\ &amp;\equiv \cos^3\theta-3\sin^2\theta\cos\theta\\
                                &amp;\equiv \cos^3\theta-3(1-\cos^2\theta)\cos\theta\\ &amp;\equiv 4\cos^3\theta-3\cos\theta \end{aligned}\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ta1 : 3*sin(t) - 4*sin(t)^3;
ta2 : 4*cos(t)^3 - 3*cos(t);

/*Writes the answers in terms of the arbitrary variable aa. In the prt we will substitute sin(t)=aa, and cos(t)=aa respectively, to ensure they don't use any other forms.*/
ta1_aa:3*aa-4*aa^3;
ta2_aa:4*aa^3-3*aa;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans1_aa:subst(sin(t)=aa,ans1)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    <span class="multilang" lang="en">
        Your answer is not equivalent to \(\sin(3t)\)
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist nicht äquivalent zu \(\sin(3t)\)
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1_aa</sans>
                <tans>ta1_aa</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your answer is equivalent to \(sin(3t)\), but you have not given it in the correct form.
</span>
<span class="multilang" lang="de">
    Ihre Antwort entspricht \(sin(3t)\), aber Sie haben sie nicht in der richtigen Form gegeben.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ans2_aa:subst(cos(t)=aa,ans2)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p dir="ltr" style="text-align: left;">
    <span class="multilang" lang="en">
        Your answer is not equivalent to \(\cos(3t)\)
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist nicht äquivalent zu \(\cos(3t)\)
    </span>
    <br/>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2_aa</sans>
                <tans>ta2_aa</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Your answer is equivalent to \(cos(3t)\), but you have not given it in the correct form.
</span>
<span class="multilang" lang="de">
    Ihre Antwort entspricht \(cos(3t)\), aber Sie haben sie nicht in der richtigen Form angegeben.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>sin(3*t)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>cos(3*t)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
