<quiz>
    <question type="stack">
        <name>
            <text>4.3.1.2 Erkunden Sie Pythagoras [Explore Pythagoras]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Using the exact values, evaluate \(\sin^2 \theta + \cos^2 \theta\) for the following angles:
    </span>
    <span class="multilang" lang="no">
        Bruk de eksakte verdiene, evaluer \(\sin^2 \theta + \cos^2 \theta\) for følgende vinkler:
    </span>
    <span class="multilang" lang="de">
        Unter Verwendung der genauen Werte, bestimmen Sie \(\sin^2 \theta + \cos^2 \theta\) für die folgenden Winkel:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                For
            </span>
            <span class="multilang" lang="de">
                Für
            </span>
            \(\theta = \displaystyle{{@a1@}}\),   \(\sin^2 \theta + \cos^2 \theta =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For
            </span>
            <span class="multilang" lang="de">
                Für
            </span>
            \(\theta = \displaystyle{{@a2@}}\),   \(\sin^2 \theta + \cos^2 \theta=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For
            </span>
            <span class="multilang" lang="de">
                Für
            </span>
            \(\theta = \displaystyle{{@a3@}}\),   \(\sin^2 \theta + \cos^2 \theta =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Using a table or otherwise, we can find the exact values of \(\sin \theta\) and \(\cos \theta\) for each of these angles.
    </span>
    <span class="multilang" lang="no">
        Ved å bruke en tabell eller på en annen måte kan vi finne de eksakte verdiene til \(\sin \theta\) og \(\cos \theta\) for hver av disse vinklene.
    </span>
    <span class="multilang" lang="de">
        Unter Verwendung einer Tabelle oder auf andere Weise können wir die genauen Werte von \(\sin \theta\) und \(\cos \theta\) für jeden dieser Winkel finden.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(\theta = {@a1@}\), we have \(\sin\theta = {@sin(a1)@}\) and \(\cos\theta = {@cos(a1)@}\). Then
            </span>
            <span class="multilang" lang="no">
                For \(\theta = {@a1@}\), har vi \(\sin\theta = {@sin(a1)@}\) og \(\cos\theta = {@cos(a1)@}\) så
            </span>
            <span class="multilang" lang="de">
                Für \(\theta = {@a1@}\) erhalten wir \(\sin\theta = {@sin(a1)@}\) und \(\cos\theta = {@cos(a1)@}\). Dann
            </span>
        </p>
        \[\sin^2 \theta + \cos^2 \theta = \left({@sin(a1)@}\right)^2 + \left({@cos(a1)@}\right)^2 = 1.\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(\theta = {@a2@}\), we have \(\sin\theta = {@sin(a2)@}\) and \(\cos\theta = {@cos(a2)@}\). Then
            </span>
            <span class="multilang" lang="no">
                For \(\theta = {@a2@}\), har vi \(\sin\theta = {@sin(a2)@}\) og \(\cos\theta = {@cos(a2)@}\) så
            </span>
            <span class="multilang" lang="de">
                Für \(\theta = {@a2@}\) erhalten wir \(\sin\theta = {@sin(a2)@}\) und \(\cos\theta = {@cos(a2)@}\). Dann
            </span>
        </p>
        \[\sin^2 \theta + \cos^2 \theta = \left({@sin(a2)@}\right)^2 + \left({@cos(a2)@}\right)^2 = 1.\]
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(\theta = {@a3@}\), we have \(\sin\theta = {@sin(a3)@}\) and \(\cos\theta = {@cos(a3)@}\). Then
            </span>
            <span class="multilang" lang="no">
                For \(\theta = {@a3@}\), har vi \(\sin\theta = {@sin(a3)@}\) og \(\cos\theta = {@cos(a3)@}\) så
            </span>
            <span class="multilang" lang="de">
                Für \(\theta = {@a3@}\) erhalten wir \(\sin\theta = {@sin(a3)@}\) und \(\cos\theta = {@cos(a3)@}\). Dann
            </span>
        </p>
        \[\sin^2 \theta + \cos^2 \theta = \left({@sin(a3)@}\right)^2 + \left({@cos(a3)@}\right)^2 = 1.\]
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>[a1, a2, a3] : rand_selection([pi/6, pi/4, pi/3, pi/2, 2*pi/3, 3*pi/4, 5*pi/6],3);
ta1 : 1;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(\theta = {@a1@}, {@a2@}, {@a3@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1429169204</deployedseed>
        <deployedseed>282609484</deployedseed>
        <deployedseed>403995505</deployedseed>
        <deployedseed>676528546</deployedseed>
        <deployedseed>1049514125</deployedseed>
        <deployedseed>424036899</deployedseed>
        <deployedseed>1628645030</deployedseed>
        <deployedseed>313226745</deployedseed>
        <deployedseed>867170503</deployedseed>
        <deployedseed>376361763</deployedseed>
        <deployedseed>617173558</deployedseed>
        <deployedseed>1971870891</deployedseed>
        <deployedseed>514798907</deployedseed>
        <deployedseed>1121851772</deployedseed>
        <deployedseed>78633435</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>a1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>a2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>a3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
