<quiz>
    <question type="stack">
        <name>
            <text>2.7.2.6 Die Pole rationaler Funktionen [The poles of rational functions]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h4 class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</h4>
<p>
    <span class="multilang" lang="en">
        It is particularly important for engineers to find values of the independent variable for which the denominator is zero. These values are known as the
        <b>
            poles
        </b>
        of the rational function.
    </span>
    <span class="multilang" lang="no">
        Det er spesielt viktig for ingeniører å finne verdier for den uavhengige variabelen der nevneren er null. Disse verdiene er kjent som
        <b>
            polene
        </b>
        til den rasjonelle funksjonen.
    </span>
    <span class="multilang" lang="de">
        Für Ingenieure ist es besonders wichtig, Werte der unabhängigen Variablen zu finden, deren Nenner Null ist. Diese Werte werden als
        <b>
            Pole
        </b>
        der rationalen Funktion bezeichnet.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Find the poles of the following rational functions (in case the function has more than one pole, write the solution in the form
        <code>
            {pole1, pole2,...}
        </code>
        , including the curly braces):
    </span>
    <span class="multilang" lang="no">
        Finn polene til følgende rasjonelle funksjoner (hvis funksjonen har mer enn én pol, skriv løsningen på formen
        <code>
            {pol1, pol2,...}
        </code>
        , inkludert krøllete klammeparenteser):
    </span>
    <span class="multilang" lang="de">
        Finden Sie die Pole der folgenden rationalen Funktionen (falls die Funktion mehr als einen Pol hat, schreiben Sie die Lösung in der Form
        <code>
            {pole1, pole2,...}
        </code>
        , einschließlich der geschweiften Klammern):
    </span>
</p>
<br/>
<ol class="HELM_parts">
    <li>
        \(\displaystyle{@f1@}\)
        <span class="multilang" lang="en">
            has poles
        </span>
        <span class="multilang" lang="no">
            har poler
        </span>
        <span class="multilang" lang="de">
            hat die Pole
        </span>
        \(t=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]].
        <br/>
    </li>
    <li>
        \(\displaystyle{@f2@}\)
        <span class="multilang" lang="en">
            has poles
        </span>
        <span class="multilang" lang="no">
            har poler
        </span>
        <span class="multilang" lang="de">
            hat die Pole
        </span>
        \(s=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]].
        <br/>
    </li>
    <li>
        \(\displaystyle{@f3@}\)
        <span class="multilang" lang="en">
            has poles
        </span>
        <span class="multilang" lang="no">
            har poler
        </span>
        <span class="multilang" lang="de">
            hat die Pole
        </span>
        \(x=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]].
        <br/>
    </li>
    <li>
        \(\displaystyle{@f4@}\)
        <span class="multilang" lang="en">
            has poles
        </span>
        <span class="multilang" lang="no">
            har poler
        </span>
        <span class="multilang" lang="de">
            hat die Pole
        </span>
        \(x=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]].
        <br/>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The poles of a rational function are the points for which the denominator becomes zero and, hence, the function diverges.
    </span>
    <span class="multilang" lang="no">
        Polene til en rasjonell funksjon er punktene der nevneren blir null, og dermed, divergerer funksjonen.
    </span>
    <span class="multilang" lang="de">
        Die Pole einer rationalen Funktion sind die Punkte, an denen der Nenner Null wird und die Funktion daher divergiert.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The poles of the rational functions given are:
    </span>
    <span class="multilang" lang="no">
        Polene til de gitte rasjonelle funksjonene er:
    </span>
    <span class="multilang" lang="de">
        Die Pole der gegebenen rationalen Funktionen sind:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            \(\displaystyle{@ta1@}\), in fact for this value of \(t\) we have that the denominator is zero: \((t-{@rand2@})=({@rand2@}-{@rand2@})=0\)
        </span>
        <span class="multilang" lang="no">
            \(\displaystyle{@ta1@}\), for denne verdien av \(t\) har vi at nevneren er null: \((t-{@rand2@})=({@rand2@}- {@rand2@})=0\)
        </span>
        <span class="multilang" lang="de">
            \(\displaystyle{@ta1@}\), denn der Nenner ist für diesen Wert von \(t\) Null: \((t-{@rand2@})=({@rand2@}-{@rand2@})=0\)
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            \(\displaystyle{@ta2@}\), in fact if we substitute in the denominator the value \(s={@rand4@}\) or \(s=-{@rand4@}\), the denominator is zero.
        </span>
        <span class="multilang" lang="no">
            \(\displaystyle{@ta2@}\), hvis vi erstatter verdien \(s={@rand4@}\) eller \(s=-{@rand4@}\) i nevneren, er nevneren null.
        </span>
        <span class="multilang" lang="de">
            \(\displaystyle{@ta2@}\), denn wenn wir im Nenner den Wert \(s={@rand4@}\) oder \(s=-{@rand4@}\) haben, ist er Null.
        </span>
        <br/>
    </li>
    <li>
        <span class="multilang" lang="en">
            \(\displaystyle{@ta3@}\), if we substitute in the denominator \(x=-{@rand7@}\) or \(x=-({@rand8@})\), the denominator becomes zero and then these two points represent the poles of our function.
        </span>
        <span class="multilang" lang="no">
            \(\displaystyle{@ta3@}\), hvis vi i nevneren erstatter \(x=-{@rand7@}\) eller \(x=-({@rand8@})\), blir nevneren null og så representerer disse to punktene polene til funksjonen vår.
        </span>
        <span class="multilang" lang="de">
            \(\displaystyle{@ta3@}\), wenn wir im Nenner \(x=-{@rand7@}\) oder \(x=-({@rand8@})\) einsetzen, wird der Nenner Null und dann stellen diese beiden Punkte die Pole unserer Funktion dar.
        </span>
        <br/>
    </li>
    <li>
        <span class="multilang" lang="en">
            \(\displaystyle{@ta4@}\), same reasons as for the previous cases.
        </span>
        <span class="multilang" lang="no">
            \(\displaystyle{@ta4@}\), samme grunner som for de tidligere tilfellene.
        </span>
        <span class="multilang" lang="de">
            \(\displaystyle{@ta4@}\), dieselben Gründe wie bei den vorherigen Fällen.
        </span>
        <br/>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var1: t;
var2: s;
var3: x;
rand1: rand_with_prohib(-3,5,[0]);
rand2: rand_with_step(1, 8, 1);
f1: f(t)=(var1+rand1)/(var1-rand2);
rand3:rand_with_prohib(-9,9,[0]);
rand4: rand_with_step(1,9,1);
f2: F(s)=(var2+rand3)/((var2+rand4)*(var2-rand4));
rand5: rand_with_prohib(-3, 3,[0]);
rand6: rand_with_step(1,6,1);
rand7:rand_with_step(1,10,1);
rand8: rand_with_prohib(-4,4,[0]);
f3: r(x)=(rand5*var3+rand6)/((var3+rand7)*(var3+rand8));
rand9: rand_with_step(1,4,1);
f4: f(x)=(var3-rand9^2)/(var3^2-rand9^2);

ta1:rand2;
ta2: {rand4,-rand4};
ta3: {-rand7,-rand8};
ta4:{rand9, -rand9};
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[{@f1@}, t = {@ta1@}\]
\[{@f2@}, t = {@ta2@}\]
\[{@f3@}, t = {@ta3@}\]
\[{@f4@}, t = {@ta4@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>ta4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1253366826</deployedseed>
        <deployedseed>958387628</deployedseed>
        <deployedseed>1573428619</deployedseed>
        <deployedseed>1477964473</deployedseed>
        <deployedseed>672035810</deployedseed>
        <deployedseed>365101807</deployedseed>
        <deployedseed>1349544365</deployedseed>
        <deployedseed>1393537675</deployedseed>
        <deployedseed>1869079653</deployedseed>
        <deployedseed>1039398906</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(rand2+5,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev({rand4+5,-rand4+5},simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ev({-rand7+5,-rand8+5},simp)</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ev({rand9+5, -rand9+5},simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
