<quiz>
    <question type="stack">
        <name>
            <text>2.7.1.7 ist p(z) konstant mit Variable z [is p(z) constant with variable z]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Is \({@Pz@}\) a polynomial of degree \(0\) with independent variable \(z\)?
    </span>
    <span class="multilang" lang="no">
        Er \({@Pz@}\) en polynom av grad \(0\) med uavhengig variabel \(z\)?
    </span>
    <span class="multilang" lang="de">
        Ist \({@Pz@}\) ein Polynom vom Grad \(0\) mit unabhängiger Variable \(z\)?
    </span>
    [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The constant function can always be considered as a polynomial of degree zero because it is a function of the form \(P(z)={@random1@}z^0\).
    </span>
    <span class="multilang" lang="no">
        Den konstanten funksjonen kan alltid defineres som en polynom med grad \(0\) siden da vil den være på formen \(P(z)={@random1@}z^0\).
    </span>
    <span class="multilang" lang="de">
        Die konstante Funktion kann immer als Polynom vom Grad Null betrachtet werden, da sie eine Funktion der Form \(P(z)={@random1@}z^0\) ist.
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Note that since we have \(P(z)\) then we know that \(z\) is the independent variable.
    </span>
    <span class="multilang" lang="no">
        Merk at siden vi har \(P(z)\), så vet vi at \(z\) er den uavhengige variabelen.
    </span>
    <span class="multilang" lang="de">
        Beachten Sie, dass wir wissen, da wir \(P(z)\) haben, dass \(z\) die unabhängige Variable ist.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[random1: rand_with_step(-13,13,1);
Pz: P(z)=random1;
ta1: [["Yes",true],["No",false],["Not enough information is given",false]];
]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>{@Pz@} {@mcq_correct(ta1)@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans1]</sans>
                <tans>mcq_correct(ta1)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1233952342</deployedseed>
        <deployedseed>1134830515</deployedseed>
        <deployedseed>999265101</deployedseed>
        <deployedseed>895796710</deployedseed>
        <deployedseed>1835588678</deployedseed>
        <deployedseed>201308534</deployedseed>
        <deployedseed>1493689188</deployedseed>
        <deployedseed>1234977607</deployedseed>
        <deployedseed>582534369</deployedseed>
        <deployedseed>658484283</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["Yes"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["No"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
