<quiz>
    <question type="stack">
        <name>
            <text>2.7.2.9 Was ist ein Pol [What is a pole]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h4 class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</h4>
<p>
    <span class="multilang" lang="en">
        What is the pole of a rational function?
    </span>
    <span class="multilang" lang="no">
        Hva er polen til en rasjonell funksjon?
    </span>
    <span class="multilang" lang="de">
        Was ist der Pol einer rationalen Funktion?
    </span>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        It is not necessary to repeat what a pole of a rational function is at this stage. However, we can clarify some concept appearing between the options. An extremum of a function is any point at which the value of a function is largest (a maximum) or smallest (a minimum). The line tangent to the curve at a point at infinity is called asymptote.
    </span>
    <span class="multilang" lang="no">
        Det er ikke nødvendig å gjenta hva en pol til en rasjonell funksjon er på dette stadiet. Imidlertid kan vi avklare et konsept som vises mellom alternativene. Et ekstremalpunkt av en funksjon er et hvilket som helst punkt der verdien av en funksjon er størst (et maksimum) eller minste (et minimum). Linjen som tangerer kurven i et uendelig punkt kalles asymptote.
    </span>
    <span class="multilang" lang="de">
        Es ist an dieser Stelle nicht notwendig zu wiederholen, was ein Pol einer rationalen Funktion ist. Wir können jedoch einige Konzepte klären, die zwischen den Optionen erscheinen. Ein Extremum einer Funktion ist jeder Punkt, an dem der Wert einer Funktion am größten (ein Maximum) oder am kleinsten (ein Minimum) ist. Die Tangente an die Kurve in einem unendlichen Punkt wird als Asymptote bezeichnet.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[ta1: [["It is a value of the independent variable which makes the numerator zero", false],["It is a value of the independent variable which makes the denominator zero", true], ["It is a line which is tangent to the curve at a point at infinity", false], ["It is an extremum of the function", false]];
]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans1]</sans>
                <tans>mcq_correct(ta1)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["It is a value of the independent variable which makes the denominator zero"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["It is a value of the independent variable which makes the numerator zero"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["It is a line which is tangent to the curve at a point at infinity"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["It is an extremum of the function"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
