<quiz>
    <question type="stack">
        <name>
            <text>2.7.1.3 Polynomgrad oder nicht polynom [Polynomial degree or not poly]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Select the correct description associated with each of the following functions.
    </span>
    <span class="multilang" lang="no">
        Velg riktig beskrivelse knyttet til hver av de følgende funksjonene.
    </span>
    <span class="multilang" lang="de">
        Wählen Sie die richtige Beschreibung für jede der folgenden Funktionen aus.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        \({@f1@}\)  [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \({@f2@}\)  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        \(\displaystyle{@f3@}\)  [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
    <li>
        \({@f4@}\)  [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
    </li>
    <li>
        \(\displaystyle{@f5@}\)  [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A polynomial is a function of the form \(P_n=a_n x^n+a_{n-1} x^{n-1}+...+a_2 x^2+a_1 x+a_0\); its degree is given by the highest of the degrees of the polynomial's monomials with non-zero coefficients.
    </span>
    <span class="multilang" lang="no">
        Et polynom er en funksjon av formen \(P_n=a_n x^n+a_{n-1} x^{n-1}+...+a_2 x^2+a_1 x+a_0\); graden er gitt av den høyeste eksponenten i uttrykket.
    </span>
    <span class="multilang" lang="de">
        Ein Polynom ist eine Funktion der Form \(P_n=a_n x^n+a_{n-1} x^{n-1}+...+a_2 x ^2+a_1x+a_0\); sein Grad wird durch den höchsten Grad der Monome des Polynoms mit Nicht-Null-Koeffizienten angegeben.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For the given functions:
    </span>
    <span class="multilang" lang="no">
        For en gitt funksjon:
    </span>
    <span class="multilang" lang="de">
        Für die angegebenen Funktionen:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                \({@f1@}:\) {@TA1@}, in fact, \(4\) represents the highest power of \(x\).
            </span>
            <span class="multilang" lang="no">
                \({@f1@}:\) {@TA1@} sin høyeste eksponent i uttrykket er \(4\). Graden til funksjon er derfor 4.
            </span>
            <span class="multilang" lang="de">
                \({@f1@}:\) {@TA1@}, tatsächlich repräsentiert \(4\) die höchste Potenz von \(x\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \({@f2@}:\) {@TA2@}; in this case, we have the constant function, so we can interpret it as \(h(x)={@rand1@}x^0\) and we can see explicitly that the degree is zero.
            </span>
            <span class="multilang" lang="no">
                \({@f2@}:\) {@TA2@}; i dette tilfellet har vi en konstant funksjon, så vi kan tolke den som \(h(x)={@rand1@}x^0\). Graden er derfor 0.
            </span>
            <span class="multilang" lang="de">
                \({@f2@}:\) {@TA2@}; in diesem Fall haben wir die konstante Funktion, also können wir sie als \(h(x)={@rand1@}x^0\) interpretieren und wir können explizit sehen, dass der Grad Null ist.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \(\displaystyle{@f3@}:\) {@TA3@}.This function does not represent a polynomial because of the inverse powers of x, \(\dfrac{{@coeffg1@}}{{@var1^2@}}\) and \(\dfrac{{@coeffg2@}}{{@var1@}}\).
            </span>
            <span class="multilang" lang="no">
                \(\displaystyle{@f3@}:\) {@TA3@}. Denne funksjonen representerer ikke et polynom på grunn av de negative potensene til x, \(\dfrac{{@coeffg1@}}{{@var1^2 @}}\) og \(\dfrac{{@coeffg2@}}{{@var1@}}\).
            </span>
            <span class="multilang" lang="de">
                \(\displaystyle{@f3@}:\) {@TA3@}. Diese Funktion stellt wegen der inversen Potenzen von x \(\dfrac {{@coeffg1@}}{{@var1^2@}}\) und \(\dfrac{{@coeffg2@}}{{@var1@}}\), kein Polynom dar.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \({@f4@}:\) {@TA4@}. The highest of the degrees of the monomials is \(3\), the dominant monomial is \({@coefft3@}{@var2@}^3\).
            </span>
            <span class="multilang" lang="no">
                \({@f4@}:\) {@TA4@}. Den høyeste av eksponenten i funksjonen er \(3\), det dominerende leddet er derfor \({@coefft3@}{@var2@}^3\).
            </span>
            <span class="multilang" lang="de">
                \({@f4@}:\) {@TA4@}. Der höchste Grad der Monome ist \(3\), das dominante Monom ist \({@coefft3@}{@var2@}^3\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \(\displaystyle{@f5@}:\) {@TA5@}. Same case as \({@f2@}\).
            </span>
            <span class="multilang" lang="no">
                \(\displaystyle{@f5@}:\) {@TA5@}. Samme tilfelle som \({@f2@}\).
            </span>
            <span class="multilang" lang="de">
                \(\displaystyle{@f5@}:\) {@TA5@}. Gleicher Fall wie \({@f2@}\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[var1: x;
var2: t;
coeffx4: rand_with_prohib(-5,5,[0]);
coeffx3: rand_with_prohib(-9,9,[0]);
coeffx2: rand_with_prohib(-2,9,[0]);
f1: f(x)=coeffx2*var1^2+coeffx3*var1^3+coeffx4*var1^4;
coefft3: rand_with_prohib(-7,5,[0]);
coefft2: rand_with_prohib(-3,13,[0]);
rand0: rand_with_prohib(-9,9,[0]);
f4: f(t)=coefft3*var2^3+coefft2*var2^2+rand0;
coeffg1: rand_with_step(1,10,1);
coeffg2: rand_with_prohib(-3,9,[0]);
f3: g(x)=coeffg1/var1^2+coeffg2/var1;
rand1: rand_with_step(-10,10,1);
f2: h(x)=rand1;
rand2: rand_with_prohib(-8, 8,[0]);
f5: k(x)=1/rand2;

ta1:[["Not a polynomial",false],["Polynomial of degree 0 (constant)", false],["Polynomial of degree 1 (linear)",false],["Polynomial of degree 2 (quadratic)", false], ["Polynomial of degree 3 (cubic)", false],["Polynomial of degree 4 (quartic)", true]];
ta4:[["Not a polynomial",false],["Polynomial of degree 0 (constant)", false],["Polynomial of degree 1 (linear)",false],["Polynomial of degree 2 (quadratic)", false], ["Polynomial of degree 3 (cubic)", true],["Polynomial of degree 4 (quartic)", false]];
ta3:[["Not a polynomial",true],["Polynomial of degree 0 (constant)", false],["Polynomial of degree 1 (linear)",false],["Polynomial of degree 2 (quadratic)", false], ["Polynomial of degree 3 (cubic)", false],["Polynomial of degree 4 (quartic)", false]];
ta2:[["Not a polynomial",false],["Polynomial of degree 0 (constant)", true],["Polynomial of degree 1 (linear)",false],["Polynomial of degree 2 (quadratic)", false], ["Polynomial of degree 3 (cubic)", false],["Polynomial of degree 4 (quartic)", false]];
ta5:[["Not a polynomial",false],["Polynomial of degree 0 (constant)", true],["Polynomial of degree 1 (linear)",false],["Polynomial of degree 2 (quadratic)", false], ["Polynomial of degree 3 (cubic)", false],["Polynomial of degree 4 (quartic)", false]];

TA1: "Polynomial of degree 4 (quartic)";
TA4: "Polynomial of degree 3 (cubic)";
TA3: "Not a polynomial";
TA2: "Polynomial of degree 0 (constant)";
TA5: "Polynomial of degree 0 (constant)";
]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text><![CDATA[<p>\({@f1@} \)</p>
<p>\({@f2@} \)</p>
<p>\({@f3@} \)</p>
<p>\({@f4@} \)</p>
<p>\({@f5@} \)</p>]]></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>dropdown</type>
            <tans>ta2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>ta3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>dropdown</type>
            <tans>ta4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>dropdown</type>
            <tans>ta5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans1]</sans>
                <tans>mcq_correct(ta1)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans2]</sans>
                <tans>mcq_correct(ta2)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans3]</sans>
                <tans>mcq_correct(ta3)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans4]</sans>
                <tans>mcq_correct(ta4)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans5]</sans>
                <tans>mcq_correct(ta5)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>818747097</deployedseed>
        <deployedseed>2098313585</deployedseed>
        <deployedseed>1186605486</deployedseed>
        <deployedseed>395995267</deployedseed>
        <deployedseed>680737032</deployedseed>
        <deployedseed>1383827447</deployedseed>
        <deployedseed>1228868470</deployedseed>
        <deployedseed>1295517051</deployedseed>
        <deployedseed>933462638</deployedseed>
        <deployedseed>315600800</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["Polynomial of degree 4 (quartic)"]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA["Polynomial of degree 0 (constant)"]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["Not a polynomial"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value><![CDATA["Polynomial of degree 3 (cubic)"]]></value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value><![CDATA["Polynomial of degree 0 (constant)"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["Not a polynomial"]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA["Not a polynomial"]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["Polynomial of degree 0 (constant)"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value><![CDATA["Not a polynomial"]]></value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value><![CDATA["Not a polynomial"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["Polynomial of degree 0 (constant)"]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA["Polynomial of degree 2 (quadratic)"]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["Polynomial of degree 1 (linear)"]]></value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value><![CDATA["Polynomial of degree 0 (constant)"]]></value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value><![CDATA["Polynomial of degree 3 (cubic)"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
