<quiz>
    <question type="stack">
        <name>
            <text>2.7.2.10 1-1 oder viele-1 [1-1 or many-1]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h4 class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</h4>
<p>
    <span class="multilang" lang="en">
        Referring to the functions \(R_1(x)=\dfrac{x+6}{x^2+1}\), \(R_2(t)=\dfrac{t^3-1}{2t+3}\) and \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\) state if they are one-to one or many-to-one.
    </span>
    <span class="multilang" lang="no">
        Med tanke på funksjonene \(R_1(x)=\dfrac{x+6}{x^2+1}\), \(R_2(t)=\dfrac{t^3-1}{2t+3}\) og \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\) angi om de er en-til-en eller mange-til-en.
    </span>
    <span class="multilang" lang="de">
        Bezogen auf die Funktionen \(R_1(x)=\dfrac{x+6}{x^2+1}\), \(R_2(t)=\dfrac {t^3-1}{2t+3}\) und \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\) geben Sie an, ob sie Eins-zu-Eins oder Viele-zu-Eins Funktionen sind.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        \(R_1(x)=\dfrac{x+6}{x^2+1}\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \(R_2(t)=\dfrac{t^3-1}{2t+3}\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
    <li>
        \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A function is said to be many-to-one if two or more elements in the domain of the function correspond to the same element in the range of the function. You can check if the function is many-to-one using the graphical method we explain in the previous exercises, or you can proceed analytically.
    </span>
    <span class="multilang" lang="no">
        En funksjon sies å være mange-til-en hvis to eller flere elementer i funksjonens definisjonsmengde tilsvarer det samme elementet i funksjonens verdimområde. Du kan sjekke om funksjonen er mange-til-en ved å bruke den grafiske metoden vi forklarte i de forrige øvingene, eller du kan gå videre analytisk.
    </span>
    <span class="multilang" lang="de">
        Eine Funktion wird Viele-zu-Eins genannt, wenn zwei oder mehr Elemente in der Definitionsmenge der Funktion demselben Element in der Zielmenge der Funktion entsprechen. Sie können überprüfen, ob die Funktion Viele-zu-Eins ist, indem Sie die grafische Methode verwenden, die wir in den vorherigen Übungen erklärt haben, oder Sie können analytisch vorgehen.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Analytically:
    </span>
    <span class="multilang" lang="no">
        Analytisk:
    </span>
    <span class="multilang" lang="de">
        Analytisch:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            \(R_1(x)=\dfrac{x+6}{x^2+1}\): consider the value \(R_1(x)=4\), there are two values of \(x\) for which this function assumes the value \(4\), in fact, substituting, we have \(4=\dfrac{x+6}{x^2+1}\), so we have to solve the second order equation \(4x^2-x-2=0\) which has two dinstict real solutions; hence, we can conclude that for two points of the domain corresponds the same value of the function.
        </span>
        <span class="multilang" lang="no">
            \(R_1(x)=\dfrac{x+6}{x^2+1}\): med \(R_1(x)=4\), er det er to verdier av \(x\) som gir funksjonsverdien \(4\). Setter vi \(4=\dfrac{x+6}{x^2+1}\), må vi løse andreordensligningen \( 4x^2-x-2=0\) som har to forskjellige reelle løsninger; derfor kan vi konkludere med at for to punkter i definisjonsmengden tilsvarer samme verdi av funksjonen.
        </span>
        <span class="multilang" lang="de">
            \(R_1(x)=\dfrac{x+6}{x^2+1}\): Betrachten Sie den Wert \(R_1(x)=4\), gibt es zwei Werte von \(x\), für die diese Funktion den Wert \(4\) annimmt, tatsächlich erhalten wir durch Einsetzen \(4=\dfrac{x+6}{x^2+1}\), also müssen wir die Gleichung zweiter Ordnung \(4x^2-x-2=0\) lösen, die zwei verschiedene reelle Lösungen hat; Daraus können wir schließen, dass für zwei Punkte der Definitionsmenge derselbe Wert der Funktion entspricht.
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            \(R_2(t)=\dfrac{t^3-1}{2t+3}\): you can apply the same arguments we use for \(R_1(x)\).
        </span>
        <span class="multilang" lang="no">
            \(R_2(t)=\dfrac{t^3-1}{2t+3}\): du kan bruke de samme argumentene vi bruker for \(R_1(x)\).
        </span>
        <span class="multilang" lang="de">
            \(R_2(t)=\dfrac{t^3-1}{2t+3}\): Sie können dieselben Argumente anwenden, die wir für \(R_1( x)\) haben.
        </span>
    </li>
    <li>
        <span class="multilang" lang="en">
            \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\): you can apply the same arguments as for \(R_1(x)\) (and \(R_2(x)\)).
        </span>
        <span class="multilang" lang="no">
            \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\): du kan bruke de samme argumentene som \(R_1(x)\) (og \( R_2(x)\)).
        </span>
        <span class="multilang" lang="de">
            \(R_3(z)=\dfrac{2z^2+z-1}{z^2+z-2}\):Sie können dieselben Argumente anwenden wie für \(R_1(x)\) (und \(R_2(x)\))..
        </span>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[ta1: [["One-to-one", false], ["Many-to-one", true]];
ta2: [["One-to-one", false], ["Many-to-one", true]];
ta3: [["One-to-one", false], ["Many-to-one", true]];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>wrd</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>ta1</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>dropdown</type>
            <tans>ta2</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>ta3</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans1]</sans>
                <tans>mcq_correct(ta1)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans2]</sans>
                <tans>mcq_correct(ta2)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>[ans3]</sans>
                <tans>mcq_correct(ta3)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["Many-to-one"]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA["Many-to-one"]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["Many-to-one"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value><![CDATA["One-to-one"]]></value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value><![CDATA["One-to-one"]]></value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value><![CDATA["One-to-one"]]></value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
