<quiz>
    <question type="stack">
        <name>
            <text>8.1.1.3 Solving two equations in two unknowns</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Use Cramer's rule to find the solution of the system:
    </span>
    <span class="multilang" lang="de">
        Verwenden Sie die Cramer-Regel, um die Lösung des Systems zu finden:
    </span>
    <br/>
</p>
<p>
    \[ {@eqns[1]@}\ =\ {@bmat[1,1]@} , \] \[ {@eqns[2]@}\ =\ {@bmat[2,1]@} , \]
    <br/>
    \(x=\) [[input:ans1]] [[validation:ans1]]
    <br/>
    \(y=\) [[input:ans2]] [[validation:ans2]]
    <br/>
    [[feedback:Result]]
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Cramer's rule states that \(x = \frac{\Delta_{x}}{\Delta}\) where \(\Delta\) is the determinant of the matrix and \(\Delta_{x}\) is the determinant of the matrix with the x column replaced by the numbers on the right hand side of the equations and similarly
        \(y = \frac{\Delta_{y}}{\Delta}\).
    </span>
    <span class="multilang" lang="de">
        Cramers Regel besagt, dass \(x = \frac{\Delta_{x}}{\Delta}\) wobei \(\Delta\) die Determinante der Matrix ist und \(\Delta_{x}\) die Determinante der Matrix ist, bei der die Spalte x durch die Zahlen auf der rechten Seite der Gleichungen ersetzt wird und genauso bei
        \(y = \frac{\Delta_{y}}{\Delta}\).
    </span>
    <br/>
</p>
<p>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore for the system,
        <br/>
        \[{@eqns[1]@}\ =\ {@bmat[1,1]@} , \] \[ {@eqns[2]@}\ =\ {@bmat[2,1]@} , \]
        <br/>
        \[\Delta = \det{@Amat@}={@D@},\]
        <br/>
        \[\Delta_{x} = \det{@Amatx@}={@Dx@},\] and
        <br/>
        \[\Delta_{y} = \det{@Amaty@}={@Dy@},\]
    </span>
    <span class="multilang" lang="de">
        Daher gilt für das System
        <br/>
        \[{@eqns[1]@}\ =\ {@bmat[1,1]@} , \] \[ {@eqns[2]@}\ =\ {@bmat[2,1]@} , \]
        <br/>
        \[\Delta = \det{@Amat@}={@D@},\]
        <br/>
        \[\Delta_{x} = \det{@Amatx@}={@Dx@},\] und
        <br/>
        \[\Delta_{y} = \det{@Amaty@}={@ Dy@},\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        where the notation \(\det A = |A|\) is the determinant of a matrix \(A\).
    </span>
    <span class="multilang" lang="de">
        wobei die Notation \(\det A = |A|\) die Determinante einer Matrix \(A\) ist.
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Hence, \[x=\frac{@Dx@}{@D@} = {@Dx/D@},\] and
    </span>
    <span class="multilang" lang="de">
        Daher gilt \[x=\frac{@Dx@}{@D@} = {@Dx/D@},\] und
    </span>
    <br/>
</p>
<p>
    \[y=\frac{@Dy@}{@D@} = {@Dy/D@}.\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>pos:rand(2)+1;
vdet:rand(6)+2;
Au:matrix([rand_with_step(-1,1,2),ev(rand(9)-4)],[0,rand_with_step(-1,1,2)]);
Au[pos,pos]:vdet;
Bu:matrix([rand_with_step(-1,1,2),ev(rand(9)-4)],[0,rand_with_step(-1,1,2)]);
Amat:ev(transpose(Au).Bu);
bmat:matrix([rand(5)+1],[rand(5)+1]);
 xy:matrix([x],[y]);

eqns[1]:ev(Amat.xy)[1,1];
eqns[2]:ev(Amat.xy)[2,1];
truesolm:ev(invert(Amat) . bmat);
correctanswer:[truesolm[1,1],truesolm[2,1]];

D:determinant(Amat);
Amatx:copymatrix(Amat);
Amatx[1,1]:ev(bmat[1,1]);
Amatx[2,1]:ev(bmat[2,1]);
Dx:determinant(Amatx);
Amaty:copymatrix(Amat);
Amaty[1,2]:ev(bmat[1,1]);
Amaty[2,2]:ev(bmat[2,1]);
Dy:determinant(Amaty);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ {@eqns[1]@}\ =\ {@bmat[1,1]@}\] 
\[{@eqns[2]@}\ =\ {@bmat[2,1]@}\]
gives x,y = {@correctanswer@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Correct answer, well done.
    </span>
    <span class="multilang" lang="de">
        Richtige Antwort, gut gemacht.
    </span>
</p>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your answer is partially correct.
    </span>
    <span class="multilang" lang="de">
        Ihre Antwort ist teilweise richtig.
    </span>
</p>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Incorrect answer.
    </span>
    <span class="multilang" lang="de">
        Falsche Antwort.
    </span>
</p>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>correctanswer[1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>correctanswer[2]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>Result</name>
            <value>2.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>correctanswer[1]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>Result-0-T </trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your x-value is correct.
    </span>
    <span class="multilang" lang="de">
        Ihr x-Wert ist richtig.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>Result-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your x-value is incorrect.
    </span>
    <span class="multilang" lang="de">
        Ihr x-Wert ist falsch.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>correctanswer[2]</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>Result-1-T </trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your y-value is correct.
    </span>
    <span class="multilang" lang="de">
        Ihr y-Wert ist richtig.
    </span>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>+</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>Result-1-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Your y-value is incorrect.
    </span>
    <span class="multilang" lang="de">
        Ihr y-Wert ist falsch.
    </span>
</p>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>658506017</deployedseed>
        <deployedseed>1174886397</deployedseed>
        <deployedseed>1974833147</deployedseed>
        <deployedseed>937479468</deployedseed>
        <deployedseed>514296711</deployedseed>
        <deployedseed>1147402</deployedseed>
        <deployedseed>1901494326</deployedseed>
        <deployedseed>1729491826</deployedseed>
        <deployedseed>481817568</deployedseed>
        <deployedseed>1173835051</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>correctanswer[1]</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>correctanswer[2]</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>Result-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>correctanswer[1]-10</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>correctanswer[2]-10</value>
            </testinput>
            <expected>
                <name>Result</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>Result-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
