<quiz>
    <question type="stack">
        <name>
            <text>3.6.3.3 Übung_In Partialbrüchen ausdrücken (kurz) [Exercises_Express in partial fractions (short)]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Express the following in partial fractions.
    </span>
    <span class="multilang" lang="no">
        Delbrøksoppspalt brøken:
    </span>
    <span class="multilang" lang="de">
        Drücken Sie Folgendes in Partialbrüchen aus.
    </span>
</p>
<p>
    \(\displaystyle{{@WholeExpFrac3@}=}\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the expression \[\displaystyle{{@WholeExpFrac3@}}.\] The repeated linear factor is \({@ta31@}\). Therefore the partial fraction decomposition is
    </span>
    <span class="multilang" lang="no">
        Vi har uttrykket \[\displaystyle{{@WholeExpFrac3@}}.\] Den repeterte lineære faktoren er \({@ta31@}\). Derfor er den delbrøksnedbrytningen
    </span>
    <span class="multilang" lang="de">
        Wir erhalten den Ausdruck \[\displaystyle{{@WholeExpFrac3@}}.\] Der wiederholte lineare Faktor ist \({@ta31@}\). Daher ist die Partialbruchzerlegung
    </span>
    \[{@WholeExpFrac3@} = \frac{A}{{@exp11@}} + \frac{B}{({@exp11@})^2}
                = \frac{A\cdot ({@exp11@}) + B}{ ({@exp11@})^2}. \]
</p>
<p>
    <span class="multilang" lang="en">
        Equating the top line yields \[ {@expNum11@} = A\cdot ({@exp11@}) + B.\] Then by substituting values of \(x\) we can solve for \(A, \,B\) and obtain \( A = {@expA@}\) and \(B = {@expB@}.\)
    </span>
    <span class="multilang" lang="no">
        Å likestille den øverste linjen gir \[ {@expNum11@} = A\cdot ({@exp11@}) + B.\] Så ved å sette inn verdiene til \(x\) kan vi løse for \(A, \,B\) ) og få \(A = {@expA@}\) og \(B = {@expB@}.\)
    </span>
    <span class="multilang" lang="de">
        Gleichsetzen der obersten Zeile ergibt \[{@expNum11@} = A\cdot ({@exp11@}) + B.\] Dann durch Ersetzen der Werte von \(x\) können wir nach \(A, \,B\) auflösen und erhalten \( A = {@expA@}\) und \(B = {@expB@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore
    </span>
    <span class="multilang" lang="no">
        Dermed får vi:
    </span>
    <span class="multilang" lang="de">
        Daher
    </span>
    \[{@WholeExpFrac3@} = {@ta32@}.\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var:x;
var1:s;
var2:A;
var3:B;

rand1:rand_with_step(1,4,1);
rand2:rand_with_step(-4,-1,1);
exp1:rand1*var+(rand1*rand2-1);
expDenom1:expand(exp1^2);
ta31:exp1^2;
expA:rand_with_prohib(-4,4,[0]);
expB:rand_with_prohib(-4,4,[0]);
expNum1:expand(expA*exp1+expB);
wholeExp3: expNum1/expDenom1;
WholeExpFrac3: expNum1/exp1^2;
ta32:expA/exp1+expB/exp1^2;

exp11: exp1;
expNum11: expNum1;
</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[{@WholeExpFrac3@} = {@ta32@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta32</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta32</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>628667984</deployedseed>
        <deployedseed>709613715</deployedseed>
        <deployedseed>1592171129</deployedseed>
        <deployedseed>1378982455</deployedseed>
        <deployedseed>156250520</deployedseed>
        <deployedseed>2076675791</deployedseed>
        <deployedseed>1850372578</deployedseed>
        <deployedseed>882804877</deployedseed>
        <deployedseed>866864305</deployedseed>
        <deployedseed>456874969</deployedseed>
        <deployedseed>1413417506</deployedseed>
        <deployedseed>652230326</deployedseed>
        <deployedseed>1565786609</deployedseed>
        <deployedseed>261171551</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans3</name>
                <value>ta32</value>
            </testinput>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans3</name>
                <value>ev(ta32+1,simp)</value>
            </testinput>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
