<quiz>
    <question type="stack">
        <name>
            <text>7.3.2.4 Compute a cofactor</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the matrix \(A = {@matrixa@}\). Compute the cofactors of the following entries:
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie die Matrix \(A = {@matrixa@}\). Berechnen Sie die Kofaktoren der folgenden Einträge:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            The cofactor of \(a_{22}\) is:
        </span>
        <span class="multilang" lang="de">
            Der Kofaktor von \(a_{22}\) ist:
        </span>
        [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        <span class="multilang" lang="en">
            The cofactor of \(a_{23}\) is:
        </span>
        <span class="multilang" lang="de">
            Der Kofaktor von \(a_{23}\) ist:
        </span>
        [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the matrix \(A = {@matrixa@}.\)
    </span>
    <span class="multilang" lang="de">
        Wir haben die Matrix \(A = {@matrixa@}\) gegeben.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            The minor of \(a_{22}\) is the \(2\times2\) determinant \(\begin{vmatrix}a_{11} &amp; a_{13} \\ a_{31} &amp; a_{33}\end{vmatrix}=a_{11}a_{33}-a_{31}a_{13}\). In this case, the minor is: \[\begin{vmatrix}{@AA[1,1]@} &amp; {@AA[1,3]@} \\ {@AA[3,1]@}
                        &amp; {@AA[3,3]@}\end{vmatrix} = {@minorA22@} = {@ta1@}\]
        </span>
        <span class="multilang" lang="de">
            Der Minor von \(a_{22}\) ist die \(2\times2\)-Determinante \(\begin{vmatrix}a_{11} &amp; a_{13 } \\ a_{31} &amp; a_{33}\end{vmatrix}=a_{11}a_{33}-a_{31}a_{13}\). In diesem Fall ist der Minor: \[\begin{vmatrix}{@AA[1,1]@} &amp; {@AA[1,3]@} \\ {@AA[3,1]@}
                        &amp; {@AA[3,3]@}\end{vmatrix} = {@minorA22@} = {@ta1@}\]
        </span>
        <p>
        </p>
        <p>
            <span class="multilang" lang="en">
                The sign in the position of \(a_{22}\) in the array of signs is \(+\), the cofactor of \(a_{22}\) is \(A_{22}=+{@ta1@}\).
            </span>
            <span class="multilang" lang="de">
                Das Vorzeichen an der Position von \(a_{22}\) in der Vorzeichnmatrix ist \(+\), der Kofaktor von \(a_{22}\) ist \(A_{22}=+{@ta1@}\).
            </span>
        </p>
    </li>
    <li>
        <span class="multilang" lang="en">
            The minor of \(a_{22}\) is the \(2\times2\) determinant \(\begin{vmatrix}a_{11} &amp; a_{12} \\ a_{31} &amp; a_{32}\end{vmatrix}=a_{11}a_{32}-a_{31}a_{12}\). In this case, the minor is: \[\begin{vmatrix}{@AA[1,1]@} &amp; {@AA[1,2]@} \\ {@AA[3,1]@}
                        &amp; {@AA[3,2]@}\end{vmatrix} = {@minorA23@} = {@minorA23s@}\]
        </span>
        <span class="multilang" lang="de">
            Der Minor von \(a_{22}\) ist die \(2\times2\)-Determinante \(\begin{vmatrix}a_{11} &amp; a_{12 } \\ a_{31} &amp; a_{32}\end{vmatrix}=a_{11}a_{32}-a_{31}a_{12}\). In diesem Fall ist der Minor: \[\begin{vmatrix}{@AA[1,1]@} &amp; {@AA[1,2]@} \\ {@AA[3,1]@}
                        &amp; {@AA[3,2]@}\end{vmatrix} = {@minorA23@} = {@minorA23s@}\]
        </span>
        <p>
        </p>
        <p>
            <span class="multilang" lang="en">
                sign in the position of \(a_{23}\) in the array of signs is \(-\), so the cofactor of \(a_{23}\) is \(A_{23}={@cofA23@}\).
            </span>
            <span class="multilang" lang="de">
                Das Vorzeichen an der Position von \(a_{23}\) in der Vorzeichenmatrix ist \(-\), also ist der Kofaktor von \(a_{23}\) \(A_{23}={@cofA23@}\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>matrixa: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa: ev(matrixa,simp);
AA: matrixa;
ta1: ev(AA[1,1]*AA[3,3]-AA[1,3]*AA[3,1],simp);
minorA22: unary_minus_sort(AA[1,1]*AA[3,3]-AA[1,3]*AA[3,1]);
minorA23s: ev(AA[1,1]*AA[3,2]-AA[1,2]*AA[3,1],simp);
minorA23: unary_minus_sort(AA[1,1]*AA[3,2]-AA[1,2]*AA[3,1]);
cofA23 : ev(-minorA23s,simp);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(A = {@matrixa@},\quad A_{22} ={@ta1@}, \quad A_{23}= {@cofA23@}\)</text>
        </questionnote>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>MinorAndCofactor</variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant, minor</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>cofA23</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>cofA23</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2047965195</deployedseed>
        <deployedseed>2131635051</deployedseed>
        <deployedseed>1274060207</deployedseed>
        <deployedseed>2067862505</deployedseed>
        <deployedseed>1072463999</deployedseed>
        <deployedseed>470144349</deployedseed>
        <deployedseed>1889680758</deployedseed>
        <deployedseed>2043294690</deployedseed>
        <deployedseed>2081408866</deployedseed>
        <deployedseed>1561557479</deployedseed>
        <deployedseed>1983783838</deployedseed>
        <deployedseed>977011922</deployedseed>
        <deployedseed>860642649</deployedseed>
        <deployedseed>2108460331</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>cofA23</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>-ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>minorA23s</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
