<quiz>
    <question type="stack">
        <name>
            <text>7.3.2.2 Compute a 2x2 minor</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the matrix \(A = {@matrixa@}\). What is the minor of the entry \(a_{22}\)?
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie die Matrix \(A = {@matrixa@}\). Was ist der Minor des Eintrags \(a_{22}\)?
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The minor of \(a_{22}\) is:
    </span>
    <span class="multilang" lang="de">
        Der Minor von \(a_{22}\) ist:
    </span>
    [[input:ans1]][[validation:ans1]][[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The minor of \(a_{22}\) is the \(2\times2\) determinant \(\begin{vmatrix}a_{11} &amp; a_{13} \\ a_{31} &amp; a_{33}\end{vmatrix}=a_{11}a_{33}-a_{31}a_{13}\).
    </span>
    <span class="multilang" lang="de">
        Der Minor von \(a_{22}\) ist die \(2\times2\)-Determinante \(\begin{vmatrix}a_{11} &amp; a_{13 } \\ a_{31} &amp; a_{33}\end{vmatrix}=a_{11}a_{33}-a_{31}a_{13}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For the matrix \(A = {@matrixa@}\) the minor is:
    </span>
    <span class="multilang" lang="de">
        Für die Matrix \(A = {@matrixa@}\) lautet der Minor:
    </span>
</p>
<p>
    \(\begin{vmatrix}{@AA[1,1]@} &amp; {@AA[1,3]@} \\ {@AA[3,1]@} &amp; {@AA[3,3]@}\end{vmatrix} = {@CC@} = {@ta1@}.\)
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>matrixa: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa: ev(matrixa,simp);

AA: matrixa;
relevantminor: matrix([AA[1,1], AA[1,3]], [AA[3,1], AA[3,3]]);

ta1: ev(determinant(relevantminor),simp);

simp: false;
CC: determinant(relevantminor);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(A = {@matrixa@}, {@ta1@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>MinorAndCofactor</variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant, minor</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>401428587</deployedseed>
        <deployedseed>256470495</deployedseed>
        <deployedseed>1580423024</deployedseed>
        <deployedseed>311223496</deployedseed>
        <deployedseed>717670310</deployedseed>
        <deployedseed>478689502</deployedseed>
        <deployedseed>1948238459</deployedseed>
        <deployedseed>975517033</deployedseed>
        <deployedseed>491816791</deployedseed>
        <deployedseed>1330136640</deployedseed>
        <deployedseed>67200360</deployedseed>
        <deployedseed>1294055332</deployedseed>
        <deployedseed>1284459475</deployedseed>
        <deployedseed>1969144389</deployedseed>
        <deployedseed>2131644959</deployedseed>
        <deployedseed>2131917626</deployedseed>
        <deployedseed>691664144</deployedseed>
        <deployedseed>138140317</deployedseed>
        <deployedseed>1729697927</deployedseed>
        <deployedseed>1066017578</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>-ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
