<quiz>
    <question type="stack">
        <name>
            <text>7.3.4.3 Exercise 3</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Find the cofactors of \(x,y,z\) in the following determinant:
    </span>
    <span class="multilang" lang="de">
        Finden Sie die Kofaktoren von \(x,y,z\) in der folgenden Determinante:
    </span>
</p>
\[\begin{vmatrix} {@AA[1,1]@} &amp; {@AA[1,2]@} &amp; {@AA[1,3]@} \\ {@AA[2,1]@} &amp; {@AA[2,2]@} &amp; {@AA[2,3]@} \\ x &amp; y &amp; z \end{vmatrix}\]
<p>
    <span class="multilang" lang="en">
        Cofactor of \(x\):
    </span>
    <span class="multilang" lang="de">
        Kofaktor von \(x\):
    </span>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
<p>
    <span class="multilang" lang="en">
        Cofactor of \(y\):
    </span>
    <span class="multilang" lang="de">
        Kofaktor von \(y\):
    </span>
    [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</p>
<p>
    <span class="multilang" lang="en">
        Cofactor of \(z\):
    </span>
    <span class="multilang" lang="de">
        Kofaktor von \(z\):
    </span>
    [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="en">
    obtain the cofactor of an element of a \(3 \times 3\) matrix we multiply the minor of that element by the corresponding sign from the \(3 \times 3\) array of signs.
    <br/>
    We are given the determinant: \(\begin{vmatrix} {@AA[1,1]@} &amp; {@AA[1,2]@} &amp;
{@AA[1,3]@} \\ {@AA[2,1]@} &amp; {@AA[2,2]@} &amp; {@AA[2,3]@} \\ x &amp; y &amp; z \end{vmatrix}.\)
</span>
<span class="multilang" lang="de">
    Ermitteln Sie den Kofaktor eines Elements einer \(3 \times 3\)-Matrix. Wir multiplizieren den Minor dieses Elements mit dem entsprechenden Vorzeichen aus der \(3 \times 3\) Vorzeichenmatrix.
    <br/>
    Wir erhalten die Determinante: \(\begin{vmatrix} {@AA[1,1]@} &amp; {@AA[1,2]@} &amp;
{@AA[1,3]@} \\ {@AA[2,1]@} &amp; {@AA[2,2]@} &amp; {@AA[2,3]@} \\ x &amp; J &amp; z \end{vmatrix}.\)
</span>
<br/>
<p>
    <span class="multilang" lang="en">
        For \(x\), the sign is \(+\), and the minor is \(\det {@MMx@}\). Hence the cofactor is \({@ta1@}.\)
    </span>
    <span class="multilang" lang="de">
        Für \(x\) ist das Vorzeichen \(+\) und der Minor ist \(\det {@MMx@}\). Daher ist der Kofaktor \({@ta1@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(y\), the sign is \(-\), and the minor is \(\det {@MMy@}\). Hence the cofactor is \({@ta2@}.\)
    </span>
    <span class="multilang" lang="de">
        Für \(y\) ist das Vorzeichen \(-\) und der Minor ist \(\det {@MMy@}\). Daher ist der Kofaktor \({@ta2@}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(z\), the sign is \(+\) again, and the minor is \(\det {@MMz@}\). Hence the cofactor is \({@ta3@}.\)
    </span>
    <span class="multilang" lang="de">
        Für \(z\) ist das Vorzeichen wieder \(+\) und der Minor ist \(\det {@MMz@}\). Daher ist der Kofaktor \({@ta3@}.\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>matrixa: matrix([rand_with_step(-5,5,1), rand_with_step(-5,5,1), rand_with_step(-5,5,1)],[rand_with_step(-5,5,1), rand_with_step(-5,5,1), rand_with_step(-5,5,1)])
matrixa: ev(matrixa,simp);

AA: matrixa;
MMx: matrix([AA[1,2], AA[1,3]], [AA[2,2], AA[2,3]]);
MMy: matrix([AA[1,1], AA[1,3]], [AA[2,1], AA[2,3]]);
MMz: matrix([AA[1,1], AA[1,2]], [AA[2,1], AA[2,2]]);

ta1: ev(determinant(MMx),simp)
ta2: ev(-determinant(MMy),simp)
ta3: ev(determinant(MMz),simp)</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\({@AA@}, \, x:{@ta1@}, \, y:{@ta2@}, \, z:{@ta3@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>815348399</deployedseed>
        <deployedseed>807206414</deployedseed>
        <deployedseed>1359982381</deployedseed>
        <deployedseed>415533922</deployedseed>
        <deployedseed>1666427994</deployedseed>
        <deployedseed>1519906486</deployedseed>
        <deployedseed>1881655594</deployedseed>
        <deployedseed>1203042910</deployedseed>
        <deployedseed>240042115</deployedseed>
        <deployedseed>770303117</deployedseed>
        <deployedseed>736720263</deployedseed>
        <deployedseed>101230784</deployedseed>
        <deployedseed>1543056679</deployedseed>
        <deployedseed>53117613</deployedseed>
        <deployedseed>938344926</deployedseed>
        <deployedseed>1208117384</deployedseed>
        <deployedseed>283144539</deployedseed>
        <deployedseed>392739448</deployedseed>
        <deployedseed>1733957553</deployedseed>
        <deployedseed>250556807</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(ta2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ev(ta3+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
