<quiz>
    <question type="stack">
        <name>
            <text>2.5.1.9 Steigung und y-Achsenabschnitt identifizieren [identify gradient and y intercept]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For each line, identify the gradient and the vertical intercept.
    </span>
    <span class="multilang" lang="no">
        For hver linje, identifiser stigningstallet og skjæringspunktet med y-aksen.
    </span>
    <span class="multilang" lang="de">
        Identifizieren Sie für jede Gerade die Steigung und den vertikalen Schnittpunkt.
    </span>
</p>
<table class="HELM_table">
    <caption>
    </caption>
    <thead>
        <tr>
            <th scope="col" style="width: 10em;">
                <span class="multilang" lang="en">
                    Line
                </span>
                <span class="multilang" lang="no">
                    Linje
                </span>
                <span class="multilang" lang="de">
                    Gerade
                </span>
            </th>
            <th scope="col">
                <span class="multilang" lang="en">
                    Gradient
                </span>
                <span class="multilang" lang="no">
                    Stigningstall
                </span>
                <span class="multilang" lang="de">
                    Steigung
                </span>
            </th>
            <th scope="col">
                <span class="multilang" lang="en">
                    Vertical intercept
                </span>
                <span class="multilang" lang="no">
                    Skjæringspunkt med y-aksen
                </span>
                <span class="multilang" lang="de">
                    Vertikaler Schnittpunkt
                </span>
                <br/>
            </th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>
                \({@line1@}\)
            </td>
            <td>
                [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
            </td>
            <td>
                [[input:ans6]] [[validation:ans6]][[feedback:prt6]]
            </td>
        </tr>
        <tr>
            <td>
                \({@line2@}\)
            </td>
            <td>
                [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
            </td>
            <td>
                [[input:ans7]] [[validation:ans7]][[feedback:prt7]]
            </td>
        </tr>
        <tr>
            <td>
                \({@line3@}\)
            </td>
            <td>
                [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
            </td>
            <td>
                [[input:ans8]] [[validation:ans8]][[feedback:prt8]]
            </td>
        </tr>
        <tr>
            <td>
                \({@line4@}\)
            </td>
            <td>
                [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
            </td>
            <td>
                [[input:ans9]] [[validation:ans9]][[feedback:prt9]]
            </td>
        </tr>
        <tr>
            <td>
                \({@line5@}\)
            </td>
            <td>
                [[input:ans5]] [[validation:ans5]][[feedback:prt5]]
            </td>
            <td>
                [[input:ans10]] [[validation:ans10]][[feedback:prt10]]
            </td>
        </tr>
    </tbody>
</table>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        As already mentioned in the other questions, the gradient is the factor in front of the independent variable, while the vertical intercept is constant term \(b\) in the general expression of the straight line \(y=ax+b\).
    </span>
    <span class="multilang" lang="no">
        Som allerede nevnt i de andre spørsmålene er stigningstallet konstanten foran den uavhengige variabelen, mens skjæringspunktet med y-aksen er det konstante leddet \(b\) i det generelle uttrykket for den rette linjen \(y=ax+b\).
    </span>
    <span class="multilang" lang="de">
        Wie bereits in den anderen Fragen erwähnt, ist die Steigung der Faktor vor der unabhängigen Variablen, während der vertikale Schnittpunkt die Konstante \(b\) im allgemeinen Ausdruck von der Gerade \(y=ax+b\) ist.
    </span>
</p>
<table class="HELM_table">
    <thead>
        <tr>
            <th scope="col">
                <span class="multilang" lang="en">
                    Line
                </span>
                <span class="multilang" lang="no">
                    Linje
                </span>
                <span class="multilang" lang="de">
                    Gerade
                </span>
            </th>
            <th scope="col">
                <span class="multilang" lang="en">
                    Gradient
                </span>
                <span class="multilang" lang="no">
                    Stigningstall
                </span>
                <span class="multilang" lang="de">
                    Steigung
                </span>
            </th>
            <th scope="col">
                <span class="multilang" lang="en">
                    Vertical intercept
                </span>
                <span class="multilang" lang="no">
                    Skjæringspunkt med y-aksen
                </span>
                <span class="multilang" lang="de">
                    Vertikaler Schnittpunkt
                </span>
            </th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td align="leftr">
                \({@line1@}\)
            </td>
            <td align="center">
                {@grad1@}
            </td>
            <td align="center">
                {@vertint1@}
            </td>
        </tr>
        <tr>
            <td align="left">
                \({@line2@}\)
            </td>
            <td align="center">
                {@grad2@}
            </td>
            <td align="center">
                {@vertint2@}
            </td>
        </tr>
        <tr>
            <td align="left">
                \({@line3@}\)
            </td>
            <td align="center">
                {@grad3@}
            </td>
            <td align="center">
                {@vertint3@}
            </td>
        </tr>
        <tr>
            <td align="left">
                \({@line4@}\)
            </td>
            <td align="center">
                {@grad4@}
            </td>
            <td align="center">
                {@vertint4@}
            </td>
        </tr>
        <tr>
            <td align="left">
                \({@line5@}\)
            </td>
            <td align="center">
                {@grad5@}
            </td>
            <td align="center">
                {@vertint5@}
            </td>
        </tr>
    </tbody>
</table>
<p>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var: x;
grad1: rand_with_step(1,12,1);
vertint1: rand_with_step(1, 23, 1);
line1: y=grad1*var+vertint1;
grad2: 0;
vertint2: rand_with_prohib(-10, 10, [0]);
line2: y=vertint2;
grad3: rand_with_step(-15, -1, 1);
line3: y=grad3*var;
vertint3: 0;
grad4: rand_with_step(-25, -2, 1);
vertint4: rand_with_step(-6, -1, 1);
line4: y=vertint4+grad4*var;
grad5: m;
vertint5: c;
line5: y=grad5*var+vertint5;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[{@line1@}, {@grad1@}, {@vertint1@}\]
\[ {@line2@}, {@grad2@}, {@vertint2@}\]
\[{@line3@}, {@grad3@}, {@vertint3@}\]
\[ {@line4@}, {@grad4@}, {@vertint4@}\]
\[ {@line5@}, {@grad5@}, {@vertint5@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>grad1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans10</name>
            <type>algebraic</type>
            <tans>vertint5</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>grad2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>grad3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>numerical</type>
            <tans>grad4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>grad5</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>numerical</type>
            <tans>vertint1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans7</name>
            <type>numerical</type>
            <tans>vertint2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans8</name>
            <type>numerical</type>
            <tans>vertint3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans9</name>
            <type>numerical</type>
            <tans>vertint4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>grad1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt10</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans10</sans>
                <tans>vertint5</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt10-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt10-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans2</sans>
                <tans>grad2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans3</sans>
                <tans>grad3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans4</sans>
                <tans>grad4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans5</sans>
                <tans>grad5</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans6</sans>
                <tans>vertint1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt7</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans7</sans>
                <tans>vertint2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt7-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt8</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans8</sans>
                <tans>vertint3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt8-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt8-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt9</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans9</sans>
                <tans>vertint4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt9-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt9-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1606734682</deployedseed>
        <deployedseed>862891664</deployedseed>
        <deployedseed>1309316084</deployedseed>
        <deployedseed>818976776</deployedseed>
        <deployedseed>1198452241</deployedseed>
        <deployedseed>1523382512</deployedseed>
        <deployedseed>273006868</deployedseed>
        <deployedseed>650961215</deployedseed>
        <deployedseed>363424015</deployedseed>
        <deployedseed>663518958</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>grad1</value>
            </testinput>
            <testinput>
                <name>ans10</name>
                <value>vertint5</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>grad2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>grad3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>grad4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>grad5</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>vertint1</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>vertint2</value>
            </testinput>
            <testinput>
                <name>ans8</name>
                <value>vertint3</value>
            </testinput>
            <testinput>
                <name>ans9</name>
                <value>vertint4</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt10</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt10-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt7-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt8</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt8-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt9</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt9-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(grad1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans10</name>
                <value>ev(vertint5+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(grad2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ev(grad3+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ev(grad4+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ev(grad5+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ev(vertint1+1, simp)</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>ev(vertint2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans8</name>
                <value>ev(vertint3+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans9</name>
                <value>ec(vertint4+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt10</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt10-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt7-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt8</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt8-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt9</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt9-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
