<quiz>
    <question type="stack">
        <name>
            <text>2.5.1.8 Als Gerade oder anderes bestimmen [Classify as line or other]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        State if the following expressions represent straight lines or other function.
    </span>
    <span class="multilang" lang="no">
        Oppgi om følgende uttrykk representerer rette linjer, eller andre funksjoner.
    </span>
    <span class="multilang" lang="de">
        Geben Sie an, ob die folgenden Ausdrücke Geraden oder andere Funktionen darstellen.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \(\displaystyle{@vq1@}\) : [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \(\displaystyle{@vq2@}\) : [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            \(\displaystyle{@vq3@}\) : [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
            \(\displaystyle{@vq4@}\) : [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
        </p>
    </li>
    <li>
        <p>
            \(\displaystyle{@vq5@}\) : [[input:ans5]] [[validation:ans5]][[feedback:prt5]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        A straight line has always to be of the form \(f(x)=y=ax+b.\)
    </span>
    <span class="multilang" lang="no">
        En rett linje må alltid ha formen \(f(x)=y=ax+b.\)
    </span>
    <span class="multilang" lang="de">
        Eine Gerade muss immer die Form \(f(x)=y=ax+b\) haben.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Terms like \(x^{\frac{c}{d}}\), with \(c\) and \(d\) integers, or higher powers of the independent variable \(x\) are not allowed.
    </span>
    <span class="multilang" lang="no">
        Uttrykk som \(x^{\frac{c}{d}}\), der \(c\) og \(d\) er heltall, eller høyere eksponenter av den uavhengige variabelen \(x\) er ikke tillatt.
    </span>
    <span class="multilang" lang="de">
        Terme wie \(x^{\frac{c}{d}}\), mit ganzen Zahlen \(c\) und \(d\) oder höheren Potenzen von der unabhängige Variable \(x\) sind nicht erlaubt.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\displaystyle {@f2@}\) does not represent a straight line because it is of the form \(x^{\frac{c}{d}}.\)
    </span>
    <span class="multilang" lang="no">
        \(\displaystyle{@f2@}\) representerer ikke en rett linje fordi den er på formen \(x^{\frac{c}{d}}.\)
    </span>
    <span class="multilang" lang="de">
        \(\displaystyle {@f2@}\) stellt keine Gerade dar, da sie die Form \(x^{\frac{c}{d}}\) hat.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\displaystyle {@f3@}\) does not represent a straight line because it is an inverse power of \(x.\)
    </span>
    <span class="multilang" lang="no">
        \(\displaystyle {@f3@} \) representerer ikke en rett linje siden det er en invers kraft av \(x. \)
    </span>
    <span class="multilang" lang="de">
        \(\displaystyle {@f3@}\) stellt keine Gerade dar, da es sich um einen inversen Exponenten von \(x\) handelt.
    </span>
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[avail:[1,2,3,4,5];
qs: rand_selection(avail,5);
idx1:qs[1];
idx2:qs[2];
idx3:qs[3];
idx4:qs[4];
idx5:qs[5];

var: x;
coeffx1: rand_with_step(-5,8,1);
rand0: rand_with_prohib(-9,9,[0]);
rand1: rand_with_step(1,7,1);
rand2: rand_with_step(2,9,1);
rand3: rand_with_step(-11,12,1); 
rand4:rand_with_step(-11,12,1); 
rand5:rand_with_step(1,12,1);
coeffx2:  rand_with_prohib(-10,10,[0]);
f1: y=var*coeffx1+rand0;
pow: rand1/rand5;
ff2: var^(pow);
f2: y=ff2;
f3: y=(var)^(-rand2);
f4: y=rand3;
f5: y=rand4+coeffx2*var;

qdes:[f2, f3, f4, f1, f5];
lst:["Other","Other","Straight line","Straight line","Straight line"];
S_list: [true, true, false, false, false];
V_list: [false, false, true, true, true];

vq1: qdes[idx1];
vq2: qdes[idx2];
vq3: qdes[idx3];
vq4: qdes[idx4];
vq5: qdes[idx5];

tav1:[["Other",S_list[idx1]],["Straight line",V_list[idx1]]];
ta1: lst[idx1];
tav2:[["Other",S_list[idx2]],["Straight line",V_list[idx2]]];
ta2: lst[idx2];
tav3:[["Other",S_list[idx3]],["Straight line",V_list[idx3]]];
ta3: lst[idx3];
tav4:[["Other",S_list[idx4]],["Straight line",V_list[idx4]]];
ta4: lst[idx4];
tav5:[["Other",S_list[idx5]],["Straight line",V_list[idx5]]];
ta5: lst[idx5];

 test1: if ta1 = "Straight line"
      then "Other"
      else "Straight line";
test2: if ta2 = "Straight line"
      then "Other"
      else "Straight line";
test3: if ta3 = "Straight line"
      then "Other"
      else "Straight line";
test4: if ta4 = "Straight line"
      then "Other"
      else "Straight line";
test5: if ta5 = "Straight line"
      then "Other"
      else "Straight line";]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>a: {@vq1@}, {@ta1@}
b: {@vq2@}, {@ta2@}
c: {@vq3@}, {@ta3@}
d: {@vq4@}, {@ta4@}
d: {@vq5@}, {@ta5@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>tav1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>dropdown</type>
            <tans>tav2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>dropdown</type>
            <tans>tav3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>dropdown</type>
            <tans>tav4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>dropdown</type>
            <tans>tav5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>ta5</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>613015920</deployedseed>
        <deployedseed>1679128829</deployedseed>
        <deployedseed>2046806979</deployedseed>
        <deployedseed>78184573</deployedseed>
        <deployedseed>906028117</deployedseed>
        <deployedseed>1818923896</deployedseed>
        <deployedseed>508422685</deployedseed>
        <deployedseed>866551085</deployedseed>
        <deployedseed>964377089</deployedseed>
        <deployedseed>1884406391</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta5</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>test1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>test2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>test3</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>test4</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>test5</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
