<quiz>
    <question type="stack">
        <name>
            <text>2.5.3.5 Der Abstand zwischen zwei Punkten und Aufgabe 1 [The distance between two points and task1]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        What is the distance between the points \(A=({@x1@}, {@y1@})\) and \(B=({@x2@}, {@y2@})\)?
    </span>
    <span class="multilang" lang="no">
        Hva er avstanden mellom punktene \(A = ({@x1@}, {@y1@}) \) og \(B = ({@x2@}, {@y2@}) \)?
    </span>
    <span class="multilang" lang="de">
        Wie groß ist der Abstand zwischen den Punkten \(A=({@x1@}, {@y1@})\) und \(B=({@x2@} , {@y2@})\)?
    </span>
</p>
<p>
    [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Using the equation to compute the distance between two points given in the keypoint \(\sqrt{(x_2-x_1)^2+(y_2-y_1)^2}\), we have
    </span>
    <span class="multilang" lang="no">
        Ved å bruke ligningen for å avstanden mellom to punkter gitt i nøkkelpunktet\(\sqrt{(x_2-x_1)^2+(y_2-y_1)^2} \), har vi
    </span>
    <span class="multilang" lang="de">
        Wenn wir die im Schlüsselpunkt gegebene Gleichung zur Berechnung des Abstands zwischen zwei Punkten \(\sqrt{(x_2-x_1)^2+(y_2-y_1)^2}\) verwenden, erhalten wir
    </span>
    \[\sqrt{(x_B-x_A)^2+(y_B-y_A)^2}= \sqrt{({@y2@}-({@y1@}))^2+({@x2@}-({@x1@}))^2}
                = {@solution@}\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>x1: rand_with_prohib(-9,9,[0]);
x2: rand_with_prohib(-9,9,[x1]);
y1: rand_with_step(-9,9,1);
y2: rand_with_prohib(-9,9,[y1]);
solution: ((y2-y1)^2+(x2-x1)^2)^(1/2);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>Points \(A=({@x1@}, {@y1@}), \, B=({@x2@}, {@y2@}),\) distance = {@solution@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>solution</tans>
            <boxsize>6</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>+,-</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>solution</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1113092984</deployedseed>
        <deployedseed>1153608495</deployedseed>
        <deployedseed>242660088</deployedseed>
        <deployedseed>466289036</deployedseed>
        <deployedseed>467250572</deployedseed>
        <deployedseed>96554523</deployedseed>
        <deployedseed>1104426622</deployedseed>
        <deployedseed>1139627702</deployedseed>
        <deployedseed>1173104785</deployedseed>
        <deployedseed>291710280</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>solution</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(2*solution,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
