<quiz>
    <question type="stack">
        <name>
            <text>7.2.7.2 Exercise 2</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Let \(A = {@A@}\) and \(B = {@B@}\).
    </span>
    <span class="multilang" lang="de">
        Sei \(A = {@A@}\) und \(B = {@B@}\).
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Is it true that \(AB=BA\)?
        </span>
        <span class="multilang" lang="de">
            Ist es wahr, dass \(AB=BA\)?
        </span>
        [[input:ans1]][[validation:ans1]][[feedback:prt1]]
    </li>
    <li>
        <span class="multilang" lang="en">
            Given that \(AB = \begin{bmatrix}\cos{z} &amp; \sin{z} \\ -\sin{z} &amp; \cos{z}\end{bmatrix}\), what is \(z\) in terms of \(x\) and \(y\)?
        </span>
        <span class="multilang" lang="de">
            Da \(AB = \begin{bmatrix}\cos{z} &amp; \sin{z} \\ -\sin{z} &amp; \cos{z }\end{bmatrix}\), was ist \(z\) in Bezug auf \(x\) und \(y\)?
        </span>
        [[input:ans2]][[validation:ans2]][[feedback:prt2]]
        <p>
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the matrices \(A = {@A@}\) and \(B = {@B@}\).
    </span>
    <span class="multilang" lang="de">
        Uns sind die Matrizen \(A = {@A@}\) und \(B = {@B@}\) gegeben.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                First, we are asked whether \(AB = BA\), that is, whether \(A\) and \(B\) commute. In order to answer this question, we must find what \(AB\) and \(BA\) are.
            </span>
            <span class="multilang" lang="de">
                Zunächst werden wir gefragt, ob \(AB = BA\), also ob \(A\) und \(B\) kommutieren. Um diese Frage zu beantworten, müssen wir herausfinden, was \(AB\) und \(BA\) sind.
            </span>
        </p>
        \[\begin{aligned} AB &amp;= {@ABsteps@} \\ &amp;= {@AB@} \end{aligned}\] \[\begin{aligned} BA &amp;= {@BAsteps@} \\ &amp;= {@BA@} \end{aligned}\]
        <p>
            <span class="multilang" lang="en">
                Checking each entry at a time, we see that \(AB\) is in fact equal to \(BA\).
            </span>
            <span class="multilang" lang="de">
                Wenn wir jeden Eintrag einzeln prüfen, sehen wir, dass \(AB\) tatsächlich gleich \(BA\) ist.
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                We already computed the value of \(AB\) in part (a). However, it was not fully simplified.
            </span>
            <span class="multilang" lang="de">
                Wir haben den Wert von \(AB\) bereits in Teil (a) berechnet. Es wurde jedoch nicht vollständig vereinfacht.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Recall the trigonometric identities for the sum of angles:
            </span>
            <span class="multilang" lang="de">
                Erinnern Sie sich an die trigonometrischen Identitäten für die Winkelsumme:
            </span>
        </p>
        \[\begin{aligned} \sin(x + y) &amp;= \sin x \cos y + \sin y \cos x \\ \cos(x + y) &amp;= \cos x \cos y - \sin x \sin y \end{aligned}\]
        <p>
            <span class="multilang" lang="en">
                Using these, we can rewrite our expression for \(AB\) as
            </span>
            <span class="multilang" lang="de">
                Mit diesen können wir unseren Ausdruck für \(AB\) umschreiben als
            </span>
        </p>
        \[{@ABfinal@}\]
        <p>
            <span class="multilang" lang="en">
                Therefore, the value of \(z\) we are asked for is \(x + y\).
            </span>
            <span class="multilang" lang="de">
                Daher ist der Wert von \(z\), nach dem wir gefragt werden, \(x + y\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[A: matrix([cos(x), sin(x)], [-sin(x), cos(x)]);
B: ev(A, x=y);
ta : x + y;

/* Needed for solution */
AT: transpose(A);
BT: transpose(B);
ABsteps: zeromatrix(matrix_size(A)[1], matrix_size(B)[2]);
BAsteps: zeromatrix(matrix_size(B)[1], matrix_size(A)[2]);

simp: false;
for a: 1 thru matrix_size(ABsteps)[1] do
    for b: 1 thru matrix_size(ABsteps)[2] do
        ABsteps[ev(a,simp), ev(b,simp)]: apply("+", zip_with("*", A[ev(a,simp)], BT[ev(b,simp)]));
AB: ev(ABsteps, simp);

for a: 1 thru matrix_size(BAsteps)[1] do
    for b: 1 thru matrix_size(BAsteps)[2] do
        BAsteps[ev(a,simp), ev(b,simp)]: apply("+", zip_with("*", B[ev(a,simp)], AT[ev(b,simp)]));
BA: ev(BAsteps, simp);

simp: true;
ABfinal: trigreduce(A.B);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(\theta+\phi\)</text>
        </questionnote>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>[[true,true,Yes],[false,false,No]]</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>x+y</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>acos, asin</forbidwords>
            <allowwords>theta, phi</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>true</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>true</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>first(mcq_incorrect([[true,true,Yes],[false,false,No]]))</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>x-y</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
