<quiz>
    <question type="stack">
        <name>
            <text>7.2.7.5 Exercise 5</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \(A\) is any matrix, then the matrix \(AA^{T}\) is symmetric. Verify this in the case \(A = {@matrixa@}\) by computing \(AA^{T}\):
    </span>
    <span class="multilang" lang="de">
        Wenn \(A\) eine beliebige Matrix ist, dann ist die Matrix \(AA^{T}\) symmetrisch. Überprüfen Sie dies im Fall \(A = {@matrixa@}\), indem Sie \(AA^{T}\) berechnen:
    </span>
</p>
<div>
    \(AA^{T} = \) [[input:ans1]][[validation:ans1]][[feedback:prt1]]
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the matrix \(A = {@AA@}\) and are asked to verify that \(AA^T\) is symmetric by computing it.
    </span>
    <span class="multilang" lang="de">
        Uns ist die Matrix \(A = {@AA@}\) gegeben und wir sollen überprüfen, ob \(AA^T\) symmetrisch ist, indem wir sie berechnen.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        First, we write down the tranpose of \(A\).
    </span>
    <span class="multilang" lang="de">
        Zuerst schreiben wir die Tranponierte von \(A\).
    </span>
    auf
</p>
\[A^T = {@AAT@}\]
<p>
    <span class="multilang" lang="en">
        Now we perform the multiplication.
    </span>
    <span class="multilang" lang="de">
        Nun führen wir die Multiplikation durch.
    </span>
</p>
\[\begin{aligned} AA^T &amp;= {@PP@} \\ &amp;= {@PPs@} \end{aligned}\]
<p>
    <span class="multilang" lang="en">
        Indeed, this matrix is symmetric about its main diagonal.
    </span>
    <span class="multilang" lang="de">
        Tatsächlich ist diese Matrix symmetrisch um ihre Hauptdiagonale.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[matrixa: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)], [rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa: ev(matrixa,simp);

ta1: ev(matrixa.transpose(matrixa),simp);

/* what follows is only for general feedback */

AA: ev(matrixa,simp);
AAT: ev(transpose(matrixa),simp);

PP: zeromatrix (first(matrix_size(AA)), first(matrix_size(AA)));
SS: for a:1 thru first(matrix_size(AA)) do for b:1 thru first(matrix_size(AA)) do PP[ev(a,simp),ev(b,simp)]:apply("+",zip_with("*",AA[ev(a,simp)],AA[ev(b,simp)]));
PPs: ev(PP,simp);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(A = {@AA@}, AA^{T} = {@PPs@}\)</text>
        </questionnote>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta1</tans>
            <boxsize>1</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1333204668</deployedseed>
        <deployedseed>1655111437</deployedseed>
        <deployedseed>1075659029</deployedseed>
        <deployedseed>691537217</deployedseed>
        <deployedseed>147583968</deployedseed>
        <deployedseed>1550441960</deployedseed>
        <deployedseed>2088256393</deployedseed>
        <deployedseed>940032131</deployedseed>
        <deployedseed>1874743048</deployedseed>
        <deployedseed>105974638</deployedseed>
        <deployedseed>560633518</deployedseed>
        <deployedseed>1311951236</deployedseed>
        <deployedseed>2024798351</deployedseed>
        <deployedseed>2103382398</deployedseed>
        <deployedseed>1382044390</deployedseed>
        <deployedseed>614624283</deployedseed>
        <deployedseed>1966106599</deployedseed>
        <deployedseed>1211628417</deployedseed>
        <deployedseed>239342529</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>AA</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
