<quiz>
    <question type="stack">
        <name>
            <text>7.2.3.1 Some surprising results, and Task 1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h3>
    <span class="multilang" lang="en">
        Some surprising results
    </span>
    <span class="multilang" lang="de">
        Einige überraschende Ergebnisse
    </span>
</h3>
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        We have already seen how to calculate a product \(AB\) where, for example, \[ A = {@matrixa@}, \qquad B = {@matrixb@} \] In this case the product is \[AB = {@ev(matrixa.matrixb,simp)@}\] In this exercise you will determine the product \(BA,\) i.e., with
        the matrices in the reverse order.
    </span>
    <span class="multilang" lang="de">
        Wir haben bereits gesehen, wie man ein Produkt \(AB\) berechnet, wobei zum Beispiel \[ A = {@matrixa@}, \qquad B = {@matrixb@} \] In diesem Fall ist das Produkt \[AB = {@ev(matrixa.matrixb,simp)@}\] In dieser Übung bestimmen Sie das Produkt \(BA,\) z.B., mit
        den Matrizen in umgekehrter Reihenfolge.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Form the products of:
    </span>
    <span class="multilang" lang="de">
        Bilden Sie die Produkte von:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        row 1 of \(B\) and column 1 of \(A\)
    </span>
    <span class="multilang" lang="de">
        Zeile 1 von \(B\) und Spalte 1 von \(A\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        row 1 of \(B\) and column 2 of \(A\)
    </span>
    <span class="multilang" lang="de">
        Zeile 1 von \(B\) und Spalte 2 von \(A\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        row 2 of \(B\) and column 1 of \(A\)
    </span>
    <span class="multilang" lang="de">
        Zeile 2 von \(B\) und Spalte 1 von \(A\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        row 2 of \(B\) and column 2 of \(A\)
    </span>
    <span class="multilang" lang="de">
        Zeile 2 von \(B\) und Spalte 2 von \(A\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now enter the matrix \(BA\):
    </span>
    <span class="multilang" lang="de">
        Geben Sie nun die Matrix \(BA\) ein:
    </span>
</p>
<div>
    [[input:ans1]][[validation:ans1]][[feedback:prt1]]
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the matrices \(A = {@matrixa@}\) and \(B = {@matrixb@}\) and are asked to compute the product \(BA\). This is given by:
    </span>
    <span class="multilang" lang="de">
        Wir erhalten die Matrizen \(A = {@matrixa@}\) und \(B = {@matrixb@}\) und sollen das Produkt \( BA\) berechnen. Dies ist gegeben durch:
    </span>
</p>
\[BA = {@CC@} = {@DD@}\]
<p>
    <span class="multilang" lang="en">
        Note how our result for \(BA\) is not equal to \(AB\).
    </span>
    <span class="multilang" lang="de">
        Beachten Sie, dass unser Ergebnis für \(BA\) nicht gleich \(AB\) ist.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[matrixa: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa: ev(matrixa,simp);
matrixb: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixb: ev(matrixb,simp);
correctanswer: ev(matrixb . matrixa,simp);

/* these need to be renamed! */
AA: matrixb;
BBT: transpose(matrixa);
CC: zeromatrix (first(matrix_size(AA)), first(matrix_size(BBT)));
SS: for a:1 thru first(matrix_size(AA)) do for b:1 thru first(matrix_size(BBT)) do CC[ev(a,simp),ev(b,simp)]:apply("+",zip_with("*",AA[ev(a,simp)],BBT[ev(b,simp)]));
DD: ev(CC,simp);

ta1: correctanswer]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(A = {@matrixa@}, B = {@matrixb@}, BA ={@correctanswer@}\)</text>
        </questionnote>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta1</tans>
            <boxsize>1</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>determinant</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>95628754</deployedseed>
        <deployedseed>1775091312</deployedseed>
        <deployedseed>245330457</deployedseed>
        <deployedseed>1965837673</deployedseed>
        <deployedseed>1389803148</deployedseed>
        <deployedseed>1829180109</deployedseed>
        <deployedseed>1372851919</deployedseed>
        <deployedseed>1931232441</deployedseed>
        <deployedseed>1277226495</deployedseed>
        <deployedseed>721254620</deployedseed>
        <deployedseed>839270455</deployedseed>
        <deployedseed>655998488</deployedseed>
        <deployedseed>1129965067</deployedseed>
        <deployedseed>1480151935</deployedseed>
        <deployedseed>1603726730</deployedseed>
        <deployedseed>1577850419</deployedseed>
        <deployedseed>1646849911</deployedseed>
        <deployedseed>357615259</deployedseed>
        <deployedseed>717019481</deployedseed>
        <deployedseed>606166384</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(transpose(matrixa.matrixb),simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
