<quiz>
    <question type="stack">
        <name>
            <text>9.2.1.3 Column vector notation</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Given \(\underline{p} = {@p_vec@}\) and \(\underline{q} = {@q_vec@}\), express \(\underline{p}\) and \(\underline{q}\) as column vectors.
    </span>
    <span class="multilang" lang="de">
        Gegeben sind \(\underline{p} = {@p_vec@}\) und \(\underline{q} = {@q_vec@}\), drücken Sie \(\underline{p}\) und \(\underline{q}\) als Spaltenvektoren aus.
    </span>
    <br/>
</p>
<div>
    \(\underline{p} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</div>
<div>
    <br/>
</div>
<div>
    \(\underline{q} =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</div>
<div>
    <br/>
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        In column vector notation we arrange the elements with the \(x\)-component at the top and the \(y\)-component below.
    </span>
    <span class="multilang" lang="de">
        In Spalten-Vektor-Notation ordnen wir die Elemente mit der \(x\)-Komponente oben und der \(y\)-Komponente unten an.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        So, \[\underline{p} = {@p_vec@} = {@ta1@}\]
    </span>
    <span class="multilang" lang="de">
        Also \[\underline{p} = {@p_vec@} = {@ta1@}\]
    </span>
</p>
<p>
    \[\underline{q} = {@q_vec@} = {@ta2@}\]
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stack_reset_vars(true);
stackvectortex(ex):= block(sconcat("{\\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
ijk_vec: matrix([i],[j]);
selection: [0,1,2,3,4,5,6,7,8,9,10];
p_sel: rand_selection(selection,2);
p_i: p_sel[1];
p_j: p_sel[2];
p_mat:matrix([p_i],[p_j]);
p_vec:texboldatoms(dotproduct(ijk_vec,p_mat));

q_sel: rand_selection(selection,2);
q_i: q_sel[1];
q_j: q_sel[2];
q_mat:matrix([q_i],[q_j]);
q_vec:texboldatoms(dotproduct(ijk_vec,q_mat));

ta1:matrix([p_i],[p_j]);
ta2:matrix([q_i],[q_j]);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(\underline{p} = {@p_vec@}\) and \(\underline{q} = {@q_vec@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>matrix</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>matrix([p_j],[p_i])</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Remember that the first element of a column vector represents the component in the \(x\) direction and the second is the component in the \(y\) axis.
    </span>
    <span class="multilang" lang="de">
        Denken Sie daran, dass das erste Element eines Spaltenvektors die Komponente in der \(x\)-Richtung darstellt und das Zweite die Komponente in der \(y\)-Achse.
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>matrix([q_j],[q_i])</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Remember that the first element of a column vector represents the component in the \(x\) direction and the second is the component in the \(y\) axis.
    </span>
    <span class="multilang" lang="de">
        Denken Sie daran, dass das erste Element eines Spaltenvektors die Komponente in der \(x\)-Richtung darstellt und das Zweite die Komponente in der \(y\)-Achse.
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1035191399</deployedseed>
        <deployedseed>1197021736</deployedseed>
        <deployedseed>1395238984</deployedseed>
        <deployedseed>569798231</deployedseed>
        <deployedseed>898479552</deployedseed>
        <deployedseed>1405119110</deployedseed>
        <deployedseed>1302621987</deployedseed>
        <deployedseed>811918157</deployedseed>
        <deployedseed>1830153782</deployedseed>
        <deployedseed>1965565836</deployedseed>
        <deployedseed>786527033</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>matrix([p_j],[p_i])</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>matrix([q_j],[q_i])</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
