<quiz>
    <question type="stack">
        <name>
            <text>9.2.2.2 modulus of vectors (cartesian 3D)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Find the modulus of each of the following vectors (writing your answer in the form of a surd, i.e. \(\sqrt{x}\)):
    </span>
    <span class="multilang" lang="de">
        Berechnen Sie den Betrag jedes der folgenden Vektoren (schreiben Sie Ihre Antwort in Wurzelschreibweise, d. h. \(\sqrt{x}\)):
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \(\underline{r} = {@a_vec@}, \quad |\underline{r}| = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \(\underline{s} = {@b_vec@}, \quad |\underline{s}| = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            \(\underline{t} = {@coefc1@}\underline{i} + {@coefc2@}\underline{j} - {@coefc3@}\underline{k}, \quad |\underline{t}| = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Considering the general case in part (c), for the vector, \(\underline{t} = {@coefc1@}\underline{i} + {@coefc2@}\underline{j} - {@coefc3@}\underline{k}\) we calculate the modulus by applying Pythagoras' theorem: \(|\underline{t}| = \sqrt{{@coefc1@}^2
        + {@coefc2@}^2 + (-{@coefc3@})^2}\).
    </span>
    <span class="multilang" lang="de">
        Unter Berücksichtigung des allgemeinen Falls in Teil (c) für den Vektor \(\underline{t} = {@coefc1@}\underline{i} + {@coefc2@}\underline{j} - {@coefc3@}\underline{k}\), berechnen wir den Betrag, indem wir den Satz des Pythagoras anwenden: \(|\underline{t}| = \sqrt{{@coefc1@}^2
        + {@coefc2@}^2 + (-{@coefc3@})^2}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        By applying this we get:
    </span>
    <span class="multilang" lang="de">
        Indem wir dies anwenden, erhalten wir:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(\underline{r} = {@a_vec@}, \quad |\underline{r}| = {@ta1@} .\)
    </span>
    <span class="multilang" lang="de">
        Für \(\underline{r} = {@a_vec@}, \quad |\underline{r}| = {@ta1@} .\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        For \(\underline{s} = -{@b_vec@}, \quad |\underline{s}| = {@ta2@}.\)
    </span>
    <span class="multilang" lang="de">
        Für \(\underline{s} = -{@b_vec@}, \quad |\underline{s}| = {@ta2@}.\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[stack_reset_vars(true);
stackvectortex(ex):= block(sconcat("{\\underline{", tex1(first(args(ex))), "}}"));
ordergreat(i,j,k);
ijk_vec: matrix([i],[j],[k]);

a_2_sign:2*rand(2)-1;
b_2_sign:2*rand(2)-1;
a_3_sign:2*rand(2)-1;
b_3_sign:2*rand(2)-1;

sel: [2,3,4,5,6,7,8,9,10,11,12,13,14,15,16,17,18,19,20];
picks:rand_selection(sel,6);
coefa1:picks[1];
coefa2:picks[2]*a_2_sign;
coefa3:picks[3]*a_3_sign;
a_mat:matrix([coefa1],[coefa2],[coefa3])
a_vec: texboldatoms(dotproduct(ijk_vec,a_mat));

coefb1:picks[4];
coefb2:picks[5]*b_2_sign;
coefb3:picks[6]*b_3_sign;
b_mat:matrix([coefb1],[coefb2],[coefb3])
b_vec: texboldatoms(dotproduct(ijk_vec,b_mat));

coefc1: a;
coefc2: b;
coefc3: c;

ta1:sqrt(coefa1^2+coefa2^2+coefa3^2);
ta2:sqrt(coefb1^2+coefb2^2+coefb3^2);
ta3:sqrt(coefc1^2+coefc2^2+coefc3^2);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[\mathbf{r} = {@a_vec@}, \quad |\mathbf{r}| = {@ta1@} \]
\[\mathbf{s} = {@b_vec@}, \quad |\mathbf{s}| = {@ta2@}\]
\[\mathbf{t} = {@coefc1@}\mathbf{i} + {@coefc2@}\mathbf{j} - {@coefc3@}\mathbf{k}, \quad |\mathbf{t}| = {@ta3@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>+</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>+</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>neg_wrong:sqrt(coefa1^2 + (a_2_sign*(coefa2^2))+(a_3_sign*(coefa3^2)));</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>neg_wrong</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        When we calculate the modulus of the vector we take \(\sqrt{({@coefa1@})^2 + ({@coefa2@})^2 + ({@coefa3@})^2}\).
    </span>
    <span class="multilang" lang="de">
        Wenn wir den Betrag des Vektors berechnen, nehmen wir \(\sqrt{({@coefa1@})^2 + ({@coefa2@})^2 + ({@coefa3@})^2}\).
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>neg_wrongb:sqrt(coefb1^2 + (b_2_sign*(coefb2^2))+(b_3_sign*(coefb3^2)));</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>neg_wrongb</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        When we calculate the modulus of the vector we take \(\sqrt{({@coefb1@})^2 + ({@coefb2@})^2 + {@coefb3@}^2}\).
    </span>
    <span class="multilang" lang="de">
        Wenn wir den Betrag des Vektors berechnen, nehmen wir \(\sqrt{({@coefb1@})^2 + ({@coefb2@})^2 + {@coefb3@}^2}\).
    </span>
    <br/>
</p>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>sqrt(coefc1^2+coefc2^2 - coefc3^2)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    When we calculate the modulus of the vector we take \(\sqrt{{@coefc1@}^2 + {@coefc2@}^2 + (-{@coefc3@})^2}\).
</span>
<span class="multilang" lang="de">
    Wenn wir den Betrag des Vektors berechnen, nehmen wir \(\sqrt{{@coefc1@}^2 + {@coefc2@}^2 + (-{@coefc3@})^2}\).
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1759603625</deployedseed>
        <deployedseed>11743682</deployedseed>
        <deployedseed>672085597</deployedseed>
        <deployedseed>2110923551</deployedseed>
        <deployedseed>234590933</deployedseed>
        <deployedseed>1817988351</deployedseed>
        <deployedseed>1574340206</deployedseed>
        <deployedseed>347631795</deployedseed>
        <deployedseed>415051493</deployedseed>
        <deployedseed>911398956</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(sqrt(coefa1^2 + (a_2_sign*coefa2^2) + (a_3_sign*coefa3^2)),simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(sqrt(coefb1^2 + (b_2_sign*coefb2^2) + (b_3_sign*coefb3^2)),simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>sqrt(coefc1^2 - coefc2^2 - coefc3^2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>sqrt(coefc1^2 + coefc2^2 - coefc3^2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
