<quiz>
    <question type="stack">
        <name>
            <text>3.5.3.4 Lineare und quadratische Ungleichungen (gekürzt) [Linear and quadratic inequalities (pruned)]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In questions (a)-(c) solve each inequality using the graph of the function:
    </span>
    <span class="multilang" lang="no">
        I spørsmål (a)-(c) løs hver ulikhet ved å bruke grafen til funksjonen:
    </span>
    <span class="multilang" lang="de">
        Lösen Sie in den Fragen (a)-(c) jede Ungleichung mit dem Graphen der Funktion:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        \({@lineq1@}\lt 0\)
        <p>
            \(x&lt;\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
    </li>
    <li>
        \({@lineq2@}\gt 0\)
        <p>
            \(x&gt;\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
        </p>
    </li>
    <li>
        \({@eqt@}\lt 0\)
        <p>
            \(x&lt;\) [[input:ans5]] [[validation:ans5]] , \(x&gt;\) [[input:ans6]] [[validation:ans6]][[feedback:prt6]] [[feedback:prt5]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            The graph crosses the \({@var@}\)-axis when \({@var@}=\frac{-({@coef1@})}{{@varcoef1@}}\). The graph of \(y={@lineq1@}\) is negative for all \(x\) less than this value.
        </span>
        <span class="multilang" lang="no">
            Grafen krysser \({@var@}\)-aksen når \({@var@}=\frac{-({@coef1@})}{{@varcoef1@}}\). Grafen til \(y={@lineq1@}\) er negativ for alle \(x\) mindre enn denne verdien.
        </span>
        <span class="multilang" lang="de">
            Der Graph schneidet die \({@var@}\)-Achse, wenn \({@var@}=\frac{-({@coef1@})}{ {@varcoef1@}}\). Der Graph von \(y={@lineq1@}\) ist negativ für alle \(x\) kleiner als dieser Wert.
        </span>
        <figure class="HELM_figure">
            {@plot1@}
        </figure>
    </li>
    <li>
        <span class="multilang" lang="en">
            The graph crosses the \({@var@}\)-axis when \({@var@}=\frac{{@-coef2@}}{{@varcoef2@}}\). The graph of \(y={@lineq2@}\) is positive for all \(x\) greater than this value.
        </span>
        <span class="multilang" lang="no">
            Grafen krysser \({@var@}\)-aksen når \({@var@}=\frac{{@-coef2@}}{{@varcoef2@}}\). Grafen til \(y={@lineq2@}\) er positiv for alle \(x\) større enn denne verdien.
        </span>
        <span class="multilang" lang="de">
            Der Graph schneidet die \({@var@}\)-Achse, wenn \({@var@}=\frac{{@-coef2@}}{{@ varcoef2@}}\). Der Graph von \(y={@lineq2@}\) ist positiv für alle \(x\) größer als dieser Wert.
        </span>
        <figure class="HELM_figure">
            {@plot2@}
        </figure>
    </li>
    <li>
        <span class="multilang" lang="en">
            The graph crosses the \({@var@}\)-axis when \({@var@}={@quadfirst@}\) and \({@var@}={@quadsec@}\). The graph of \(y={@eqt@}\) is negative when \({@var@}\in ({@quadfirst@},{@quadsec@})\), so the correct answer is:
        </span>
        <span class="multilang" lang="no">
            Grafen krysser \({@var@}\)-aksen når \({@var@}={@quadfirst@}\) og \({@var@}={@quadsec@}\). Grafen til \(y={@eqt@}\) er negativ når \({@var@}\in ({@quadfirst@},{@quadsec@})\), så det riktige svaret er:
        </span>
        <span class="multilang" lang="de">
            Der Graph schneidet die \({@var@}\)-Achse, wenn \({@var@}={@quadfirst@}\) und \({@var @}={@quadsec@}\). Der Graph von \(y={@eqt@}\) ist negativ, wenn \({@var@}\in ({@quadfirst@},{@quadsec@})\), also lautet die richtige Antwort:
        </span>
        \({@var@}
                                \lt {@quadsec@}\)
        <span class="multilang" lang="en">
            and
        </span>
        <span class="multilang" lang="no">
            og
        </span>
        <span class="multilang" lang="de">
            und
        </span>
        \({@var@} \gt {@quadfirst@}\).
        <figure class="HELM_figure">
            {@plot3@}
        </figure>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>var:x;
varcoef1: rand_with_step(1,9,1);
coef1: rand_with_prohib(-5,5,[0]);
lineq1: varcoef1*var+coef1;
varcoef2: rand_with_step(1,9,1);
coef2: rand_with_prohib(-5,5,[0]);
lineq2: varcoef2*var+coef2;


correctans1: -coef1/varcoef1;
correctans2: -coef2/varcoef2;


quadfirst: rand_with_step(-5,-1,1);
quadsec: rand(5)+1;
eqt: expand((var-quadfirst)*(var-quadsec));
correctans5: quadsec;
correctans6: quadfirst;

plot1: plot(lineq1, [x,-coef1/varcoef1-1.5, -coef1/varcoef1+1.5], [box, false],[yx_ratio, 1],[axes, solid]);

plot2: plot(lineq2, [x,-abs(-coef2/varcoef2)-0.2, abs(-coef2/varcoef2)+0.2], [box, false],[yx_ratio, 1],[axes, solid]);


plot3: plot(eqt, [x,correctans6-1, correctans5+1], [box, false],[yx_ratio, 1],[axes, solid]);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text><![CDATA[\({@lineq1@}\lt 0 \ \mbox{ans: x<} {@correctans1@}\)
\({@lineq2@}\lt 0 \ \mbox{ans: x<} {@correctans2@}\)
\({@eqt@}\lt 0 \ \mbox{ans: x<} {@correctans5@} \ \mbox{ans: x>} {@correctans6@}\)
]]></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>correctans1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>correctans2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>correctans5</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>correctans6</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>correctans1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    The correct answer is \({@correctans1@}\).
</span>
<span class="multilang" lang="de">
    Die richtige Antwort ist \({@correctans1@}\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>correctans2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    The correct answer is \({@correctans2@}\).
</span>
<span class="multilang" lang="de">
    Die richtige Antwort ist \({@correctans2@}\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>correctans5</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    The correct answer is \({@correctans5@}\).
</span>
<span class="multilang" lang="de">
    Die richtige Antwort ist \({@correctans5@}\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>correctans6</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    The correct answer is \({@correctans6@}\).
</span>
<span class="multilang" lang="de">
    Die richtige Antwort ist \({@correctans6@}\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>104574154</deployedseed>
        <deployedseed>781460552</deployedseed>
        <deployedseed>387999321</deployedseed>
        <deployedseed>2041283710</deployedseed>
        <deployedseed>1585452832</deployedseed>
        <deployedseed>2031499108</deployedseed>
        <deployedseed>1664711463</deployedseed>
        <deployedseed>1298125062</deployedseed>
        <deployedseed>1245444807</deployedseed>
        <deployedseed>1804433803</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>correctans1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>correctans2</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>correctans5</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>correctans6</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(correctans1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(correctans2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ev(correctans5+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ev(correctans6+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
