<quiz>
    <question type="stack">
        <name>
            <text>2.3.3.2 Demonstrieren Sie x^4 Viele-zu-Eins [Demonstrate x^4 many-to-one]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Illustrate why \(y = x^4\) is a many-to-one function by providing suitable examples for the following:
    </span>
    <span class="multilang" lang="de">
        Verdeutlichen Sie, warum \(y = x^4\) eine Viele-zu-Eins-Funktion ist, indem Sie geeignete Beispiele für Folgendes geben:
    </span>
</p>
<p>
    \( x_1 = \) [[input:ansx1]] [[validation:ansx1]]
</p>
<p>
    \( x_2 = \) [[input:ansx2]] [[validation:ansx2]]
</p>
<p>
    \( y_1 = f(x_1) =\) [[input:ansy1]] [[validation:ansy1]]
</p>
<p>
    \( y_2 = f(x_2) =\) [[input:ansy2]] [[validation:ansy2]]
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        Pick \(x_1\) and \(x_2\) values which help to explain why \(y = x^4\) is many-to-one. For example, we have \({@tax1@} \neq {@tax2@}\) and yet the value of the function at these two points is {@tay1@}.
    </span>
    <span class="multilang" lang="de">
        Wählen Sie \(x_1\)- und \(x_2\)-Werte aus, die erklären, warum \(y = x^4\) Viele-zu-Eins ist. Zum Beispiel haben wir \({@tax1@} \neq {@tax2@}\) und dennoch ist der Wert der Funktion an diesen beiden Punkten {@tay1@}.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>p:x^4;
tax1:2;
tax2:-2;
tay1:ev(p,x=tax1);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ansx1</name>
            <type>numerical</type>
            <tans>tax1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ansx2</name>
            <type>numerical</type>
            <tans>tax2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ansy1</name>
            <type>numerical</type>
            <tans>tay1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ansy2</name>
            <type>numerical</type>
            <tans>tay1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ansx1</sans>
                <tans>ansx2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You must supply distinct \(x_1\) and \(x_2\)!
</span>
<span class="multilang" lang="de">
    Sie müssen unterschiedliche \(x_1\) und \(x_2\) angeben!
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ansy1</sans>
                <tans>ev(p,x=ansx1)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You don't appear to have evaluated \(y_1={@ansy1@} \neq f(x_1) = {@ev(p,x=ansx1)@}\), so you must have done something wrong!
</span>
<span class="multilang" lang="en">
    Sie scheinen \(y_1={@ansy1@} \neq f(x_1) = {@ev(p,x=ansx1)@}\) nicht ausgewertet zu haben, also müssen Sie etwas falsch gemacht haben!
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ansy2</sans>
                <tans>ev(p,x=ansx2)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty></truepenalty>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    You don't appear to have evaluated \(y_2={@ansy2@} \neq f(x_2) = {@ev(p,x=ansx2)@}\), so you must have done something wrong!
</span>
<span class="multilang" lang="de">
    Sie scheinen \(y_2={@ansy2@} \neq f(x_2) = {@ev(p,x=ansx2)@}\) nicht ausgewertet zu haben, also müssen Sie etwas falsch gemacht haben!
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <answertest>AlgEquiv</answertest>
                <sans>ev(p,x=ansx1)</sans>
                <tans>ev(p,x=ansx2)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Yes, you have \(x_1\neq x_2\) but \(y_1=y_2={@ev(p,x=ansx1)@}.\)
</span>
<span class="multilang" lang="de">
    Ja, Sie schreiben \(x_1\neq x_2\), es gilt aber \(y_1=y_2={@ev(p,x=ansx1)@}.\)
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    No, you have \(x_1\neq x_2\) but \({@ev(p,x=ansx1)@}=y_1\neq y_2={@ev(p,x=ansx2)@}.\)
</span>
<span class="multilang" lang="de">
    Nein, Sie schreiben \(x_1\neq x_2\), es gilt aber \({@ev(p,x=ansx1)@}=y_1\neq y_2={@ev( p,x=ansx2)@}.\)
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ansx1</name>
                <value>tax1</value>
            </testinput>
            <testinput>
                <name>ansx2</name>
                <value>tax2</value>
            </testinput>
            <testinput>
                <name>ansy1</name>
                <value>tay1</value>
            </testinput>
            <testinput>
                <name>ansy2</name>
                <value>tay1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-4-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ansx1</name>
                <value>tax1</value>
            </testinput>
            <testinput>
                <name>ansx2</name>
                <value>tax1</value>
            </testinput>
            <testinput>
                <name>ansy1</name>
                <value>tay1</value>
            </testinput>
            <testinput>
                <name>ansy2</name>
                <value>tay1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
