<quiz>
    <question type="stack">
        <name>
            <text>2.3.3.5 Funktionen umkehren [Invert functions]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Find the inverse of each of the following functions:
    </span>
    <span class="multilang" lang="de">
        Finden Sie die Umkehrung jeder der folgenden Funktionen:
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \(f(x)= {@aa@} \Rightarrow f^{-1}(x) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \(f(x)= {@bb@} \Rightarrow f^{-1}(x) = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            \(f(x)= {@cc@} \Rightarrow f^{-1}(x) = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
            \(\displaystyle{f(x)= {@dd@} \Rightarrow f^{-1}(x) = }\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                If \(f(x)= {@aa@}\) then we can write \( x =\displaystyle \frac{y-{@a1@}}{@a2@}\) and \( f^{-1}(x)= \displaystyle {@Ansa@}.\)
            </span>
            <span class="multilang" lang="de">
                Wenn \(f(x)= {@aa@}\) dann können wir schreiben \( x =\displaystyle \frac{y-{@a1@}}{@a2@}\) und \( f^{-1}(x)= \displaystyle {@Ansa@}.\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(f(x)= {@bb@}\) we take \( f^{-1}(x) = {@Ansb@}.\)
            </span>
            <span class="multilang" lang="de">
                Für \(f(x)= {@bb@}\) erhalten wir \( f^{-1}(x) = {@Ansb@}.\)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(f(x)= {@cc@}\) we take \( f^{-1}(x) = \displaystyle {@Ansc@}.\)
            </span>
            <span class="multilang" lang="de">
                Für \(f(x)= {@cc@}\) erhalten wir \( f^{-1}(x) = \displaystyle {@Ansc@}. \)
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                For \(\displaystyle{f(x)= {@dd@}}\) we take \( f^{-1}(x) =\displaystyle  {@Ansd@}.\)
            </span>
            <span class="multilang" lang="de">
                Für \(\displaystyle{f(x)= {@dd@}}\) erhalten wir \( f^{-1}(x) =\displaystyle  {@Ansd@}.\)
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: 1+rand(9);
a2: 2+rand(8);
aa: a2*x + a1;
Ansa: (x - a1)/a2;

b1: rand([-1,1]);
bb: b1*x;
Ansb: (1/b1)*x;

c1: 21+rand(20);
cc: -b1*c1*x;
Ansc: -(1/(b1*c1))*x;

d1: 1+rand(4);
d2: 1+rand(4);
dd: d1/(x + d2);
Ansd: (d1 - d2*x)/x;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[f(x) = {@aa@} \Rightarrow f^{-1}(x) = {@Ansa@} \]
\[f(x) = {@bb@} \Rightarrow f^{-1}(x) = {@Ansb@} \]
\[f(x) = {@cc@} \Rightarrow f^{-1}(x) = {@Ansc@} \]
\[f(x) = {@dd@} \Rightarrow f^{-1}(x) = {@Ansd@} \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>Ansa</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>Ansb</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>Ansc</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>Ansd</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>Ansa</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>Ansb</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>Ansc</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>Ansd</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1589588607</deployedseed>
        <deployedseed>30770722</deployedseed>
        <deployedseed>601348234</deployedseed>
        <deployedseed>740225704</deployedseed>
        <deployedseed>415177568</deployedseed>
        <deployedseed>293092557</deployedseed>
        <deployedseed>1164018457</deployedseed>
        <deployedseed>1631347340</deployedseed>
        <deployedseed>678705173</deployedseed>
        <deployedseed>832316823</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>Ansa</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>Ansb</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>Ansc</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>Ansd</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>Ansa+1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>Ansb+1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>Ansc+1</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>Ansd+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
