<quiz>
    <question type="stack">
        <name>
            <text>2.4.4.4 Häufigkeit der Periode [Frequency from period]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="no">
        Øving
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        A periodic function has period \({\displaystyle{@float(pp)@}}\) seconds. How many times will the pattern in the graph repeat over an interval of {@tt@} seconds?
    </span>
    <span class="multilang" lang="no">
        En periodisk funksjon har periode \({\displaystyle{@float(pp)@}} \) sekunder. Hvor mange ganger vil mønsteret i grafen gjenta seg over et intervall på {@tt@} sekunder?
    </span>
    <span class="multilang" lang="de">
        Eine periodische Funktion hat eine Periode von \({\displaystyle {@float (pp) @}} \) Sekunden. Wie oft wiederholt sich das Muster im Graph in einem Intervall von {@ tt @} Sekunden?
    </span>
</p>
<p>
    \( N = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The periodic function has period \({\displaystyle{@float(pp)@}}\) seconds.
    </span>
    <span class="multilang" lang="no">
        Den periodiske funksjonen har perioden \({\displaystyle{@float(pp)@}}\) sekunder.
    </span>
    <span class="multilang" lang="de">
        Die periodische Funktion hat eine Periode von \({\displaystyle {@float (pp) @}}\) Sekunden.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The frequency (the number of repetitions per time unit) is \( f = \frac{1}{T} = {@1/pp@}\, \textrm{s}^{-1} = {@1/pp@}\, \textrm{Hz}.\)
    </span>
    <span class="multilang" lang="no">
        Frekvensen (antall repetisjoner per tidsenhet) er \(f = \frac{1}{T} = {@1/pp@} \, \textrm{s}^{-1} = {@1/pp @} \, \textrm{Hz}. \)
    </span>
    <span class="multilang" lang="de">
        Die Frequenz (die Anzahl der Wiederholungen pro Zeiteinheit) ist \(f = \frac {1} {T} = {@ 1 / pp @}\, \textrm { s} ^ {- 1} = {@ 1 / pp @}\, \textrm {Hz}.\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        So in the time interval of \( {@tt@} \,\textrm{s}\) the pattern repeats \( N = {@1/pp@}\, \textrm{s}^{-1} \times {@tt@} \, \textrm{s} = {@Ans@}\) times.
    </span>
    <span class="multilang" lang="no">
        Så i tidsintervallet for \({@tt@} \, \textrm{s} \) gjentas mønsteret \(N = {@1/pp@} \, \textrm{s}^{-1} \times {@tt@} \, \textrm{s} = {@Ans@} \) ganger.
    </span>
    <span class="multilang" lang="de">
        Also im Zeitintervall von \( {@ tt @} \, \textrm {s}\) wiederholt sich das Muster \(N = {@ 1/pp @} \, \textrm{s}^{- 1} \times {@tt@} \, \textrm{s} = {@Ans@} \) mal.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>pp: 10^-(1+rand(3));
tt: rand([10,20,30,40,50,50,60,80]);

Ans: tt/pp;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(T ={@pp@}, t={@tt@}, N = {@Ans@}\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>Ans</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>CasEqual</answertest>
                <sans>ans1</sans>
                <tans>Ans</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>474454178</deployedseed>
        <deployedseed>2103682428</deployedseed>
        <deployedseed>26316965</deployedseed>
        <deployedseed>771985559</deployedseed>
        <deployedseed>748981128</deployedseed>
        <deployedseed>745917705</deployedseed>
        <deployedseed>398645102</deployedseed>
        <deployedseed>1867571400</deployedseed>
        <deployedseed>68311927</deployedseed>
        <deployedseed>1038670467</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>Ans</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>pp*tt</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
