<quiz>
    <question type="stack">
        <name>
            <text>9.3.2.2 The scalar product</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \(\underline{ p}={@p1@}\underline{ i}-{@p2@}\underline{ j}+7\underline{ k}\) and \(\underline{ q}=6\underline{ i}-\underline{ j}+{@q3@}\underline{ k}\),
                then 
                \(\underline{ p}\cdot\underline{ q}\) is:
    </span>
    <span class="multilang" lang="de">
        Wenn \(\underline{ p}={@p1@}\underline{ i}-{@p2@}\underline{ j}+7\underline{ k}\) und \(\underline{ q}=6\underline{ i}-\underline{ j}+{@q3@}\underline{ k}\),
                dann ist 
                \(\underline{ p}\cdot\underline{ q}\):
    </span>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We use the formula given in a previous Key point:
    </span>
    <span class="multilang" lang="de">
        Wir verwenden die in einem früheren Schlüsselpunkt angegebene Formel:
    </span>
</p>
<p>
    \(\underline{p}\cdot\underline{q}=({@p1@})(6) + (-{@p2@})(-1) + (7)({@q3@}) = {@6*p1@}+{@p2@}+{@7*q3@}={@ta@}\)
</p>
<p>
    <span class="multilang" lang="en">
        Note again that the result is a scalar: there are no \(\underline{ i}\)'s,  \(\underline{ j}\)'s, or \(\underline{ k}\)'s in the answer.
    </span>
    <span class="multilang" lang="de">
        Beachten Sie noch einmal, dass das Ergebnis ein Skalar ist: es gibt keine \(\underline{ i}\)'s,  \(\underline{ j}\)'s oder \(\underline{ k}\)'s in der Antwort.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>p1: 2 + rand(6);
p2: 2 + rand(6);
q3: 2 + rand(6);
ta: p1*6 + p2 + 7*q3;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ ({@p1@}\underline{ i}-{@p2@}\underline{ j}+7\underline{ k}) \cdot (6\underline{ i}-\underline{ j}+{@q3@}\underline{ k}) = {@ta@}. \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer. Well done!
</span>
<span class="multilang" lang="de">
    Richtige Antwort. Gut gemacht!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1339471212</deployedseed>
        <deployedseed>1006240315</deployedseed>
        <deployedseed>1602091065</deployedseed>
        <deployedseed>411016882</deployedseed>
        <deployedseed>1802144806</deployedseed>
        <deployedseed>520364281</deployedseed>
        <deployedseed>541646464</deployedseed>
        <deployedseed>451782586</deployedseed>
        <deployedseed>1290763240</deployedseed>
        <deployedseed>1429828326</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>40</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
