<quiz>
    <question type="stack">
        <name>
            <text>9.3.4.7 Work done by a force</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The work done by a force \(\underline{F}\) in moving a body through a displacement \(\underline{r}\) is given by \(\underline{F}\cdot\underline{r}\).
                Find the work done by the force \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) if it causes a body to move from the point with coordinates \((1,1,2)\) to the point \(({@p2x@},3,5)\).
    </span>
    <span class="multilang" lang="de">
        Die von einer Kraft \(\underline{F}\) verrichtete Arbeit bei der Bewegung eines Körpers durch eine Verschiebung \(\underline{r}\) ist gegeben durch
                \(\underline{F}\cdot\underline{r}\). Finden Sie die von der Kraft \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) verrichtete Arbeit, wenn sie einen Körper vom Punkt mit den Koordinaten
                \((1,1,2)\) zum Punkt \(({@p2x @},3,5)\) bewegt.
    </span>
</p>
<p>
    [[input:ans1]]
    <span class="multilang" lang="en">
        units
    </span>
    <span class="multilang" lang="de">
        Einheiten
    </span>
    [[validation:ans1]]
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        First, we proceed to compute the displacement vector \(\underline{r}\) subtracting the final point \(({@p2x@},3,5)\) minus the initial point \((1,1,2)\):
    </span>
    <span class="multilang" lang="de">
        Zunächst berechnen wir den Verschiebungsvektor \(\underline{r}\) und subtrahieren den Endpunkt \(({@p2x@},3,5)\) minus Anfangspunkt \((1,1,2)\):
    </span>
</p>
<p>
    \(\underline{r} = ({@p2x@},3,5) - (1,1,2) = {@p2x -1@}\underline{i} + 2\underline{j} + 3\underline{k}\)
</p>
<p>
    <span class="multilang" lang="en">
        Then, we proceed to compute the work done by the force \(\underline{F}=3\underline{i}+{@f2@}\underline{k}\) using the formula given in the statement:
    </span>
    <span class="multilang" lang="de">
        Dann fahren wir fort, die von der Kraft \(\underline{F}=3\underline{i}+{@f2@}\underline{k} \) geleistete Arbeit mit der in der Anweisung angegebenen Formel zu berechnen:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(W = \underline{F}\cdot\underline{r} = (3)({@p2x -1@})+(0)(2)+({@f2@})(3) = {@ta@} \) units.
    </span>
    <span class="multilang" lang="de">
        \(W = \underline{F}\cdot\underline{r} = (3)({@p2x -1@})+(0)(2)+ ({@f2@})(3) = {@ta@} \) Einheiten.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>f2: 3 + rand(4)
p2x: 3 + rand(4)
ta: 3*(p2x -1) + f2*3</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\(W = (3\underline{i}+{@f2@}\underline{k}) \cdot (({@p2x@}-1)\underline{i} + 2 \underline{j} + 3 \underline{k}) = {@ta@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer.
</span>
<span class="multilang" lang="de">
    Richtige Antwort.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1596035697</deployedseed>
        <deployedseed>428490799</deployedseed>
        <deployedseed>24406443</deployedseed>
        <deployedseed>1578315134</deployedseed>
        <deployedseed>1013805574</deployedseed>
        <deployedseed>108770404</deployedseed>
        <deployedseed>1460721703</deployedseed>
        <deployedseed>223450315</deployedseed>
        <deployedseed>526907617</deployedseed>
        <deployedseed>492674192</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>20</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
