<quiz>
    <question type="stack">
        <name>
            <text>9.3.4.1 Computing the scalar product of two vectors</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If \(\underline{a}={@a1@}\underline{i}-{@a2@}\underline{j}\) and \(\underline{b}=3\underline{i}+{@b2@}\underline{j}\) find \(\underline{a}\cdot\underline{b}\)
    </span>
    <span class="multilang" lang="de">
        Wenn \(\underline{a}={@a1@}\underline{i}-{@a2@}\underline{ j}\) und \(\underline{b}=3\underline{i}+{@b2@}\underline{j}\) finden Sie \(\underline{a}\cdot\underline{b}\)
    </span>
</p>
<p>
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
<p>
    [[feedback:prt1]]
</p>
<span class="multilang" lang="en">
    Now find \(\underline{b}\cdot\underline{a}\)
</span>
<span class="multilang" lang="de">
    Finden Sie jetzt \(\underline{b}\cdot\underline{a}\)
</span>
<p>
</p>
<p>
</p>
<p>
    [[input:ans2]] [[validation:ans2]]
</p>
<p>
    [[feedback:prt2]]
</p>
<p>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We use the formula given previously:
    </span>
    <span class="multilang" lang="de">
        Wir verwenden die zuvor angegebene Formel:
    </span>
</p>
<p>
    \(\underline{a}\cdot\underline{b}=({@a1@})(3) + (-{@a2@})({@b2@}) = {@3*a1@}-{@a2*b2@}={@3*a1 - a2*b2@}\)
</p>
<p>
    <span class="multilang" lang="en">
        And also remember that \(\underline{a}\cdot\underline{b} = \underline{b}\cdot\underline{a}\) so both answers are the same.
    </span>
    <span class="multilang" lang="de">
        Und denken Sie auch daran, dass \(\underline{a}\cdot\underline{b} = \underline{b}\cdot\underline{a}\) also beide Antworten sind das gleiche.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: 2 + rand(4);
a2: 2 + rand(4);
b2: 2 + rand(4);
ta: 3*a1 - a2*b2;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ \underline{a}\cdot\underline{b}= \underline{a}\cdot\underline{b} = ({@a1@}\underline{i}-{@a2@}\underline{j}) \cdot (3\underline{i}+{@b2@}\underline{j}) = {@ta@}  \]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer.
</span>
<span class="multilang" lang="de">
    Richtige Antwort.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>703601437</deployedseed>
        <deployedseed>634839485</deployedseed>
        <deployedseed>2065096916</deployedseed>
        <deployedseed>621805973</deployedseed>
        <deployedseed>1009654782</deployedseed>
        <deployedseed>878404379</deployedseed>
        <deployedseed>1385260672</deployedseed>
        <deployedseed>952685503</deployedseed>
        <deployedseed>1105128643</deployedseed>
        <deployedseed>1161583249</deployedseed>
        <deployedseed>608791516</deployedseed>
        <deployedseed>1786104960</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(ta+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
