<quiz>
    <question type="stack">
        <name>
            <text>9.3.4.6 True or False</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The following vectors \({@p1@}\underline{i}+{@p2@}\underline{j}\) and \({@q1@}\underline{i}-{@q2@}\underline{j}\) perpendicular. True or False?
    </span>
    <span class="multilang" lang="de">
        Die folgenden Vektoren \({@p1@}\underline{i}+{@p2@}\underline{j}\) und \({@q1@}\underline{i}-{@q2@}\underline{j}\) stehen senkrecht zueinander.
                Richtig oder falsch?
    </span>
</p>
<p>
    [[input:ans1]] [[validation:ans1]]
</p>
<p>
    [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        If the scalar product of these two vectors is zero, then both vectors are perpendicular. Otherwise, they are not.
    </span>
    <span class="multilang" lang="de">
        Wenn das Skalarprodukt dieser beiden Vektoren null ist, dann stehen beide Vektoren senkrecht zueinander. Andernfalls sind sie es nicht.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In this case,
    </span>
    <span class="multilang" lang="de">
        In diesem Fall
    </span>
</p>
\[({@p1@}\underline{i} + {@p2@}\underline{j}) \cdot ({@q1@}\underline{i} + {@q2@}\underline{j}) = {@prod@} = {@prods@}.\]
<p>
    <span class="multilang" lang="en">
        Therefore, the two vectors are {@clas@}.
    </span>
    <span class="multilang" lang="de">
        Daher stehen die beiden Vektoren {@clasde@} zueinander.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[p1: 2 + rand(2);
p2: 2 + rand(2);
q1: 2 + 2*rand(2);
q2: 2 + rand(2);
ta: is(p1*q1 - p2*q2 = 0);

/* Needed for solution */
simp: false;
prod: p1 * q1 + p2 * q2;
prods: ev(prod, simp);
clas: if ta = true
  then "perpendicular"
  else "not perpendicular";
clasde: if ta = true
  then "senkrecht"
  else "nicht senkrecht";]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>Is \( {@p1@}\underline{i}+{@p2@}\underline{j} \cdot  {@q1@}\underline{i}-{@q2@}\underline{j}  = 0\)? {@ta@}</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer.
</span>
<span class="multilang" lang="de">
    Richtige Antwort.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>boolean</type>
            <tans>ta</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>PropLogic</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>857602432</deployedseed>
        <deployedseed>1188613156</deployedseed>
        <deployedseed>246979991</deployedseed>
        <deployedseed>178036310</deployedseed>
        <deployedseed>2015206763</deployedseed>
        <deployedseed>1545416990</deployedseed>
        <deployedseed>1057868805</deployedseed>
        <deployedseed>1994237709</deployedseed>
        <deployedseed>1073126526</deployedseed>
        <deployedseed>1561753378</deployedseed>
        <deployedseed>919952378</deployedseed>
        <deployedseed>1900424729</deployedseed>
        <deployedseed>870781741</deployedseed>
        <deployedseed>698903587</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(not p1*q1 - p2*q2= 0,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
