<quiz>
    <question type="stack">
        <name>
            <text>1.1.6.2 Berechnung der Veränderung mit Einheiten [Calculating Change with Units]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The pressure \(p\) in a reaction vessel changes from {@au1@} to {@au2@}. Write down the value of \(\delta p\).
    </span>
    <span class="multilang" lang="de">
        Der Druck \(p\) in einem Reaktionsgefäß ändert sich von {@au1@} auf {@au2@}. Notieren Sie den Wert von \(\delta p\).
    </span>
</p>
<p>
    <br/>
</p>
<p>
    \(\delta p =\) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        \(\delta p\) stands for the change in the value of p hence we have:
    </span>
    <span class="multilang" lang="de">
        \(\delta p\) steht für die Änderung des Wertes von p, also haben wir:
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        \(\qquad \delta p = {@au2@} - {@au1@} = {@tau@}\)
    </span>
    <span class="multilang" lang="de">
        \(\qquad \delta p = {@au2@} - {@au1@} = {@tau@}\)
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>stack_unit_si_declare(true);

a1: rand([34,35,36]);
au1:stackunits(a1, Pa);
a2: rand(10)+39;
au2:stackunits(a2, Pa);

tau: stackunits(a2-a1, Pa);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\(\delta p = {@au2@} - {@au1@} = {@tau@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>units</type>
            <tans>tau</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>UnitsStrictRelative</answertest>
                <sans>ans1</sans>
                <tans>tau</tans>
                <testoptions>0</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1558446885</deployedseed>
        <deployedseed>1559242441</deployedseed>
        <deployedseed>113399676</deployedseed>
        <deployedseed>1991351493</deployedseed>
        <deployedseed>1011989010</deployedseed>
        <deployedseed>1347782777</deployedseed>
        <deployedseed>396599599</deployedseed>
        <deployedseed>701458948</deployedseed>
        <deployedseed>1451975064</deployedseed>
        <deployedseed>1371363943</deployedseed>
        <deployedseed>2133880934</deployedseed>
        <deployedseed>1950249308</deployedseed>
        <deployedseed>1270206314</deployedseed>
        <deployedseed>640656554</deployedseed>
        <deployedseed>1349011307</deployedseed>
        <deployedseed>220660749</deployedseed>
        <deployedseed>1734317938</deployedseed>
        <deployedseed>1923882075</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>tau</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>a2-a1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore></expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
