<quiz>
    <question type="stack">
        <name>
            <text>8.3.2.1 General Gauss Elimination + Final exercise on back substitution</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<h5>
    <span class="multilang" lang="en">
        The general system of three simultaneous linear equations
    </span>
    <span class="multilang" lang="de">
        Das allgemeine System von drei simultanen linearen Gleichungen
    </span>
</h5>
<p>
    <span class="multilang" lang="en">
        In the previous section we met systems of equations which could be solved by back-substitution alone. In this Section we meet systems which are not so amenable and where preliminary work must be done before back-substitution can be used.
    </span>
    <span class="multilang" lang="de">
        Im vorherigen Abschnitt haben wir Gleichungssysteme kennengelernt, die allein durch Rücksubstitution gelöst werden können. In diesem Abschnitt treffen wir auf Systeme, die nicht so zugänglich sind und bei denen Vorarbeiten durchgeführt werden müssen, bevor die Rücksubstitution eingesetzt werden kann.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Consider the system \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                2x_1 &amp; - &amp; x_2 &amp; - &amp; 3x_3 &amp; = &amp; 3 \\
                4x_1 &amp; + &amp; 5x_2 &amp; - &amp; x_3 &amp; = &amp; 7 \\
\end{matrix}\] We will use the solution method known as
        <strong>
            Gauss elimination
        </strong>
        , which has three stages. In the first stage the equations are written in matrix form. In the second stage the matrix equations are replaced
                by a system of equations having the same solution but which are in
        <strong>
            triangular form
        </strong>
        . In the final stage the new system is solved by
        <strong>
            back-substitution
        </strong>
        .
    </span>
    <span class="multilang" lang="de">
        Betrachten Sie das System \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                2x_1 &amp; - &amp; x_2 &amp; - &amp; 3x_3 &amp; = &amp; 3 \\
                4x_1 &amp; + &amp; 5x_2 &amp; - &amp; x_3 &amp; = &amp; 7 \\
\end{matrix}\] Wir verwenden die als
        <strong>
            Gauss-Eliminierung
        </strong>
        bekannte Lösungsmethode, die aus drei Stufen besteht. Im ersten Schritt werden die Gleichungen in Matrixform geschrieben. Im zweiten Schritt werden die Matrixgleichungen ersetzt
                durch ein Gleichungssystem, das die gleiche Lösung hat, aber in
        <strong>
            dreieckiger Form
        </strong>
        vorliegt. Im letzten Schritt wird das neue System durch
        <strong>
            Rücksubstitution
        </strong>
        gelöst.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        <strong>
            Stage 1
        </strong>
        : Matrix Formulation
    </span>
    <span class="multilang" lang="de">
        <strong>
            Stufe 1
        </strong>
        : Matrixformulierung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The first step is to write the equations in matrix form:
    </span>
    <span class="multilang" lang="de">
        Der erste Schritt besteht darin, die Gleichungen in Matrixform zu schreiben:
    </span>
</p>
<p>
    <span class="math display">
        \[\left[\begin{matrix}
                1&amp;3&amp;5\\ 
                2&amp;-1&amp;-3\\
                4&amp;5&amp;-1 \\
\end{matrix}\right]
\left[\begin{matrix}
                x_1\\
                x_2\\
                x_3 \\
\end{matrix}\right] = 
\left[\begin{matrix} 
                14\\
                3\\
                7\\
\end{matrix}\right].\]
    </span>
    <span class="multilang" lang="en">
        Then, for conciseness, we combine the matrix of coefficients with the column vector of right-hand sides to produce the
        <strong>
            augmented matrix
        </strong>
        :
    </span>
    <span class="multilang" lang="de">
        Dann kombinieren wir der Kürze halber die Koeffizientenmatrix mit dem Spaltenvektor der rechten Seiten, um die
        <strong>
            erweiterte Matrix
        </strong>
        zu erstellen:
    </span>
</p>
<p>
    \[\left[\begin{array}{rrr|r} 1 &amp; 3 &amp; 5 &amp; 14 \\ 2 &amp; -1 &amp; -3 &amp; 3 \\ 4 &amp; 5 &amp; -1 &amp; 7 \\ \end{array}\right].\]
</p>
<p>
    <span class="multilang" lang="en">
        If the general system of equations is written \(A X =B\) then the augmented matrix is written \([A|B]\). Hence the first equation \[x_1+3x_2+5x_3=14\]    is replaced by the first row of the augmented matrix, \[1 \qquad 3\qquad 5\qquad |\qquad 14\] and so on. Stage 1 has now been completed. We will next triangularise the matrix of coefficients by means of
        <strong>
            row operations
        </strong>
        .
                There are three possible row operations:
    </span>
    <span class="multilang" lang="de">
        Wenn das allgemeine Gleichungssystem als \(A X =B\) geschrieben wird, dann wird die erweiterte Matrix als \([A|B]\) geschrieben. Daher wird die erste Gleichung \[x_1+3x_2+5x_3=14\] durch die erste Zeile der erweiterten Matrix ersetzt, \[1 \qquad 3\qquad 5\qquad |\qquad 14\] und so weiter. Stufe 1 ist nun abgeschlossen. Als nächstes triangularisieren wir die Koeffizientenmatrix mittels
        <strong>
            Zeilenoperationen
        </strong>
        .
                Es gibt drei mögliche Zeilenoperationen:
    </span>
</p>
<ul>
    <li>
        <p>
            <span class="multilang" lang="en">
                interchange two rows;
            </span>
            <span class="multilang" lang="de">
                Zwei Zeilen vertauschen;
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                multiply or divide a row by a non-zero constant factor;
            </span>
            <span class="multilang" lang="de">
                Eine Zeile mit einem konstanten Faktor ungleich Null multiplizieren oder dividieren;
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                add to, or subtract from, one row a multiple of another row.
            </span>
            <span class="multilang" lang="de">
                Zu einer Zeile ein Vielfaches einer anderen Zeile addieren oder davon subtrahieren.
            </span>
        </p>
    </li>
</ul>
<p>
    <span class="multilang" lang="en">
        Note that interchanging two rows of the augmented matrix is equivalent to interchanging the two corresponding equations. The shorthand notation we use is introduced by example. To interchange row 1 and row 3 we write \(R1\leftrightarrow R3\).
                To divide row 2 by 5 we write \(R2\div 5\). To add three times row 1 to row 2, we write \(R2+3R1\). In the Task which follows you will see where these annotations are placed.
    </span>
    <span class="multilang" lang="de">
        Beachten Sie, dass das Vertauschen von zwei Zeilen der erweiterten Matrix dem Vertauschen der beiden entsprechenden Gleichungen entspricht. Die von uns verwendete Kurzschreibweise wird anhand eines Beispiels eingeführt. Um Zeile 1 und Zeile 3 zu vertauschen, schreiben wir \(R1\leftrightarrow R3\).
                Um Zeile 2 durch 5 zu teilen, schreiben wir \(R2\div 5\). Um dreimal Zeile 1 zu Zeile 2 zu addieren, schreiben wir \(R2+3R1\). In der folgenden Aufgabe sehen Sie, wo diese Anmerkungen platziert sind.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Note that these operations neither create nor destroy solutions so that at every step the system of equations has the same solution as the original system.
    </span>
    <span class="multilang" lang="de">
        Beachten Sie, dass diese Operationen weder Lösungen erstellen noch zerstören, sodass das Gleichungssystem bei jedem Schritt die gleiche Lösung wie das ursprüngliche System hat.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        <strong>
            Stage 2
        </strong>
        : Triangularisation
    </span>
    <span class="multilang" lang="de">
        <strong>
            Stufe 2
        </strong>
        : Triangularisierung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The second stage proceeds by first eliminating \(x_1\) from the second and third equations using row operations.
    </span>
    <span class="multilang" lang="de">
        In der zweiten Stufe wird zunächst \(x_1\) mithilfe von Zeilenoperationen aus der zweiten und dritten Gleichung eliminiert.
    </span>
</p>
<p>
    <span class="math display">
        \[\left[\begin{array}{rrr|r} 
                1 &amp; 3 &amp; 5 &amp; 14 \\
                2 &amp; -1 &amp; -3 &amp; 3 \\
                4 &amp; 5 &amp; -1 &amp; 7 
\end{array}\right]
\begin{array}{c}
                \\ R2-2 \times R1 \\
                R3-4 \times R1 \\
\end{array}
\Rightarrow 
\left[\begin{array}{rrr|r}
                1 &amp; 3 &amp; 5 &amp; 14 \\
                0 &amp; -7 &amp; -13 &amp; -25 \\
                0 &amp; -7 &amp; -21 &amp; -49 \\
\end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In the above we have subtracted twice row (equation) 1 from row (equation) 2.
    </span>
    <span class="multilang" lang="de">
        Oben haben wir Zeile (Gleichung) 1 zweimal von Zeile (Gleichung) 2 subtrahiert.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In full these operations would be written, respectively, as \[(2x_1-x_2-3x_3)-2(x_1+3x_2+5x_3)=3-2\times 14 \qquad \text{or} \qquad -7x_2-13x_3=-25\] and \[(4x_1+5x_2-x_3)-4(x_1+3x_2+5x_3)=7-4\times 14 \qquad \text{or} \qquad -7x_2-21x_3=-49.\]
    </span>
    <span class="multilang" lang="de">
        Vollständig würden diese Operationen jeweils als \[(2x_1-x_2-3x_3)-2(x_1+3x_2+5x_3)=3-2\times 14 \qquad \text{oder} \qquad -7x_2-13x_3=-25\] und \[(4x_1+5x_2-x_3)-4(x_1+3x_2+5x_3)=7-4\times 14 \qquad \text{oder} \qquad -7x_2-21x_3=-49\] geschrieben werden.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now since all the elements in rows 2 and 3 are negative we multiply throughout by \(-1\), this is not really necessary but it will clean up the matrix a bit: \[\left[\begin{array}{rrr|r} 
                1 &amp; 3 &amp; 5 &amp; 14 \\
                0 &amp; -7 &amp; -13 &amp; -25 \\
                0 &amp; -7 &amp; -21 &amp; -49 
\end{array}\right]
\begin{array}{c}
                \\ 
                R2\times (-1) \\
                R3\times (-1) \\
\end{array}
\Rightarrow
\left[\begin{array}{rrr|r}
                1 &amp; 3 &amp; 5 &amp; 14 \\
                0 &amp; 7 &amp; 13 &amp; 25 \\
                0 &amp; 7 &amp; 21 &amp; 49 
\end{array}\right]\]
    </span>
    <span class="multilang" lang="de">
        Da nun alle Elemente in den Zeilen 2 und 3 negativ sind, multiplizieren wir durchgehend mit \(-1\). Dies ist nicht wirklich notwendig, aber es wird die Matrix ein wenig aufräumen: \[\left[\begin{array}{rrr|r} 
                1 &amp; 3 &amp; 5 &amp; 14 \\
                0 &amp; -7 &amp; -13 &amp; -25 \\
                0 &amp; -7 &amp; -21 &amp; -49 
\end{array}\right]
\begin{array}{c}
                \\ 
                R2\times (-1) \\
                R3\times (-1) \\
\end{array}
\Rightarrow
\left[\begin{array}{rrr|r}
                1 &amp; 3 &amp; 5 &amp; 14 \\
                0 &amp; 7 &amp; 13 &amp; 25 \\
                0 &amp; 7 &amp; 21 &amp; 49 
\end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Finally, we eliminate \(x_2\) from the third equation by subtracting equation 2 from equation 3, i.e., \(R3-R2\):
    </span>
    <span class="multilang" lang="de">
        Schließlich eliminieren wir \(x_2\) aus der dritten Gleichung, indem wir Gleichung 2 von Gleichung 3 subtrahieren, d. h. \(R3-R2\):
    </span>
    <span class="math display">
        \[\left[\begin{array}{rrr|r}
                1&amp;3&amp;5&amp;14\\
                0&amp;7&amp;13&amp;25\\
                0&amp;7&amp;21&amp;49 \\
\end{array}\right]
\begin{array}{c}
                \\
                \\
                R3-R2 \\
\end{array}
\Rightarrow
\left[\begin{array}{rrr|r}
                1&amp;3&amp;5&amp;14\\
                0&amp;7&amp;13&amp;25\\
                0&amp;0&amp;8&amp;24 \\
\end{array}\right]\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The system is now in triangular form.
    </span>
    <span class="multilang" lang="de">
        Das System hat jetzt die Dreiecksform.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        <strong>
            Stage 3
        </strong>
        : Back Substitution
    </span>
    <span class="multilang" lang="de">
        <strong>
            Stufe 3
        </strong>
        : Rücksubstitution
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Here we solve the equations from bottom to top. At each step of the back substitution process we encounter equations which only have a
        <strong>
            single
        </strong>
        unknown and so can be easily solved.
    </span>
    <span class="multilang" lang="de">
        Hier lösen wir die Gleichungen von unten nach oben. Bei jedem Schritt des Rücksubstitutionsprozesses stoßen wir auf Gleichungen, die nur eine einzige Unbekannte haben und daher leicht gelöst werden können.
    </span>
</p>
<hr/>
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Complete the solution to the above system by back-substitution.
    </span>
    <span class="multilang" lang="de">
        Vervollständigen Sie die Lösung des obigen Systems durch Rücksubstitution.
    </span>
</p>
<div>
    \(x=\) [[input:ans1]] [[validation:ans1]]
</div>
<div>
    <br/>
</div>
<p>
    <span class="multilang" lang="en">
        Check that these values satisfy the original system of equations.
    </span>
    <span class="multilang" lang="de">
        Überprüfen Sie, ob diese Werte dem ursprünglichen Gleichungssystem entsprechen.
    </span>
</p>
<p>
    [[feedback:prt1]]
    <br/>
</p>
<hr/>
<p>
    <span class="multilang" lang="en">
        You should always check your solution in the original equations to ensure that the equations balance.
    </span>
    <span class="multilang" lang="de">
        Sie sollten Ihre Lösung immer in den Originalgleichungen überprüfen, um sicherzustellen, dass die Gleichungen ausgewogen sind.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Again we emphasise that we chose a particular set of procedures in Stage 1. This was chosen mainly to keep the arithmetic simple by delaying the introduction of fractions. Sometimes we are courageous and take fewer, harder steps.
    </span>
    <span class="multilang" lang="de">
        Wir betonen noch einmal, dass wir in Stufe 1 eine bestimmte Menge von Verfahren ausgewählt haben. Dies wurde hauptsächlich gewählt, um die Arithmetik einfach zu halten, indem die Einführung von Brüchen verzögert wurde. Manchmal sind wir mutig und unternehmen weniger, dafür härtere Schritte.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        An important point to note is that when in Stage 2 we wrote \(R2-2R1\) against row 2; what we meant is that row 2 is replaced by the combination (row \(2)-2\times\)(row 1).
    </span>
    <span class="multilang" lang="de">
        Ein wichtiger Punkt ist, dass wir in Stufe 2 \(R2-2R1\) für Zeile 2 geschrieben haben; Was wir meinten ist, dass Zeile 2 durch die Kombination (Zeile \(2)-2\times\)(Zeile 1) ersetzt wird.
    </span>
</p>
<div class="HELM_keypoint">
    <h4>
        <span class="multilang" lang="en">
            Key Point: Elementary Row Operations
        </span>
        <span class="multilang" lang="de">
            Schlüsselpunkt: Elementare Zeilenoperationen
        </span>
    </h4>
    <p>
        <span class="multilang" lang="en">
            The operation \[\text{ row } i-\alpha \times \text{ row }\ j\] means replace
            <strong>
                row \(i\)
            </strong>
            by the combination \[\text{ row } i-\alpha \times \text{ row } j.\]
        </span>
        <span class="multilang" lang="de">
            Die Operation \[\text{ Zeile } i-\alpha \times \text{ Zeile }\ j\] bedeutet Ersetzen von
            <strong>
                Zeile \(i\)
            </strong>
            durch die Kombination \[\text{ Zeile } i-\alpha \times \text{ Zeile } j.\]
        </span>
    </p>
</div>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        In full the equations are \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                &amp; &amp; 7x_2 &amp; + &amp; 13x_3 &amp; = &amp; 25 \\
                &amp; &amp; &amp; &amp; 8x_3 &amp; = &amp; 24 \\
\end{matrix}\] From the last equation we see that \(x_3=3\). Substituting this value into the second equation gives \[7x_2+39=25 \qquad \text{or} \qquad 7x_2=-14 \qquad \text{so that} \qquad x_2=-2\] Finally, using these values for \(x_2\) and \(x_3\) in equation 1 gives \[x_1-6+15=14.\] Hence \(x_1=5\). The solution
                is therefore \([x_1,x_2,x_3]^T=[5,-2,3]^T\). You can check that these values satisfy the original system of equations.
    </span>
    <span class="multilang" lang="de">
        Die vollständigen Gleichungen lauten \[\begin{matrix}
                x_1 &amp; + &amp; 3x_2 &amp; + &amp; 5x_3 &amp; = &amp; 14 \\
                &amp; &amp; 7x_2 &amp; + &amp; 13x_3 &amp; = &amp; 25 \\
                &amp; &amp; &amp; &amp; 8x_3 &amp; = &amp; 24 \\
\end{matrix}\] Aus der letzten Gleichung sehen wir, dass \(x_3=3\). Das Einsetzen dieses Werts in die zweite Gleichung ergibt \[7x_2+39=25 \qquad \text{oder} \qquad 7x_2=-14 \qquad \text{sodass} \qquad x_2=-2\] Schließlich werden diese Werte für verwendet \(x_2\) und \(x_3\) in Gleichung 1 ergibt \[x_1-6+15=14.\] Daher \(x_1=5\). Die Lösung
                ist also \([x_1,x_2,x_3]^T=[5,-2,3]^T\). Sie können überprüfen, ob diese Werte das ursprüngliche Gleichungssystem erüllen.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ordergreat(x, y, z);
/* Input the system as in the example*/
a11 : 1;
a12 : 3;
a13 : 5;
a22 : 7;
a23 : 13;
a33 : 8;
matA : matrix([a11, a12, a13], [0, a22, a23], [0, 0, a33]);
vecB : matrix([14], [25], [24]);
vecX : matrix([x], [y], [z]);

/* Construct the solution */
aux_LHS : matA.vecX;
eq1 : aux_LHS[1][1] = vecB[1][1];
eq2 : aux_LHS[2][1] = vecB[2][1];
eq3 : aux_LHS[3][1] = vecB[3][1];
ta_aux : algsys([eq1, eq2, eq3], [x, y, z]);
ta : matrix([rhs(ta_aux[1][1])], [rhs(ta_aux[1][2])], [rhs(ta_aux[1][3])])</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Well done!
</span>
<span class="multilang" lang="de">
    Gut gemacht!
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="en">
    Be careful, your answer doesn't satisfy the system \[ {@matA@} {@ans1@} = {@matA . ans1@} \neq {@vecB@} \].
</span>
<span class="multilang" lang="de">
    Seien Sie vorsichtig, Ihre Antwort erfüllt nicht das System \[ {@matA@} {@ans1@} = {@matA . ans1@} \neq {@vecB@} \].
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(2*ta,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
