<quiz>
    <question type="stack">
        <name>
            <text>4.5.4.4 Übung 4 Wassertiefe [Exercise 4 Water depth]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The depth \(y\) of water at the entrance to a small harbour at time \(t\ (\mathrm{h})\) is
    </span>
    <span class="multilang" lang="de">
        Die Wassertiefe \(y\) am Eingang eines kleinen Hafens zur Zeit \(t\ (\mathrm{h})\) beträgt
    </span>
</p>
\[ y=a\sin\left( b\left(t-\dfrac{\pi}{2}\right)\right)+k,\]
<p>
    <span class="multilang" lang="en">
        where \(k\ (\mathrm{m})\) is the average depth.
    </span>
    <span class="multilang" lang="de">
        wobei \(k\ (\mathrm{m})\) die durchschnittliche Tiefe ist.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If the tidal period is \({@na1t@}\ \mathrm{h}\) and the depths at high tide and low tide are \( {@na1h@}\ \mathrm{m}\) and \( {@na1l@}\ \mathrm{m}\) respectively, obtain \(a,\ b,\ k\).
    </span>
    <span class="multilang" lang="de">
        Wenn die Gezeitenperiode \({@na1t@}\ \mathrm{h}\) ist und die Tiefen bei Flut und Ebbe \( {@na1h@} \ \mathrm{m}\) bzw. \( {@na1l@}\ \mathrm{m}\), erhalten wir \(a,\ b,\ k\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please write "pi" for the value \( \pi\).
    </span>
    <span class="multilang" lang="de">
        Bitte schreiben Sie "pi" für den Wert \( \pi\).
    </span>
</p>
<p>
    \(a=\) [[input:ans1]] \(\mathrm{m}\)[[validation:ans1]] [[feedback:prt1]]
</p>
<p>
    \(b=\) [[input:ans2]] \(\mathrm{rad}\,\mathrm{h}^{-1}\) [[validation:ans2]] [[feedback:prt2]]
</p>
<p>
    \(k=\) [[input:ans3]] \(\mathrm{m}\) [[validation:ans3]] [[feedback:prt3]]
    <br/>
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We have \(y=a\sin\left( b\left(t-\dfrac{\pi}{2}\right)\right)+k.\)
    </span>
    <span class="multilang" lang="de">
        Gegeben ist \(y=a\sin\left( b\left(t-\dfrac{\pi}{2}\right)\right)+k.\)
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Since the period is \(\dfrac{2\pi}{b}= {@na1t@} \text{ hr},\) we obtain \(b= \displaystyle {@ta1b@}\ \mathrm{rad}\,\mathrm{h}^{-1}\).
    </span>
    <span class="multilang" lang="de">
        Da der Punkt \(\dfrac{2\pi}{b}= {@na1t@} \text{ hr},\) ist, erhalten wir \(b= \displaystyle {@ta1b@}\ \mathrm{rad}\,\mathrm{h}^{-1}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Furthermore,  \[y_{\max}=a+k \quad \text{and} \quad y_{\min}=-a+k\] and hence  \[a+k= {@na1h@} \quad \text{and} \quad -a+k={@na1l@}.\]
    </span>
    <span class="multilang" lang="de">
        Darüber hinaus  \[y_{\max}=a+k \quad \text{und} \quad y_{\min}=-a+k\] und somit  \[a+k= {@na1h@} \quad \text{und} \quad -a+k={@na1l@}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        If we sum these two equations, we obtain \(k=\displaystyle{{@ta1k@}} \ \mathrm{m}\), and if we subtract them, we find \(a=\displaystyle{{@ta1a@}} \ \mathrm{m}\).
    </span>
    <span class="multilang" lang="de">
        Wenn wir diese beiden Gleichungen summieren, erhalten wir \(k=\displaystyle{{@ta1k@}} \ \mathrm{m}\), und wenn wir sie subtrahieren, erhalten wir \(a=\displaystyle{{@ta1a@}} \ \mathrm{m}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Therefore, the function that models the depth of the water is
    </span>
    <span class="multilang" lang="de">
        Daher ist die Funktion, die die Wassertiefe modelliert,
    </span>
</p>
\[y=\displaystyle{{@ta1a@} \sin \left({@ta1b@}\left(t-\dfrac{\pi}{2}\right)\right)+{@ta1k@}}.\]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>na1t: rand(10)+2;
ta1b: (2*pi)/na1t;
na1l: rand(8)+2;
na1h: na1l+ rand(20);
ta1k: (na1l+na1h)/2;
ta1a:  ta1k-na1l;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ y=a\sin b\left(t-\dfrac{\pi}{2}\right)+k\]
\[a={@ta1a@}, \quad b={@ta1b@}, \quad k={@ta1k@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1a</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta1b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta1k</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1a</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta1b</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta1k</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1112334714</deployedseed>
        <deployedseed>1122743144</deployedseed>
        <deployedseed>693514011</deployedseed>
        <deployedseed>1261227194</deployedseed>
        <deployedseed>1766854782</deployedseed>
        <deployedseed>135363704</deployedseed>
        <deployedseed>980677847</deployedseed>
        <deployedseed>503271250</deployedseed>
        <deployedseed>385380930</deployedseed>
        <deployedseed>2081209704</deployedseed>
        <deployedseed>1701213164</deployedseed>
        <deployedseed>1385794456</deployedseed>
        <deployedseed>1429780616</deployedseed>
        <deployedseed>1340958527</deployedseed>
        <deployedseed>1663614042</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1a</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta1b</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta1k</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1k+na1l</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>na1t</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>na1l+na1h</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
