<quiz>
    <question type="stack">
        <name>
            <text>4.5.4.5 Übung 5 Temp [Exercise 5 Temp]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The Fahrenheit temperature at a certain location over one complete day is modelled by \[ \quad F(t)={@na1a@}+{@na1b@}\sin\left( \dfrac{\pi}{{@na1c@}}(t-{@na1d@})\right)\qquad 0\le  t\le 24\] where \(t\) is time in hours after midnight.
    </span>
    <span class="multilang" lang="de">
        Die Fahrenheit-Temperatur an einem bestimmten Ort über einen ganzen Tag wird modelliert durch \[ \quad F(t)={@na1a@}+{@na1b@}\sin\left( \dfrac{\pi}{{@na1c@}}(t-{@na1d@})\right)\qquad 0\le  t\le 24\] wobei \(t\) die Zeit in Stunden nach Mitternacht ist.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                What are the temperatures at \(8\) am and \(12\) noon?
            </span>
            <span class="multilang" lang="de">
                Wie hoch sind die Temperaturen um \(8\) Uhr morgens und \(12\) Uhr mittags?
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Give your answer to \(1\) decimal place.
            </span>
            <span class="multilang" lang="de">
                Geben Sie Ihre Antwort auf \(1\) Dezimalstelle.
            </span>
        </p>
        <p>
            \( F(8) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
        <br/>
        <p>
            \( F(12) = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                At what time(s) that day is the temperature \({@na1a@}\phantom{1}^\circ\mathrm{F}\)?
            </span>
            <span class="multilang" lang="de">
                Zu welcher Tageszeit ist die Temperatur \({@na1a@}\phantom{1}^\circ\mathrm{F}\)?
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Write your answer as a set
                <code>
                    {t1,t2,...}
                </code>
                .
            </span>
            <span class="multilang" lang="de">
                Schreiben Sie Ihre Antwort als Satz
                <code>
                    {t1,t2,...}
                </code>
                .
            </span>
        </p>
        <p>
            \( t = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Obtain the maximum and minimum temperatures and the times at which they occur.
            </span>
            <span class="multilang" lang="de">
                Ermitteln Sie die maximalen und minimalen Temperaturen und die Zeiten, zu denen sie auftreten.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                The maximum temperature is
            </span>
            <span class="multilang" lang="de">
                Die Höchsttemperatur beträgt
            </span>
            [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
            <span class="multilang" lang="en">
                \(\phantom{1}^\circ\mathrm{F}\) which occurs at \(t_{max} = \)
            </span>
            <span class="multilang" lang="de">
                \(\phantom{1}^\circ\mathrm{F}\), die bei \(t_{max} = \)
            </span>
            [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]
            <span class="multilang" lang="en">
            </span>
            <span class="multilang" lang="de">
                auftritt
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                The minimum temperature is
            </span>
            <span class="multilang" lang="de">
                Die Mindesttemperatur beträgt
            </span>
            [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]
            <span class="multilang" lang="en">
                \(\phantom{1}^\circ\mathrm{F}\) which occurs at \(t_{min} = \)
            </span>
            <span class="multilang" lang="de">
                \(\phantom{1}^\circ\mathrm{F}\), die bei \(t_{min} = \)
            </span>
            [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]
            <span class="multilang" lang="en">
            </span>
            <span class="multilang" lang="de">
                auftritt
            </span>
        </p>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            <span class="multilang" lang="en">
                When \(t=8\), then the temperature is \({@na1a@}+{@na1b@}\sin \dfrac{\pi}{{@na1c@}}(8-{@na1d@})={@ta1@}^\circ \,F.\)
            </span>
            <span class="multilang" lang="de">
                Wenn \(t=8\), dann ist die Temperatur \({@na1a@}+{@na1b@}\sin \dfrac{\pi}{{@ na1c@}}(8-{@na1d@})={@ta1@}^\circ \,F.\)
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                And for \(t=12\), we have \({@na1a@}+{@na1b@}\sin \dfrac{\pi}{{@na1c@}}(12-{@na1d@})={@ta2@}^\circ \, F\).
            </span>
            <span class="multilang" lang="de">
                Und für \(t=12\) erhalten wir \({@na1a@}+{@na1b@}\sin \dfrac{\pi}{{@na1c @}}(12-{@na1d@})={@ta2@}^\circ \, F\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                \(F(t)={@na1a@}\) when  \( \dfrac{\pi}{{@na1c@}}(t-{@na1d@})=0,\ \pi,\ \ 2\pi,\ldots\).
            </span>
            <span class="multilang" lang="de">
                \(F(t)={@na1a@}\) wenn  \( \dfrac{\pi}{{@na1c@}}(t-{@na1d@})=0,\ \pi,\ \ 2\pi,\ldots\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Therefore,  \(t-{@na1d@}=0,{@na1c@},{@na12@},\ldots\) and hence,  \(t={@na1d@},{@ta1t1@},{@ta1t2@},\ldots\) hours.
            </span>
            <span class="multilang" lang="de">
                Daher  \(t-{@na1d@}=0,{@na1c@},{@na12@},\ldots\) und somit  \(t={@na1d@},{@ta1t1@},{@ta1t2@},\ldots\) Stunden.
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Since \( t\le 24\), we obtain \(t={@na1d@} \, ({@na1d@}.00 \, \mathrm{am}) \) and \(t={@ta1t1@} \, ({@ta1t1s@}.00 \, \mathrm{pm}).\)
            </span>
            <span class="multilang" lang="de">
                Da \( t\le 24\) erhalten wir \(t={@na1d@}\) und \(t={@ta1t1@}\).
            </span>
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Maximum temperature is \({@ta3maxval@}^\circ\) F when  \(\dfrac{\pi}{{@na1c@}}(t-{@na1d@})=\dfrac{\pi}{2}\) i.e. at \(t={@ta3max@}\).
            </span>
            <span class="multilang" lang="de">
                Die maximale Temperatur beträgt \({@ta3maxval@}^\circ\) F, wenn  \(\dfrac{\pi}{{@na1c@}}(t-{@na1d@})=\dfrac{\pi}{2}\), also bei \(t={@ta3max@}\).
            </span>
        </p>
        <p>
            <span class="multilang" lang="en">
                Minimum temperature is \({@ta3minval@}^\circ\)F when  \(\dfrac{\pi}{{@na1c@}}(t-{@na1d@})=\dfrac{3\pi}{2}\) i.e. at \(t={@ta3min@}\).
            </span>
            <span class="multilang" lang="de">
                Die Mindesttemperatur beträgt \({@ta3minval@}^\circ\)F, wenn  \(\dfrac{\pi}{{@na1c@}}(t -{@na1d@})=\dfrac{3\pi}{2}\), also bei \(t={@ta3min@}\).
            </span>
        </p>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>na1a: rand(30)+40;
na1b: rand(15)+1;
na1c: 2*(rand(3)+6);
na12: 2*na1c;
na1d: rand(2)+7;
ta1t1: na1d+na1c;
ta111: set(na1d,ta1t1);
ta1112: set(na1d);
ta1113: set(ta1t1);
ta1t1s: ta1t1-12;
ta1t2: na1d+na12;

ta1: decimalplaces(float(na1a+na1b*sin((pi/na1c)*(8-na1d))),1);
ta2: decimalplaces(float(na1a+na1b*sin((pi/na1c)*(12-na1d))),1);
ta3maxval: na1a+na1b;
ta3minval: na1a-na1b;
ta3max: na1c/2+na1d;
ta3minold: -(na1c/2)+na1d;
ta3minold: 3*(na1c/2)+na1d-24;
ta3min: -(na1c/2)+na1d;
ca: NA</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ \quad F(t)={@na1a@}+{@na1b@}\sin \dfrac{\pi}{{@na1c@}}(t-{@na1d@})\qquad 0\le  t\le 24\]
\[ F(8)= {@ta1@}, \quad F(12)= {@ta2@} \]
\[ t_1= {@ta111@} \]
\[  T_{max}={@ta3maxval@}, \quad t_{max}= {@ta3max@}\]
\[  T_{min}={@ta3minval@}, \quad t_{min}= {@ta3min@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>maxdp:1, mindp:1</options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options>maxdp:1, mindp:1</options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta111</tans>
            <boxsize>8</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{?,?}</syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans4</name>
            <type>numerical</type>
            <tans>ta3maxval</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans5</name>
            <type>numerical</type>
            <tans>ta3max</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans6</name>
            <type>numerical</type>
            <tans>ta3minval</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans7</name>
            <type>numerical</type>
            <tans>ta3min</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>1</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>NumDecPlaces</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions>1</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>Sets</answertest>
                <sans>ans3</sans>
                <tans>ta111</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>Sets</answertest>
                <sans>intersect(ans3, ta111)</sans>
                <tans>ta1112</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>Sets</answertest>
                <sans>intersect(ans3, ta111)</sans>
                <tans>ta1113</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans4</sans>
                <tans>ta3maxval</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans5</sans>
                <tans>ta3max</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans6</sans>
                <tans>ta3minval</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt7</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans7</sans>
                <tans>ta3min</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt7-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans7</sans>
                <tans>ta3minold</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt7-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1018989699</deployedseed>
        <deployedseed>1816448532</deployedseed>
        <deployedseed>818983487</deployedseed>
        <deployedseed>255938086</deployedseed>
        <deployedseed>2131688418</deployedseed>
        <deployedseed>745957536</deployedseed>
        <deployedseed>742815408</deployedseed>
        <deployedseed>998473530</deployedseed>
        <deployedseed>1694348903</deployedseed>
        <deployedseed>584545500</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta111</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>ta3maxval</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>ta3max</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>ta3minval</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>ta3min</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt5-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt6-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt7-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta1+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ev(ta2+1,simp)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>{}</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>0</value>
            </testinput>
            <testinput>
                <name>ans5</name>
                <value>0</value>
            </testinput>
            <testinput>
                <name>ans6</name>
                <value>na1a+1</value>
            </testinput>
            <testinput>
                <name>ans7</name>
                <value>na1a+1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt3-3-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt4</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt4-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt5</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt5-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt6</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt6-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt7</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt7-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
