<quiz>
    <question type="stack">
        <name>
            <text>4.5.2.2 Berechnen Sie die Zeitverschiebung [Calculate the time shift]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Calculate the time shift between the waves
    </span>
    <span class="multilang" lang="de">
        Berechnen Sie die Zeitverschiebung zwischen den Kurven
    </span>
</p>
\[\begin{aligned} x_1 &amp;= {@naa@}\cos ({@nat@} t) \\ x_2 &amp;= {@naa@}\cos \left({@nat@} t +{@nas@}\right) \end{aligned}\]
<p>
    <span class="multilang" lang="en">
        where the time \(t\) is in seconds.
    </span>
    <span class="multilang" lang="de">
        wobei \(t\) die Zeit in Sekunden ist.
    </span>
</p>
<p>
    [[input:ans1]] \(\mathrm{s}\) [[validation:ans1]] [[feedback:prt1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        We are given the waves
    </span>
    <span class="multilang" lang="de">
        Gegeben sind die Kurven
    </span>
</p>
\[\begin{aligned} x_1 &amp;= {@naa@}\cos ({@nat@} t) \\ x_2 &amp;= {@naa@}\cos \left({@nat@} t +{@nas@}\right) \end{aligned}\]
<p>
    <span class="multilang" lang="en">
        Note firstly that the waves have the same amplitude {@naa@} and angular frequency \({@nat@}\) (corresponding to a common period \(\dfrac{2\pi}{{@nat@}}= \displaystyle{{@natt@}} \text{ s}\)) .
    </span>
    <span class="multilang" lang="de">
        Beachten Sie zunächst, dass die Kurven die gleiche Amplitude {@naa@} und Winkelfrequenz \({@nat@}\) haben (entsprechend einer gemeinsamen Periode \(\dfrac{ 2\pi}{{@nat@}}= \displaystyle{{@natt@}} \text{ s}\)) .
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Now \(\cos \left({@nat@} t+\displaystyle{{@nas@}}\right) = \cos \left({@nat@}\left( t +\displaystyle{{@nats@}}\right)\right)\) so \(x_1\left( t +\displaystyle{{@nats@}}\right)=x_2(t).\)
    </span>
    <span class="multilang" lang="de">
        Jetzt gilt \(\cos \left({@nat@} t+\displaystyle{{@nas@}}\right) = \cos \left({@nat@} \left( t +\displaystyle{{@nats@}}\right)\right)\) also \(x_1\left( t +\displaystyle{{@nats@}}\right)=x_2(t).\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In other words the time shift is \(\displaystyle{{@nats@}}\ \text{s}\), the wave \(x_2\)
        <b>
            leads
        </b>
        the wave \(x_1\) by this amount. Alternatively we could say that \(x_1\)
        <b>
            lags
        </b>
        \(x_2\) by \(\displaystyle{{@nats@}} s.\)
    </span>
    <span class="multilang" lang="de">
        Mit anderen Worten, die Zeitverschiebung ist \(\displaystyle{{@nats@}}\ \text{s}\), die Kurve \(x_2\) eilt der Kurve \(x_1\) um diesen Betrag
        <b>
            voraus
        </b>
        . Alternativ könnten wir sagen, dass \(x_1\)
        <b>
            hinter
        </b>
        \(x_2\) um \(\displaystyle{{@nats@}} s\) liegt.
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>naa: rand(5)+1;
nat: (rand(10)+1)*pi;
nas: pi/(rand(7)+1);
natt: 2*pi/(nat);
nats: nas/nat</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>time shift \(= {@nats@} s\)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>nats</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>nats</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>-nats</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>860100309</deployedseed>
        <deployedseed>1686941969</deployedseed>
        <deployedseed>1763433499</deployedseed>
        <deployedseed>628685542</deployedseed>
        <deployedseed>1303124280</deployedseed>
        <deployedseed>1884341729</deployedseed>
        <deployedseed>1017417140</deployedseed>
        <deployedseed>1646584130</deployedseed>
        <deployedseed>1789016378</deployedseed>
        <deployedseed>137173424</deployedseed>
        <deployedseed>581966465</deployedseed>
        <deployedseed>1056726628</deployedseed>
        <deployedseed>739587670</deployedseed>
        <deployedseed>876663158</deployedseed>
        <deployedseed>1209582072</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>nats</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>nat</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>-nats</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(nats+1,simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
