<quiz>
    <question type="stack">
        <name>
            <text>4.5.4.1 Übung 1 Amp und Phase [Exercise 1 Amp and Phase]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Write down the amplitude, R, and the period, T, of \( \displaystyle{y={@ta1@}\sin\left( {@na12@} t\right)} \).
    </span>
    <span class="multilang" lang="de">
        Schreiben Sie die Amplitude R und die Periode T von \( \displaystyle{y={@ta1@}\sin\left( {@na12@} t \right)} \).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Please write "pi" for the value \( \pi\).
    </span>
    <span class="multilang" lang="de">
        Bitte schreiben Sie "pi" für den Wert \( \pi\).
    </span>
</p>
<p>
    \( \textrm{R} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>
<p>
    \( \textrm{T} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        The wave equation \(y= \displaystyle{{@ta1@}\sin {@na12@}t} \) has amplitude \(\textrm{R} = {@ta1@} .\)
    </span>
    <span class="multilang" lang="de">
        Die Kurvengleichung \(y= \displaystyle{{@ta1@}\sin {@na12@}t} \) hat die Amplitude \(\textrm{R} = {@ta1@} .\)
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        The period is \(\textrm{T} = \dfrac{2\pi}{{@na12@}}=\displaystyle {@ta2@}\).
    </span>
    <span class="multilang" lang="de">
        Die Periode ist \(\textrm{T} = \dfrac{2\pi}{{@na12@}}=\displaystyle {@ta2@}\).
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Check: \(y(t+{@ta2@})={@ta1@}\sin ({@na12@}(t+{@ta2@}))={@ta1@}\sin ({@na12@} t+2\pi)= {@ta1@}\sin {@na12@} t=y(t)\).
    </span>
    <span class="multilang" lang="de">
        Prüfen Sie: \(y(t+{@ta2@})={@ta1@}\sin ({@na12@}(t+{@ta2@}))= {@ta1@}\sin ({@na12@} t+2\pi)= {@ta1@}\sin {@na12@} t=y(t)\).
    </span>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ta1: (rand(15)+1)/(rand(5)+1);
rr:rand(5);
na12: pi*(rr+1);
ta2: 2/(rr+1);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[ y={@ta1@}\sin\left( {@na12@} t\right) \]
\[ R={@ta1@}, \, d={@ta2@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1731391848</deployedseed>
        <deployedseed>24720440</deployedseed>
        <deployedseed>1313205285</deployedseed>
        <deployedseed>166216795</deployedseed>
        <deployedseed>493838254</deployedseed>
        <deployedseed>889037442</deployedseed>
        <deployedseed>1496942840</deployedseed>
        <deployedseed>1085325193</deployedseed>
        <deployedseed>1248217704</deployedseed>
        <deployedseed>43482512</deployedseed>
        <deployedseed>2072892909</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>na12</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>na12</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty></expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
