<quiz>
    <question type="stack">
        <name>
            <text>1.2.4.2 Potenzen vereinfachen [Simplify indices]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Use the laws of indices to simplify.
    </span>
    <span class="multilang" lang="de">
        Verwenden Sie zur Vereinfachung die Potenzgesetze.
    </span>
</p>
<ol class="HELM_parts">
    <li>
        \(\dfrac{a^{@a1@} \times a^{@a2@}}{a^{@a3@}}= \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>
        \(\dfrac{m^{@b1@} \times m^{@b2@}}{m^{@b3@}}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <span class="multilang" lang="en">
            Use the first law of indices to simplify the numerator:
        </span>
        <span class="multilang" lang="de">
            Verwenden Sie das erste Potenzgesetz, um den Zähler zu vereinfachen:
        </span>
    </li>
    <p>
        \(\dfrac{a^{@a1@} \times a^{@a2@}}{a^{@a3@}}= \dfrac{a^{{@a1@}+{@a2@}}}{a^{@a3@}}= \dfrac{a^{{@a1+a2@}}}{a^{@a3@}}.\)
    </p>
    <p>
        <span class="multilang" lang="en">
            Now use the second law to simplify the result:
        </span>
        <span class="multilang" lang="de">
            Verwenden Sie nun das zweite Gesetz, um das Ergebnis zu vereinfachen:
        </span>
    </p>
    <p>
        \(\dfrac{a^{{@a1+a2@}}}{a^{@a3@}}= a^{{@a1+a2@}-{@a3@}} = a^{@a4@}.\)
        <br/>
        <br/>
    </p>
    <li>
        <span class="multilang" lang="en">
            Use the first law of indices to simplify the numerator:
        </span>
        <span class="multilang" lang="de">
            Verwenden Sie das erste Potenzgesetz, um den Zähler zu vereinfachen:
        </span>
    </li>
    <p>
        \(\dfrac{m^{@b1@} \times m^{@b2@}}{m^{@b3@}}= \dfrac{m^{{@b1@}+{@b2@}}}{m^{@b3@}}= \dfrac{m^{{@b1+b2@}}}{m^{@b3@}}.\)
    </p>
    <p>
        <span class="multilang" lang="en">
            Now use the second law to simplify the result:
        </span>
        <span class="multilang" lang="de">
            Verwenden Sie nun das zweite Gesetz, um das Ergebnis zu vereinfachen:
        </span>
    </p>
    <p>
        \(\dfrac{m^{{@b1+b2@}}}{m^{@b3@}}= m^{{@b1+b2@}-{@b3@}} = m^{@b4@}.\)
        <br/>
        <br/>
    </p>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>a1: rand(3)+8;
a2: a1-1;
a3: rand(3)+3;
a4: a1+a2-a3;
ta1: a^a4;

b1: rand(3)+8;
b3: -1*(rand(3)+3);
b2: b3-1;
b4: b1+b2-b3;
ta2: m^b4;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\( \dfrac{a^{@a1@} \times a^{@a2@}}{a^{@a3@}}={@ta1@}, \quad  \dfrac{m^{@b1@} \times m^{@b2@}}{m^{@b3@}}={@ta2@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>-,+,/,*</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>7</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>-,+,/,*</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1676326203</deployedseed>
        <deployedseed>745121917</deployedseed>
        <deployedseed>1294475609</deployedseed>
        <deployedseed>1573694115</deployedseed>
        <deployedseed>1380103228</deployedseed>
        <deployedseed>800138020</deployedseed>
        <deployedseed>1409956981</deployedseed>
        <deployedseed>466370301</deployedseed>
        <deployedseed>1770332193</deployedseed>
        <deployedseed>305677485</deployedseed>
        <deployedseed>1548431384</deployedseed>
        <deployedseed>1583463426</deployedseed>
        <deployedseed>1366338827</deployedseed>
        <deployedseed>367047438</deployedseed>
        <deployedseed>924277274</deployedseed>
        <deployedseed>226328307</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>999999</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>999999</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
