<quiz>
    <question type="stack">
        <name>
            <text>3.4.1.5 Gleichzeitig lösen [Solve simultaneous]</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Solve the following equations to find the values of \(x\) and \(y\)
    </span>
    <span class="multilang" lang="de">
        Lösen Sie die folgenden Gleichungen, um die Werte von \(x\) und \(y\) zu erhalten
    </span>
</p>
<ol class="HELM_parts">
    <li>
        <p>
            \[{@sysl1[1,1]@}={@sysr1[1,1]@}\]
        </p>
        <p>
            \[{@sysl1[2,1]@}={@sysr1[2,1]@}\]
        </p>
        <div>
            \({@B1@}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </div>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Solve the following equations to find the values of \(x\) and \(y\)
            </span>
            <span class="multilang" lang="de">
                Lösen Sie die folgenden Gleichungen, um die Werte von \(x\) und \(y\) zu erhalten
            </span>
        </p>
        <p>
            \[{@sysl2[1,1]@}={@sysr2[1,1]@}\]
        </p>
        <p>
            \[{@sysl2[2,1]@}={@sysr2[2,1]@}\]
        </p>
        <div>
            \({@B2@}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </div>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Solve the following equations to find the values of \(x\) and \(y\)
            </span>
            <span class="multilang" lang="de">
                Lösen Sie die folgenden Gleichungen, um die Werte von \(x\) und \(y\) zu erhalten
            </span>
        </p>
        <p>
            \[{@sysl3[1,1]@}={@sysr3[1,1]@}\]
        </p>
        <p>
            \[{@sysl3[2,1]@}={@sysr3[2,1]@}\]
        </p>
        <div>
            \({@B3@}=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </div>
    </li>
</ol>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<ol class="HELM_parts">
    <li>
        <p>
            \[{@sysl1[1,1]@}={@sysr1[1,1]@}\]
        </p>
        <p>
            \[{@sysl1[2,1]@}={@sysr1[2,1]@}\]
        </p>
        <p>
            <span class="multilang" lang="en">
                We modify the equations so that x can be eliminated. For example by multiplying the first equation by {@A1[2,1]@} and the second by {@A1[1,1]@} we find:
            </span>
            <span class="multilang" lang="de">
                Wir modifizieren die Gleichungen so, dass \(x\) eliminiert werden kann. Wenn wir zum Beispiel die erste Gleichung mit {@A1[2,1]@} und die zweite mit {@A1[1,1]@} multiplizieren, erhalten wir:
            </span>
        </p>
        <p>
            \[{@expand(A1[2,1]*sysl1[1,1])@}={@A1[2,1]*sysr1[1,1]@}\]
        </p>
        <p>
            \[{@expand(A1[1,1]*sysl1[2,1])@}={@A1[1,1]*sysr1[2,1]@}\]
        </p>
        <p>
            <span class="multilang" lang="en">
                If these equations are now added we can eliminate \(x.\) Therefore
            </span>
            <span class="multilang" lang="de">
                Wenn diese Gleichungen jetzt addiert werden, können wir \(x\) streichen. Daher
            </span>
        </p>
        <p>
            \({@B1@}= {@ta1@}\)
        </p>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Solve the following equations to find the values of \(x\) and \(y\)
            </span>
            <span class="multilang" lang="de">
                Lösen Sie die folgenden Gleichungen, um die Werte von \(x\) und \(y\) zu erhalten
            </span>
        </p>
        <p>
            \[{@sysl2[1,1]@}={@sysr2[1,1]@}\]
        </p>
        <p>
            \[{@sysl2[2,1]@}={@sysr2[2,1]@}\]
        </p>
        <div>
            \({@B2@}= {@ta2@}\)
        </div>
    </li>
    <li>
        <p>
            <span class="multilang" lang="en">
                Solve the following equations to find the values of \(x\) and \(y\)
            </span>
            <span class="multilang" lang="de">
                Lösen Sie die folgenden Gleichungen, um die Werte von \(x\) und \(y\) zu erhalten
            </span>
        </p>
        <p>
            \[{@sysl3[1,1]@}={@sysr3[1,1]@}\]
        </p>
        <p>
            \[{@sysl3[2,1]@}={@sysr3[2,1]@}\]
        </p>
        <div>
            \({@B3@}= {@ta3@}\)
        </div>
    </li>
</ol>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>ordergreat( x,y);

A1:rand(matrix([1,2],[3,4]))+matrix([2,-5],[2,2]);
B1:matrix([x],[y]);
sysl1:A1.B1;
ta1:rand(matrix([3],[3]))+matrix([-6],[1]);
sysr1:A1.ta1;

A2:rand(matrix([1,3],[5,1]))+matrix([3,2],[5,-3]);
B2:matrix([x],[y]);
sysl2:A2.B2;
ta2:rand(matrix([1],[3]))+matrix([-3],[1]);
sysr2:A2.ta2;

A3:rand(matrix([1,2],[3,4]))+matrix([2,2],[2,2]);
B3:matrix([x],[y]);
sysl3:A1.B1;
ta3:rand(matrix([1],[3]))+matrix([-25],[2]);
sysr3:A3.ta3;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote>
            <text>\[{@sysl1=sysr1@}, \quad {@B1=ta1@}\]
\[{@sysl2=sysr2@}, \quad {@B2=ta2@}\]
\[{@sysl3=sysr3@}, \quad {@B3=ta3@}\]</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>ta1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans2</name>
            <type>matrix</type>
            <tans>ta2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <input>
            <name>ans3</name>
            <type>matrix</type>
            <tans>ta3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>747401</deployedseed>
        <deployedseed>2092430348</deployedseed>
        <deployedseed>1618318875</deployedseed>
        <deployedseed>1400103312</deployedseed>
        <deployedseed>1937277017</deployedseed>
        <deployedseed>241226227</deployedseed>
        <deployedseed>808865466</deployedseed>
        <deployedseed>328291974</deployedseed>
        <deployedseed>1283423058</deployedseed>
        <deployedseed>1790492456</deployedseed>
        <deployedseed>2100343285</deployedseed>
        <deployedseed>410307393</deployedseed>
        <deployedseed>704834984</deployedseed>
        <deployedseed>502057471</deployedseed>
        <deployedseed>1882147453</deployedseed>
        <deployedseed>468304492</deployedseed>
        <deployedseed>1297222069</deployedseed>
        <deployedseed>1070156288</deployedseed>
        <deployedseed>1624345601</deployedseed>
        <deployedseed>1937177746</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
