<quiz>
    <question type="stack">
        <name>
            <text>6.2.4.4 Express exp in terms of coth (exercise 2(b))</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<p class="HELM_exercise">
    <span class="multilang" lang="en">
        Exercise
    </span>
    <span class="multilang" lang="de">
        Übung
    </span>
    <br/>
</p>
<p>
    <span class="multilang" lang="en">
        Express the following in terms of \(\coth({@aa@}{@xx@})\).
    </span>
    <span class="multilang" lang="de">
        Drücken Sie Folgendes in Form von \(\coth({@aa@}{@xx@})\) aus.
    </span>
</p>
<p>
    \(  \dfrac{{@uu@} e^{{@aa@}{@xx@}}} {{@vv@} e^{{@aa@}{@xx@}} - {@vv@} e^{-{@aa@}{@xx@}}} = \) [[input:ans1]] [[validation:ans1]]
</p>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<p>
    <span class="multilang" lang="en">
        It will be worth keeping in mind that \(\coth(z)=\dfrac{\cosh(z)}{\sinh(z)}\). We start by rewriting the expression in terms of \(\sinh({@aa@}{@xx@})\) and \(\cosh({@aa@}{@xx@})\) as at least these are more like the expression we are targeting. For the
                numerator we use \[e^z=\cosh(z)+\sinh(z)\] while for denominator we can notice that this is already quite close to \(\sinh({@aa@}{@xx@})\) indeed \[ \sinh({@aa@}{@xx@})=\frac{e^{{@aa@}{@xx@}}-e^{-{@aa@}{@xx@}}}{2}.\]
    </span>
    <span class="multilang" lang="de">
        Es lohnt sich, daran zu denken, dass \(\coth(z)=\dfrac{\cosh(z)}{\sinh(z)}\). Wir beginnen damit, den Ausdruck in Form von \(\sinh({@aa@}{@xx@})\) und \(\cosh({@aa@}{@xx@})\) umzuschreiben, da diese zumindest dem von uns angestrebten Ausdruck näher kommen. Für den
Zähler verwenden wir \[e^z=\cosh(z)+\sinh(z)\], während wir für den Nenner feststellen können, dass dies bereits ziemlich ähnlich ist wie \(\sinh({@aa@}{@xx@})\), nämlich \[ \sinh({@aa@}{@xx@})=\frac{e^{{@aa@}{@xx@}}-e^{-{@aa@}{@xx@}}}{2}.\]
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        Note that instead of spotting this we could just write each \(e^z=\cosh(z)+\sinh(z)\), there would be some cancelation after using the even / odd properties of \(\cosh(z)\) and \(\sinh(z)\). The result is of course the same. It's good exercise to confirm.
    </span>
    <span class="multilang" lang="de">
        Anstatt dies zu bemerken, könnten wir auch einfach \(e^z=\cosh(z)+\sinh(z)\) schreiben, da sich die geraden und ungeraden Eigenschaften von \(\cosh(z)\) und \(\sinh(z)\) aufheben würden. Das Ergebnis ist natürlich dasselbe. Es ist eine gute Übung, um das zu bestätigen.
    </span>
</p>
<p>
    <span class="multilang" lang="en">
        In any case, we can bring this together and simplify to see
    </span>
    <span class="multilang" lang="de">
        In jedem Fall können wir dies zusammenführen und vereinfachen
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        \[\dfrac{{@uu@} e^{{@aa@}{@xx@}}} {{@vv@} e^{{@aa@}{@xx@}} - {@vv@} e^{-{@aa@}{@xx@}}} = \dfrac{{@uu@}(\cosh({@aa@}{@xx@}) + \sinh({@aa@}{@xx@}))}{{@2*vv@} \sinh({@aa@}{@xx@})} = \dfrac{{@uu@}}{{@2*vv@}} \cdot \frac{\cosh({@aa@}{@xx@})}{\sinh({@aa@}{@xx@})} + \dfrac{{@uu@}}{{@2*vv@}}  = {@ta@}. \]
    </span>
    <br/>
</p>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text>aa: rand(6)+2;
uu: rand([3,5,7]);
vv: rand([2,4,8]);
xx: rand([x,y,z,t]);
exp: cosh(aa*xx) + sinh(aa*xx);
expn: cosh(aa*xx) - sinh(aa*xx);
nn: uu*exp(aa*xx); 
dd: vv*exp(aa*xx) - vv*expn(aa*xx);
ta: (uu/(2*vv))*coth(aa*xx) + uu/(2*vv);
ta1: factor(ta);
ta2: uu*(coth(aa*xx)/(2*vv) + 1/(2*vv));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text>\( \frac{{@uu@} e^{{@aa@}{@xx@}}} {{@2*vv@} e^{{@aa@}{@xx@}} - {@vv@} e^{-{@aa@}{@xx@}}} = {@ta@} \)</text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<span class="multilang" lang="en">
    Correct answer, well done.
</span>
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht.
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<span class="multilang" lang="en">
    Your answer is partially correct.
</span>
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise richtig.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<span class="multilang" lang="en">
    Incorrect answer.
</span>
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
]]></text>
        </prtincorrect>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>e,sinh,cosh,[[BASIC-ALGEBRA]]</forbidwords>
            <allowwords></allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <answertest>EqualComAss</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>coth(aa*xx)</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta2</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty></truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty></falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1661422740</deployedseed>
        <deployedseed>1369725300</deployedseed>
        <deployedseed>827831590</deployedseed>
        <deployedseed>1998631485</deployedseed>
        <deployedseed>161896063</deployedseed>
        <deployedseed>1045713347</deployedseed>
        <deployedseed>1249745547</deployedseed>
        <deployedseed>1661989766</deployedseed>
        <deployedseed>481623722</deployedseed>
        <deployedseed>1898850484</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <testinput>
                <name>ans1</name>
                <value>ev(ta+1, simp)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-3-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>4</testcase>
            <testinput>
                <name>ans1</name>
                <value>ta2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-3-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
</quiz>
